{
 "cells": [
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "---\n",
    "date: 2024-04-29\n",
    "title: Experiments in the Irrationality of Sqrt 2 with SMT\n",
    "---\n",
    "\n",
    "A classical theorem is that $\\sqrt{2}$ is irrational. https://en.wikipedia.org/wiki/Square_root_of_2#Proofs_of_irrationality This means it cannot be written as $\\frac{p}{q}$ for integers $p$ and $q$.\n",
    "\n",
    "This and the existence of an infinite number of primes are some simple prominent theorems and reasonable test cases for a system intended to support proofs.\n",
    "\n",
    "There is an interesting article [The Seventeen Provers of the World](https://www.cs.ru.nl/~freek/comparison/comparison.pdf) compiled by Freek Wiedijk which does the irrationality of two in many different systems.\n",
    "\n",
    "I've been interested in trying to build infrastructure for a highly automated interactive proof system using off the shelf solvers and this is an test case.\n",
    "\n",
    "First off, just asking the question in a naive-ish way chokes the solvers. It is not inconceivable that they could have baked this in."
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 1,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "failed to prove\n"
     ]
    },
    {
     "ename": "Z3Exception",
     "evalue": "model is not available",
     "output_type": "error",
     "traceback": [
      "\u001b[0;31m---------------------------------------------------------------------------\u001b[0m",
      "\u001b[0;31mZ3Exception\u001b[0m                               Traceback (most recent call last)",
      "File \u001b[0;32m~/.local/lib/python3.10/site-packages/z3/z3.py:7130\u001b[0m, in \u001b[0;36mSolver.model\u001b[0;34m(self)\u001b[0m\n\u001b[1;32m   7129\u001b[0m \u001b[38;5;28;01mtry\u001b[39;00m:\n\u001b[0;32m-> 7130\u001b[0m     \u001b[38;5;28;01mreturn\u001b[39;00m ModelRef(\u001b[43mZ3_solver_get_model\u001b[49m\u001b[43m(\u001b[49m\u001b[38;5;28;43mself\u001b[39;49m\u001b[38;5;241;43m.\u001b[39;49m\u001b[43mctx\u001b[49m\u001b[38;5;241;43m.\u001b[39;49m\u001b[43mref\u001b[49m\u001b[43m(\u001b[49m\u001b[43m)\u001b[49m\u001b[43m,\u001b[49m\u001b[43m \u001b[49m\u001b[38;5;28;43mself\u001b[39;49m\u001b[38;5;241;43m.\u001b[39;49m\u001b[43msolver\u001b[49m\u001b[43m)\u001b[49m, \u001b[38;5;28mself\u001b[39m\u001b[38;5;241m.\u001b[39mctx)\n\u001b[1;32m   7131\u001b[0m \u001b[38;5;28;01mexcept\u001b[39;00m Z3Exception:\n",
      "File \u001b[0;32m~/.local/lib/python3.10/site-packages/z3/z3core.py:4065\u001b[0m, in \u001b[0;36mZ3_solver_get_model\u001b[0;34m(a0, a1, _elems)\u001b[0m\n\u001b[1;32m   4064\u001b[0m r \u001b[38;5;241m=\u001b[39m _elems\u001b[38;5;241m.\u001b[39mf(a0, a1)\n\u001b[0;32m-> 4065\u001b[0m \u001b[43m_elems\u001b[49m\u001b[38;5;241;43m.\u001b[39;49m\u001b[43mCheck\u001b[49m\u001b[43m(\u001b[49m\u001b[43ma0\u001b[49m\u001b[43m)\u001b[49m\n\u001b[1;32m   4066\u001b[0m \u001b[38;5;28;01mreturn\u001b[39;00m r\n",
      "File \u001b[0;32m~/.local/lib/python3.10/site-packages/z3/z3core.py:1475\u001b[0m, in \u001b[0;36mElementaries.Check\u001b[0;34m(self, ctx)\u001b[0m\n\u001b[1;32m   1474\u001b[0m \u001b[38;5;28;01mif\u001b[39;00m err \u001b[38;5;241m!=\u001b[39m \u001b[38;5;28mself\u001b[39m\u001b[38;5;241m.\u001b[39mOK:\n\u001b[0;32m-> 1475\u001b[0m     \u001b[38;5;28;01mraise\u001b[39;00m \u001b[38;5;28mself\u001b[39m\u001b[38;5;241m.\u001b[39mException(\u001b[38;5;28mself\u001b[39m\u001b[38;5;241m.\u001b[39mget_error_message(ctx, err))\n",
      "\u001b[0;31mZ3Exception\u001b[0m: b'there is no current model'",
      "\nDuring handling of the above exception, another exception occurred:\n",
      "\u001b[0;31mZ3Exception\u001b[0m                               Traceback (most recent call last)",
      "\u001b[0;32m/tmp/ipykernel_367156/575137406.py\u001b[0m in \u001b[0;36m?\u001b[0;34m()\u001b[0m\n\u001b[1;32m      1\u001b[0m \u001b[0;32mfrom\u001b[0m \u001b[0mz3\u001b[0m \u001b[0;32mimport\u001b[0m \u001b[0;34m*\u001b[0m\u001b[0;34m\u001b[0m\u001b[0;34m\u001b[0m\u001b[0m\n\u001b[1;32m      2\u001b[0m \u001b[0;31m# from cvc5.pythonic import * # cvc5 is a pretty good dropin. It also fails here.\u001b[0m\u001b[0;34m\u001b[0m\u001b[0;34m\u001b[0m\u001b[0m\n\u001b[1;32m      3\u001b[0m \u001b[0mn\u001b[0m\u001b[0;34m,\u001b[0m\u001b[0mm\u001b[0m \u001b[0;34m=\u001b[0m \u001b[0mInts\u001b[0m\u001b[0;34m(\u001b[0m\u001b[0;34m\"n m\"\u001b[0m\u001b[0;34m)\u001b[0m\u001b[0;34m\u001b[0m\u001b[0;34m\u001b[0m\u001b[0m\n\u001b[0;32m----> 4\u001b[0;31m \u001b[0mprove\u001b[0m\u001b[0;34m(\u001b[0m\u001b[0mNot\u001b[0m\u001b[0;34m(\u001b[0m\u001b[0mExists\u001b[0m\u001b[0;34m(\u001b[0m\u001b[0;34m[\u001b[0m\u001b[0mn\u001b[0m\u001b[0;34m,\u001b[0m\u001b[0mm\u001b[0m\u001b[0;34m]\u001b[0m\u001b[0;34m,\u001b[0m \u001b[0mAnd\u001b[0m\u001b[0;34m(\u001b[0m\u001b[0mm\u001b[0m \u001b[0;34m!=\u001b[0m \u001b[0;36m0\u001b[0m\u001b[0;34m,\u001b[0m \u001b[0mn\u001b[0m\u001b[0;34m*\u001b[0m\u001b[0mn\u001b[0m \u001b[0;34m==\u001b[0m \u001b[0;36m2\u001b[0m \u001b[0;34m*\u001b[0m \u001b[0mm\u001b[0m \u001b[0;34m*\u001b[0m \u001b[0mm\u001b[0m\u001b[0;34m)\u001b[0m\u001b[0;34m)\u001b[0m\u001b[0;34m)\u001b[0m\u001b[0;34m)\u001b[0m\u001b[0;34m\u001b[0m\u001b[0;34m\u001b[0m\u001b[0m\n\u001b[0m",
      "\u001b[0;32m~/.local/lib/python3.10/site-packages/z3/z3.py\u001b[0m in \u001b[0;36m?\u001b[0;34m(claim, show, **keywords)\u001b[0m\n\u001b[1;32m   9087\u001b[0m     \u001b[0;32mif\u001b[0m \u001b[0mr\u001b[0m \u001b[0;34m==\u001b[0m \u001b[0munsat\u001b[0m\u001b[0;34m:\u001b[0m\u001b[0;34m\u001b[0m\u001b[0;34m\u001b[0m\u001b[0m\n\u001b[1;32m   9088\u001b[0m         \u001b[0mprint\u001b[0m\u001b[0;34m(\u001b[0m\u001b[0;34m\"proved\"\u001b[0m\u001b[0;34m)\u001b[0m\u001b[0;34m\u001b[0m\u001b[0;34m\u001b[0m\u001b[0m\n\u001b[1;32m   9089\u001b[0m     \u001b[0;32melif\u001b[0m \u001b[0mr\u001b[0m \u001b[0;34m==\u001b[0m \u001b[0munknown\u001b[0m\u001b[0;34m:\u001b[0m\u001b[0;34m\u001b[0m\u001b[0;34m\u001b[0m\u001b[0m\n\u001b[1;32m   9090\u001b[0m         \u001b[0mprint\u001b[0m\u001b[0;34m(\u001b[0m\u001b[0;34m\"failed to prove\"\u001b[0m\u001b[0;34m)\u001b[0m\u001b[0;34m\u001b[0m\u001b[0;34m\u001b[0m\u001b[0m\n\u001b[0;32m-> 9091\u001b[0;31m         \u001b[0mprint\u001b[0m\u001b[0;34m(\u001b[0m\u001b[0ms\u001b[0m\u001b[0;34m.\u001b[0m\u001b[0mmodel\u001b[0m\u001b[0;34m(\u001b[0m\u001b[0;34m)\u001b[0m\u001b[0;34m)\u001b[0m\u001b[0;34m\u001b[0m\u001b[0;34m\u001b[0m\u001b[0m\n\u001b[0m\u001b[1;32m   9092\u001b[0m     \u001b[0;32melse\u001b[0m\u001b[0;34m:\u001b[0m\u001b[0;34m\u001b[0m\u001b[0;34m\u001b[0m\u001b[0m\n\u001b[1;32m   9093\u001b[0m         \u001b[0mprint\u001b[0m\u001b[0;34m(\u001b[0m\u001b[0;34m\"counterexample\"\u001b[0m\u001b[0;34m)\u001b[0m\u001b[0;34m\u001b[0m\u001b[0;34m\u001b[0m\u001b[0m\n\u001b[1;32m   9094\u001b[0m         \u001b[0mprint\u001b[0m\u001b[0;34m(\u001b[0m\u001b[0ms\u001b[0m\u001b[0;34m.\u001b[0m\u001b[0mmodel\u001b[0m\u001b[0;34m(\u001b[0m\u001b[0;34m)\u001b[0m\u001b[0;34m)\u001b[0m\u001b[0;34m\u001b[0m\u001b[0;34m\u001b[0m\u001b[0m\n",
      "\u001b[0;32m~/.local/lib/python3.10/site-packages/z3/z3.py\u001b[0m in \u001b[0;36m?\u001b[0;34m(self)\u001b[0m\n\u001b[1;32m   7128\u001b[0m         \"\"\"\n\u001b[1;32m   7129\u001b[0m         \u001b[0;32mtry\u001b[0m\u001b[0;34m:\u001b[0m\u001b[0;34m\u001b[0m\u001b[0;34m\u001b[0m\u001b[0m\n\u001b[1;32m   7130\u001b[0m             \u001b[0;32mreturn\u001b[0m \u001b[0mModelRef\u001b[0m\u001b[0;34m(\u001b[0m\u001b[0mZ3_solver_get_model\u001b[0m\u001b[0;34m(\u001b[0m\u001b[0mself\u001b[0m\u001b[0;34m.\u001b[0m\u001b[0mctx\u001b[0m\u001b[0;34m.\u001b[0m\u001b[0mref\u001b[0m\u001b[0;34m(\u001b[0m\u001b[0;34m)\u001b[0m\u001b[0;34m,\u001b[0m \u001b[0mself\u001b[0m\u001b[0;34m.\u001b[0m\u001b[0msolver\u001b[0m\u001b[0;34m)\u001b[0m\u001b[0;34m,\u001b[0m \u001b[0mself\u001b[0m\u001b[0;34m.\u001b[0m\u001b[0mctx\u001b[0m\u001b[0;34m)\u001b[0m\u001b[0;34m\u001b[0m\u001b[0;34m\u001b[0m\u001b[0m\n\u001b[1;32m   7131\u001b[0m         \u001b[0;32mexcept\u001b[0m \u001b[0mZ3Exception\u001b[0m\u001b[0;34m:\u001b[0m\u001b[0;34m\u001b[0m\u001b[0;34m\u001b[0m\u001b[0m\n\u001b[0;32m-> 7132\u001b[0;31m             \u001b[0;32mraise\u001b[0m \u001b[0mZ3Exception\u001b[0m\u001b[0;34m(\u001b[0m\u001b[0;34m\"model is not available\"\u001b[0m\u001b[0;34m)\u001b[0m\u001b[0;34m\u001b[0m\u001b[0;34m\u001b[0m\u001b[0m\n\u001b[0m",
      "\u001b[0;31mZ3Exception\u001b[0m: model is not available"
     ]
    }
   ],
   "source": [
    "from z3 import *\n",
    "# from cvc5.pythonic import * # cvc5 is a pretty good dropin. It also fails here.\n",
    "n,m = Ints(\"n m\")\n",
    "prove(Not(Exists([n,m], And(m != 0, n*n == 2 * m * m))))"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "Ok, well we may have to help it along.\n",
    "\n",
    "I took a look in Baby Rudin to see what their proof looks like. This led me down a rabbithole that was seemingly leading no where. \n",
    "\n",
    "I started proving lemmas about even and odd numbers. There were some basic facts that seemed difficult to push through.\n"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": null,
   "metadata": {},
   "outputs": [],
   "source": [
    "even = Function(\"even\", IntSort(), BoolSort())\n",
    "even_def = ForAll(x, even(x) == Exists([y], x == 2*y))\n",
    "odd = Function(\"odd\", IntSort(), BoolSort())\n",
    "odd_def = ForAll(x, odd(x) == Exists([y], x == 2*y+1))\n",
    "\n",
    "even_or_odd = ForAll(x, even(x) != odd(x))\n",
    "# Z3 can't prove this for some reason. I suppose this might be a complex diophantine equation or something.\n",
    "# CVC5 seemingly can prove this.\n",
    "prove(Implies(And(even_def, odd_def), even_or_odd))\n",
    "\n",
    "div2_even = ForAll([x], even(x) == (x == 2 * (x/2)))\n",
    "prove(Implies(even_def, div2_even))\n",
    "\n",
    "div2_odd = ForAll([x], odd(x) == (x == 2 * (x/2) + 1))\n",
    "prove(Implies(odd_def, div2_odd))"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "So instead I took a look at the Otter proof in the above paper. Otter is an automated theorem and the predecessor of [Prover9](https://www.cs.unm.edu/~mccune/prover9/manual/2009-11A/). It offers some features with more operational flavor that the modern contingent of solvers (vampire and eprover) don't.\n",
    "\n",
    "This encoding of the problem seems pretty clever. `d(x,y)` is the y is divisible by x. `m` is multiplication. I think implicitly it is working over positive integers to avoid `!= 0` constraints.\n",
    "\n",
    "Note that in CNF, `~ foo | bar` should often be read as `foo => bar` which is classically equivalent. In that case, this reads as some equations and horn clauses / conditional equations."
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 2,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "Overwriting /tmp/sqrt2.in\n"
     ]
    }
   ],
   "source": [
    "%%file /tmp/sqrt2.in\n",
    "% From the Wiedijk review. Wos, Beeson, McCune\n",
    "set(auto).\n",
    "%set(ur_res).\n",
    "%assign(max_distinct_vars, 1).\n",
    "formulas(sos).\n",
    "    x = x.  %? Why did they write this?\n",
    "    m(1,x) = x. % identity\n",
    "    m(x,1) = x.\n",
    "    m(x,m(y,z)) = m(m(x,y),z). % associativity\n",
    "    m(x,y) = m(y,x). % commutativity\n",
    "    m(x,y) != m(x,z) | y = z. % cancellation... ??? What about x = 0? Implicit positive integers only?\n",
    "    -d(x,y) | m(x,f(x,y)) = y. % this and next line define divides\n",
    "    m(x,z) != y | d(x,y).\n",
    "    -d(2,m(x,y)) | d(2,x) | d(2,y). % 2 is prime\n",
    "    m(a,a) = m(2,m(b,b)). % a/b = âˆš2\n",
    "    -d(x,a) | -d(x,b) | x = 1. % a/b is in lowest terms\n",
    "    2 != 1. % I almost forgot this!\n",
    "end_of_list."
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 4,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "============================== Prover9 ===============================\n",
      "Prover9 (64) version 2017-11A (CIIRC), November 2017.\n",
      "Process 370431 was started by philip on philip-ThinkPad-Z13-Gen-2,\n",
      "Mon Apr 29 20:47:34 2024\n",
      "The command was \"prover9 -f /tmp/sqrt2.in\".\n",
      "============================== end of head ===========================\n",
      "\n",
      "============================== INPUT =================================\n",
      "\n",
      "% Reading from file /tmp/sqrt2.in\n",
      "\n",
      "set(auto).\n",
      "    % set(auto) -> set(auto_inference).\n",
      "    % set(auto) -> set(auto_setup).\n",
      "    % set(auto_setup) -> set(predicate_elim).\n",
      "    % set(auto_setup) -> assign(eq_defs, unfold).\n",
      "    % set(auto) -> set(auto_limits).\n",
      "    % set(auto_limits) -> assign(max_weight, \"100.000\").\n",
      "    % set(auto_limits) -> assign(sos_limit, 20000).\n",
      "    % set(auto) -> set(auto_denials).\n",
      "    % set(auto) -> set(auto_process).\n",
      "\n",
      "formulas(sos).\n",
      "x = x.\n",
      "m(1,x) = x.\n",
      "m(x,1) = x.\n",
      "m(x,m(y,z)) = m(m(x,y),z).\n",
      "m(x,y) = m(y,x).\n",
      "m(x,y) != m(x,z) | y = z.\n",
      "-d(x,y) | m(x,f(x,y)) = y.\n",
      "m(x,z) != y | d(x,y).\n",
      "-d(2,m(x,y)) | d(2,x) | d(2,y).\n",
      "m(a,a) = m(2,m(b,b)).\n",
      "-d(x,a) | -d(x,b) | x = 1.\n",
      "2 != 1.\n",
      "end_of_list.\n",
      "\n",
      "============================== end of input ==========================\n",
      "\n",
      "============================== PROCESS NON-CLAUSAL FORMULAS ==========\n",
      "\n",
      "% Formulas that are not ordinary clauses:\n",
      "\n",
      "============================== end of process non-clausal formulas ===\n",
      "\n",
      "============================== PROCESS INITIAL CLAUSES ===============\n",
      "\n",
      "% Clauses before input processing:\n",
      "\n",
      "formulas(usable).\n",
      "end_of_list.\n",
      "\n",
      "formulas(sos).\n",
      "x = x.  [assumption].\n",
      "m(1,x) = x.  [assumption].\n",
      "m(x,1) = x.  [assumption].\n",
      "m(x,m(y,z)) = m(m(x,y),z).  [assumption].\n",
      "m(x,y) = m(y,x).  [assumption].\n",
      "m(x,y) != m(x,z) | y = z.  [assumption].\n",
      "-d(x,y) | m(x,f(x,y)) = y.  [assumption].\n",
      "m(x,y) != z | d(x,z).  [assumption].\n",
      "-d(2,m(x,y)) | d(2,x) | d(2,y).  [assumption].\n",
      "m(a,a) = m(2,m(b,b)).  [assumption].\n",
      "-d(x,a) | -d(x,b) | x = 1.  [assumption].\n",
      "2 != 1.  [assumption].\n",
      "end_of_list.\n",
      "\n",
      "formulas(demodulators).\n",
      "end_of_list.\n",
      "\n",
      "============================== PREDICATE ELIMINATION =================\n",
      "\n",
      "No predicates eliminated.\n",
      "\n",
      "============================== end predicate elimination =============\n",
      "\n",
      "Auto_denials:  (non-Horn, no changes).\n",
      "\n",
      "Term ordering decisions:\n",
      "Predicate symbol precedence:  predicate_order([ =, d ]).\n",
      "Function symbol precedence:  function_order([ 2, 1, a, b, m, f ]).\n",
      "After inverse_order:  (no changes).\n",
      "Unfolding symbols: (none).\n",
      "\n",
      "Auto_inference settings:\n",
      "  % set(paramodulation).  % (positive equality literals)\n",
      "  % set(binary_resolution).  % (non-Horn)\n",
      "  % set(neg_ur_resolution).  % (non-Horn, less than 100 clauses)\n",
      "\n",
      "Auto_process settings:\n",
      "  % set(factor).  % (non-Horn)\n",
      "  % set(unit_deletion).  % (non-Horn)\n",
      "\n",
      "kept:      1 m(1,x) = x.  [assumption].\n",
      "kept:      2 m(x,1) = x.  [assumption].\n",
      "           3 m(x,m(y,z)) = m(m(x,y),z).  [assumption].\n",
      "kept:      4 m(m(x,y),z) = m(x,m(y,z)).  [copy(3),flip(a)].\n",
      "\n",
      "% Operation m is commutative; C redundancy checks enabled.\n",
      "kept:      5 m(x,y) = m(y,x).  [assumption].\n",
      "kept:      6 m(x,y) != m(x,z) | y = z.  [assumption].\n",
      "kept:      7 -d(x,y) | m(x,f(x,y)) = y.  [assumption].\n",
      "kept:      8 m(x,y) != z | d(x,z).  [assumption].\n",
      "kept:      9 -d(2,m(x,y)) | d(2,x) | d(2,y).  [assumption].\n",
      "           10 m(a,a) = m(2,m(b,b)).  [assumption].\n",
      "kept:      11 m(2,m(b,b)) = m(a,a).  [copy(10),flip(a)].\n",
      "           12 -d(x,a) | -d(x,b) | x = 1.  [assumption].\n",
      "kept:      13 -d(x,a) | -d(x,b) | 1 = x.  [copy(12),flip(c)].\n",
      "           14 2 != 1.  [assumption].\n",
      "kept:      15 1 != 2.  [copy(14),flip(a)].\n",
      "kept:      16 -d(2,m(x,x)) | d(2,x).  [factor(9,b,c)].\n",
      "\n",
      "============================== end of process initial clauses ========\n",
      "\n",
      "============================== CLAUSES FOR SEARCH ====================\n",
      "\n",
      "% Clauses after input processing:\n",
      "\n",
      "formulas(usable).\n",
      "end_of_list.\n",
      "\n",
      "formulas(sos).\n",
      "1 m(1,x) = x.  [assumption].\n",
      "2 m(x,1) = x.  [assumption].\n",
      "4 m(m(x,y),z) = m(x,m(y,z)).  [copy(3),flip(a)].\n",
      "5 m(x,y) = m(y,x).  [assumption].\n",
      "6 m(x,y) != m(x,z) | y = z.  [assumption].\n",
      "7 -d(x,y) | m(x,f(x,y)) = y.  [assumption].\n",
      "8 m(x,y) != z | d(x,z).  [assumption].\n",
      "9 -d(2,m(x,y)) | d(2,x) | d(2,y).  [assumption].\n",
      "11 m(2,m(b,b)) = m(a,a).  [copy(10),flip(a)].\n",
      "13 -d(x,a) | -d(x,b) | 1 = x.  [copy(12),flip(c)].\n",
      "15 1 != 2.  [copy(14),flip(a)].\n",
      "16 -d(2,m(x,x)) | d(2,x).  [factor(9,b,c)].\n",
      "end_of_list.\n",
      "\n",
      "formulas(demodulators).\n",
      "1 m(1,x) = x.  [assumption].\n",
      "2 m(x,1) = x.  [assumption].\n",
      "4 m(m(x,y),z) = m(x,m(y,z)).  [copy(3),flip(a)].\n",
      "5 m(x,y) = m(y,x).  [assumption].\n",
      "        % (lex-dep)\n",
      "11 m(2,m(b,b)) = m(a,a).  [copy(10),flip(a)].\n",
      "end_of_list.\n",
      "\n",
      "============================== end of clauses for search =============\n",
      "\n",
      "============================== SEARCH ================================\n",
      "\n",
      "% Starting search at 0.00 seconds.\n",
      "\n",
      "given #1 (I,wt=5): 1 m(1,x) = x.  [assumption].\n",
      "\n",
      "given #2 (I,wt=5): 2 m(x,1) = x.  [assumption].\n",
      "\n",
      "given #3 (I,wt=11): 4 m(m(x,y),z) = m(x,m(y,z)).  [copy(3),flip(a)].\n",
      "\n",
      "given #4 (I,wt=7): 5 m(x,y) = m(y,x).  [assumption].\n",
      "\n",
      "% Operation m is associative-commutative; CAC redundancy checks enabled.\n",
      "\n",
      "given #5 (I,wt=10): 6 m(x,y) != m(x,z) | y = z.  [assumption].\n",
      "\n",
      "given #6 (I,wt=10): 7 -d(x,y) | m(x,f(x,y)) = y.  [assumption].\n",
      "\n",
      "given #7 (I,wt=8): 8 m(x,y) != z | d(x,z).  [assumption].\n",
      "\n",
      "given #8 (I,wt=11): 9 -d(2,m(x,y)) | d(2,x) | d(2,y).  [assumption].\n",
      "\n",
      "given #9 (I,wt=9): 11 m(2,m(b,b)) = m(a,a).  [copy(10),flip(a)].\n",
      "\n",
      "given #10 (I,wt=9): 13 -d(x,a) | -d(x,b) | 1 = x.  [copy(12),flip(c)].\n",
      "\n",
      "given #11 (I,wt=3): 15 1 != 2.  [copy(14),flip(a)].\n",
      "\n",
      "given #12 (I,wt=8): 16 -d(2,m(x,x)) | d(2,x).  [factor(9,b,c)].\n",
      "\n",
      "given #13 (A,wt=11): 17 m(x,m(y,z)) = m(y,m(x,z)).  [para(5(a,1),4(a,1,1)),rewrite([4(2)])].\n",
      "\n",
      "given #14 (F,wt=5): 38 m(x,2) != x.  [ur(6,b,15,a),rewrite([2(2)]),flip(a)].\n",
      "\n",
      "given #15 (F,wt=5): 45 m(2,x) != x.  [para(5(a,1),38(a,1))].\n",
      "\n",
      "given #16 (F,wt=7): 47 m(b,b) != m(a,a).  [para(11(a,1),45(a,1)),flip(a)].\n",
      "\n",
      "given #17 (F,wt=9): 44 m(x,m(y,2)) != m(x,y).  [ur(6,b,38,a)].\n",
      "\n",
      "given #18 (T,wt=3): 24 d(x,x).  [resolve(8,a,2,a)].\n",
      "\n",
      "given #19 (T,wt=3): 25 d(1,x).  [resolve(8,a,1,a)].\n",
      "\n",
      "given #20 (T,wt=5): 22 d(x,m(y,x)).  [resolve(8,a,5,a)].\n",
      "\n",
      "given #21 (T,wt=5): 27 d(x,m(x,y)).  [resolve(8,a,2,a(flip)),rewrite([5(3),1(3)])].\n",
      "\n",
      "given #22 (A,wt=8): 18 m(x,y) != x | 1 = y.  [para(2(a,1),6(a,1)),flip(a)].\n",
      "\n",
      "given #23 (F,wt=9): 46 m(x,m(2,y)) != m(x,y).  [ur(6,b,45,a)].\n",
      "\n",
      "given #24 (F,wt=9): 50 m(x,m(2,y)) != m(y,x).  [para(5(a,1),44(a,1)),rewrite([4(3)])].\n",
      "\n",
      "given #25 (F,wt=9): 51 m(x,m(y,2)) != m(y,x).  [para(5(a,1),44(a,2))].\n",
      "\n",
      "given #26 (F,wt=9): 66 m(2,m(x,y)) != m(y,x).  [para(5(a,1),46(a,1)),rewrite([4(3)])].\n",
      "\n",
      "given #27 (T,wt=5): 33 d(2,m(a,a)).  [resolve(11,a,8,a)].\n",
      "\n",
      "given #28 (T,wt=3): 79 d(2,a).  [resolve(33,a,16,a)].\n",
      "\n",
      "given #29 (T,wt=5): 55 f(1,x) = x.  [resolve(25,a,7,a),rewrite([1(4)])].\n",
      "\n",
      "given #30 (T,wt=6): 28 x != y | d(x,y).  [para(2(a,1),8(a,1))].\n",
      "\n",
      "given #31 (A,wt=14): 19 m(x,m(y,z)) != m(x,m(y,u)) | z = u.  [para(4(a,1),6(a,1)),rewrite([4(4)])].\n",
      "\n",
      "given #32 (F,wt=3): 82 -d(2,b).  [ur(13,a,79,a,c,15,a)].\n",
      "\n",
      "given #33 (F,wt=3): 100 b != 2.  [ur(28,b,82,a),flip(a)].\n",
      "\n",
      "given #34 (F,wt=5): 101 -d(2,m(b,b)).  [ur(16,b,82,a)].\n",
      "\n",
      "given #35 (F,wt=5): 102 m(2,x) != b.  [ur(8,b,82,a)].\n",
      "\n",
      "given #36 (T,wt=6): 53 -d(b,a) | b = 1.  [resolve(24,a,13,b),flip(b)].\n",
      "\n",
      "given #37 (T,wt=7): 39 d(x,m(y,m(x,z))).  [resolve(17,a,8,a)].\n",
      "\n",
      "given #38 (T,wt=5): 118 d(b,m(a,a)).  [para(11(a,1),39(a,2))].\n",
      "\n",
      "given #39 (T,wt=7): 54 m(x,f(x,x)) = x.  [resolve(24,a,7,a)].\n",
      "\n",
      "given #40 (A,wt=10): 20 m(x,y) != m(z,x) | z = y.  [para(5(a,1),6(a,1)),flip(a)].\n",
      "\n",
      "given #41 (F,wt=5): 105 m(b,b) != 2.  [ur(28,b,101,a),flip(a)].\n",
      "\n",
      "given #42 (F,wt=5): 111 m(x,2) != b.  [para(5(a,1),102(a,1))].\n",
      "\n",
      "given #43 (F,wt=5): 112 m(a,a) != b.  [para(11(a,1),102(a,1))].\n",
      "\n",
      "given #44 (F,wt=7): 104 m(x,b) != m(x,2).  [ur(6,b,100,a)].\n",
      "\n",
      "given #45 (T,wt=5): 120 f(x,x) = 1.  [resolve(54,a,18,a),flip(a)].\n",
      "\n",
      "given #46 (T,wt=6): 131 x != y | y = x.  [para(1(a,1),20(a,1)),rewrite([2(2)])].\n",
      "\n",
      "given #47 (T,wt=7): 57 d(x,m(y,m(z,x))).  [para(4(a,1),22(a,2))].\n",
      "\n",
      "given #48 (T,wt=7): 58 d(m(b,b),m(a,a)).  [para(11(a,1),22(a,2))].\n",
      "\n",
      "given #49 (A,wt=10): 21 m(x,y) != m(z,x) | y = z.  [para(5(a,1),6(a,2))].\n",
      "\n",
      "given #50 (F,wt=7): 107 -d(2,m(b,m(b,b))).  [ur(9,b,82,a,c,101,a)].\n",
      "\n",
      "given #51 (F,wt=7): 108 m(b,b) != m(2,x).  [ur(8,b,101,a),flip(a)].\n",
      "\n",
      "given #52 (F,wt=7): 113 m(x,m(2,y)) != b.  [para(17(a,1),102(a,1))].\n",
      "\n",
      "given #53 (F,wt=7): 122 m(b,x) != m(x,2).  [ur(20,b,100,a),flip(a)].\n",
      "\n",
      "given #54 (T,wt=7): 81 m(2,f(2,a)) = a.  [resolve(79,a,7,a)].\n",
      "\n",
      "given #55 (T,wt=5): 191 d(f(2,a),a).  [para(81(a,1),22(a,2))].\n",
      "\n",
      "given #56 (T,wt=5): 197 d(2,m(x,a)).  [para(81(a,1),39(a,2,2))].\n",
      "\n",
      "given #57 (T,wt=5): 206 d(2,m(a,x)).  [para(5(a,1),197(a,2))].\n",
      "\n",
      "given #58 (A,wt=9): 23 d(m(x,y),m(x,m(y,z))).  [resolve(8,a,4,a)].\n",
      "\n",
      "given #59 (F,wt=3): 196 b != a.  [para(81(a,1),102(a,1)),flip(a)].\n",
      "\n",
      "given #60 (F,wt=5): 190 f(2,a) != a.  [para(81(a,1),45(a,1)),flip(a)].\n",
      "\n",
      "given #61 (F,wt=5): 201 m(b,b) != a.  [para(81(a,1),108(a,2))].\n",
      "\n",
      "given #62 (F,wt=5): 202 m(x,a) != b.  [para(81(a,1),113(a,1,2))].\n",
      "\n",
      "given #63 (T,wt=6): 188 a != x | d(2,x).  [para(81(a,1),8(a,1))].\n",
      "\n",
      "given #64 (T,wt=7): 84 d(m(x,y),m(y,x)).  [resolve(28,a,5,a)].\n",
      "\n",
      "given #65 (T,wt=7): 200 d(f(2,a),m(x,a)).  [para(81(a,1),57(a,2,2))].\n",
      "\n",
      "given #66 (T,wt=7): 205 d(2,m(x,m(y,a))).  [para(4(a,1),197(a,2))].\n",
      "\n",
      "given #67 (A,wt=12): 29 m(x,m(y,z)) != u | d(m(x,y),u).  [para(4(a,1),8(a,1))].\n",
      "\n",
      "given #68 (F,wt=5): 233 m(a,x) != b.  [para(5(a,1),202(a,1))].\n",
      "\n",
      "given #69 (F,wt=3): 255 -d(a,b).  [ur(7,b,233,a)].\n",
      "\n",
      "given #70 (F,wt=7): 128 m(2,x) != m(x,b).  [ur(20,b,100,a(flip)),flip(a)].\n",
      "\n",
      "given #71 (F,wt=7): 147 m(x,m(y,2)) != b.  [para(4(a,1),111(a,1))].\n",
      "\n",
      "given #72 (T,wt=7): 208 d(2,m(x,m(a,y))).  [para(17(a,1),206(a,2))].\n",
      "\n",
      "given #73 (T,wt=7): 214 d(m(2,b),m(a,a)).  [para(11(a,1),23(a,2))].\n",
      "\n",
      "given #74 (T,wt=7): 217 d(m(x,2),m(x,a)).  [para(81(a,1),23(a,2,2))].\n",
      "\n",
      "given #75 (T,wt=7): 241 d(f(2,a),m(a,x)).  [para(5(a,1),200(a,2))].\n",
      "\n",
      "given #76 (A,wt=8): 30 m(x,y) != z | d(y,z).  [para(5(a,1),8(a,1))].\n",
      "\n",
      "given #77 (F,wt=7): 164 m(b,m(b,b)) != 2.  [ur(28,b,107,a),flip(a)].\n",
      "\n",
      "given #78 (F,wt=7): 172 m(b,b) != m(x,2).  [para(5(a,1),108(a,2))].\n",
      "\n",
      "given #79 (F,wt=7): 184 m(b,x) != m(2,x).  [para(5(a,1),122(a,2))].\n",
      "\n",
      "given #80 (F,wt=7): 218 m(a,x) != m(x,b).  [ur(21,b,196,a),flip(a)].\n",
      "\n",
      "given #81 (T,wt=7): 275 d(m(2,x),m(x,a)).  [para(5(a,1),217(a,1))].\n",
      "\n",
      "given #82 (T,wt=7): 276 d(m(x,2),m(a,x)).  [para(5(a,1),217(a,2))].\n",
      "\n",
      "given #83 (T,wt=7): 304 d(m(2,x),m(a,x)).  [para(5(a,1),275(a,2))].\n",
      "\n",
      "given #84 (T,wt=8): 37 m(a,a) != x | d(2,x).  [para(11(a,1),8(a,1))].\n",
      "\n",
      "given #85 (A,wt=15): 31 -d(2,m(x,m(y,z))) | d(2,m(x,y)) | d(2,z).  [para(4(a,1),9(a,2))].\n",
      "\n",
      "given #86 (F,wt=7): 219 m(b,x) != m(x,a).  [ur(20,b,196,a),flip(a)].\n",
      "\n",
      "given #87 (F,wt=7): 221 m(x,b) != m(x,a).  [ur(6,b,196,a)].\n",
      "\n",
      "given #88 (F,wt=7): 232 m(x,m(y,a)) != b.  [para(4(a,1),202(a,1))].\n",
      "\n",
      "given #89 (F,wt=7): 234 m(b,m(b,b)) != a.  [ur(188,b,107,a),flip(a)].\n",
      "\n",
      "given #90 (T,wt=8): 62 m(x,y) != y | 1 = x.  [para(5(a,1),18(a,1))].\n",
      "\n",
      "given #91 (T,wt=8): 192 a != 2 | f(2,a) = 1.  [para(81(a,1),18(a,1)),flip(b)].\n",
      "\n",
      "given #92 (T,wt=8): 282 a != x | d(f(2,a),x).  [para(81(a,1),30(a,1))].\n",
      "\n",
      "given #93 (T,wt=9): 59 d(m(x,y),m(x,m(z,y))).  [para(17(a,1),22(a,2))].\n",
      "\n",
      "given #94 (A,wt=14): 32 -d(2,m(x,m(y,m(x,y)))) | d(2,m(x,y)).  [factor(31,b,c)].\n",
      "\n",
      "given #95 (F,wt=7): 257 m(x,m(a,y)) != b.  [para(17(a,1),233(a,1))].\n",
      "\n",
      "given #96 (F,wt=7): 263 m(b,f(2,a)) != a.  [para(81(a,1),128(a,1)),rewrite([5(6)]),flip(a)].\n",
      "\n",
      "given #97 (F,wt=7): 301 m(b,x) != m(a,x).  [para(5(a,1),218(a,2)),flip(a)].\n",
      "\n",
      "given #98 (F,wt=9): 106 -d(2,m(b,m(b,m(b,b)))).  [ur(16,b,101,a),rewrite([17(8),5(7)])].\n",
      "\n",
      "given #99 (T,wt=9): 116 d(x,m(y,m(z,m(x,u)))).  [para(4(a,1),39(a,2))].\n",
      "\n",
      "given #100 (T,wt=7): 368 d(b,m(x,m(a,a))).  [para(11(a,1),116(a,2,2))].\n",
      "\n",
      "given #101 (T,wt=7): 371 d(b,m(a,m(a,x))).  [para(5(a,1),368(a,2)),rewrite([4(5)])].\n",
      "\n",
      "given #102 (T,wt=7): 372 d(b,m(a,m(x,a))).  [para(17(a,1),368(a,2))].\n",
      "\n",
      "given #103 (A,wt=13): 34 m(2,m(b,m(b,x))) = m(a,m(a,x)).  [para(11(a,1),4(a,1,1)),rewrite([4(4),4(9)]),flip(a)].\n",
      "\n",
      "given #104 (F,wt=9): 110 m(x,m(2,y)) != m(x,b).  [ur(6,b,102,a)].\n",
      "\n",
      "given #105 (F,wt=9): 121 m(2,m(x,y)) != m(y,b).  [ur(20,b,102,a),rewrite([4(5)]),flip(a)].\n",
      "\n",
      "given #106 (F,wt=9): 127 m(b,x) != m(x,m(2,y)).  [ur(20,b,102,a(flip)),flip(a)].\n",
      "\n",
      "given #107 (F,wt=9): 139 m(b,m(b,x)) != m(x,2).  [ur(20,b,105,a),rewrite([4(6)]),flip(a)].\n",
      "\n",
      "given #108 (T,wt=9): 158 d(x,m(y,m(z,m(u,x)))).  [para(4(a,1),57(a,2,2))].\n",
      "\n",
      "given #109 (T,wt=9): 159 d(m(b,b),m(x,m(a,a))).  [para(11(a,1),57(a,2,2))].\n",
      "\n",
      "given #110 (T,wt=9): 211 d(m(x,y),m(y,m(x,z))).  [para(5(a,1),23(a,1))].\n",
      "\n",
      "given #111 (T,wt=9): 212 d(m(x,y),m(y,m(z,x))).  [para(5(a,1),23(a,2)),rewrite([4(3)])].\n",
      "\n",
      "given #112 (A,wt=13): 35 m(b,m(b,m(x,2))) = m(x,m(a,a)).  [para(11(a,1),4(a,2,2)),rewrite([17(6),5(5)])].\n",
      "\n",
      "given #113 (F,wt=9): 141 m(x,m(b,b)) != m(x,2).  [ur(6,b,105,a)].\n",
      "\n",
      "given #114 (F,wt=7): 513 m(a,a) != m(2,2).  [para(11(a,1),141(a,1))].\n",
      "\n",
      "given #115 (F,wt=9): 142 m(2,x) != m(x,m(b,b)).  [ur(20,b,105,a(flip)),flip(a)].\n",
      "\n",
      "given #116 (F,wt=9): 143 m(x,m(2,y)) != m(y,b).  [ur(20,b,111,a),rewrite([4(5)]),flip(a)].\n",
      "\n",
      "given #117 (T,wt=9): 213 d(m(x,2),m(x,m(a,a))).  [para(11(a,1),23(a,2,2))].\n",
      "\n",
      "given #118 (T,wt=9): 240 d(f(2,a),m(x,m(y,a))).  [para(4(a,1),200(a,2))].\n",
      "\n",
      "given #119 (T,wt=9): 243 d(2,m(x,m(y,m(z,a)))).  [para(4(a,1),205(a,2,2))].\n",
      "\n",
      "given #120 (T,wt=9): 270 d(2,m(x,m(y,m(a,z)))).  [para(4(a,1),208(a,2))].\n",
      "\n",
      "given #121 (A,wt=12): 36 m(a,a) != m(2,x) | m(b,b) = x.  [para(11(a,1),6(a,1))].\n",
      "\n",
      "given #122 (F,wt=7): 548 m(a,a) != m(2,a).  [ur(36,b,201,a)].\n",
      "\n",
      "given #123 (F,wt=9): 145 m(x,m(y,2)) != m(x,b).  [ur(6,b,111,a)].\n",
      "\n",
      "given #124 (F,wt=9): 146 m(b,x) != m(x,m(y,2)).  [ur(20,b,111,a(flip)),flip(a)].\n",
      "\n",
      "given #125 (F,wt=9): 167 m(b,m(b,b)) != m(2,x).  [ur(8,b,107,a),flip(a)].\n",
      "\n",
      "given #126 (T,wt=9): 278 d(f(2,a),m(x,m(a,y))).  [para(17(a,1),241(a,2))].\n",
      "\n",
      "given #127 (T,wt=9): 305 d(m(a,a),m(a,m(b,b))).  [para(11(a,1),275(a,1)),rewrite([5(8)])].\n",
      "\n",
      "given #128 (T,wt=9): 339 d(m(x,y),m(z,m(y,x))).  [para(5(a,1),59(a,2)),rewrite([4(3)])].\n",
      "\n",
      "given #129 (T,wt=9): 342 d(m(x,f(2,a)),m(x,a)).  [para(81(a,1),59(a,2,2))].\n",
      "\n",
      "given #130 (A,wt=14): 40 m(x,m(y,z)) != m(y,u) | m(x,z) = u.  [para(17(a,1),6(a,1))].\n",
      "\n",
      "given #131 (F,wt=9): 173 m(b,b) != m(x,m(2,y)).  [para(17(a,1),108(a,2))].\n",
      "\n",
      "given #132 (F,wt=7): 639 m(b,b) != m(x,a).  [para(81(a,1),173(a,2,2))].\n",
      "\u0007-------- Proof 1 -------- \n",
      "\n",
      "============================== PROOF =================================\n",
      "\n",
      "% Proof 1 at 0.02 (+ 0.01) seconds.\n",
      "% Length of proof is 27.\n",
      "% Level of proof is 10.\n",
      "% Maximum clause weight is 12.000.\n",
      "% Given clauses 132 (13.200 givens/level).\n",
      "\n",
      "3 m(x,m(y,z)) = m(m(x,y),z).  [assumption].\n",
      "4 m(m(x,y),z) = m(x,m(y,z)).  [copy(3),flip(a)].\n",
      "5 m(x,y) = m(y,x).  [assumption].\n",
      "6 m(x,y) != m(x,z) | y = z.  [assumption].\n",
      "7 -d(x,y) | m(x,f(x,y)) = y.  [assumption].\n",
      "8 m(x,y) != z | d(x,z).  [assumption].\n",
      "9 -d(2,m(x,y)) | d(2,x) | d(2,y).  [assumption].\n",
      "10 m(a,a) = m(2,m(b,b)).  [assumption].\n",
      "11 m(2,m(b,b)) = m(a,a).  [copy(10),flip(a)].\n",
      "12 -d(x,a) | -d(x,b) | x = 1.  [assumption].\n",
      "13 -d(x,a) | -d(x,b) | 1 = x.  [copy(12),flip(c)].\n",
      "14 2 != 1.  [assumption].\n",
      "15 1 != 2.  [copy(14),flip(a)].\n",
      "16 -d(2,m(x,x)) | d(2,x).  [factor(9,b,c)].\n",
      "17 m(x,m(y,z)) = m(y,m(x,z)).  [para(5(a,1),4(a,1,1)),rewrite([4(2)])].\n",
      "33 d(2,m(a,a)).  [resolve(11,a,8,a)].\n",
      "36 m(a,a) != m(2,x) | m(b,b) = x.  [para(11(a,1),6(a,1))].\n",
      "79 d(2,a).  [resolve(33,a,16,a)].\n",
      "81 m(2,f(2,a)) = a.  [resolve(79,a,7,a)].\n",
      "82 -d(2,b).  [ur(13,a,79,a,c,15,a)].\n",
      "101 -d(2,m(b,b)).  [ur(16,b,82,a)].\n",
      "108 m(b,b) != m(2,x).  [ur(8,b,101,a),flip(a)].\n",
      "173 m(b,b) != m(x,m(2,y)).  [para(17(a,1),108(a,2))].\n",
      "186 m(2,m(f(2,a),x)) = m(a,x).  [para(81(a,1),4(a,1,1)),flip(a)].\n",
      "639 m(b,b) != m(x,a).  [para(81(a,1),173(a,2,2))].\n",
      "642 m(a,a) != m(2,m(x,a)).  [ur(36,b,639,a)].\n",
      "643 $F.  [resolve(642,a,186,a(flip))].\n",
      "\n",
      "============================== end of proof ==========================\n",
      "\n",
      "============================== STATISTICS ============================\n",
      "\n",
      "Given=132. Generated=2152. Kept=638. proofs=1.\n",
      "Usable=127. Sos=485. Demods=56. Limbo=1, Disabled=36. Hints=0.\n",
      "Kept_by_rule=0, Deleted_by_rule=0.\n",
      "Forward_subsumed=1514. Back_subsumed=15.\n",
      "Sos_limit_deleted=0. Sos_displaced=0. Sos_removed=0.\n",
      "New_demodulators=58 (2 lex), Back_demodulated=9. Back_unit_deleted=0.\n",
      "Demod_attempts=18123. Demod_rewrites=1480.\n",
      "Res_instance_prunes=0. Para_instance_prunes=0. Basic_paramod_prunes=0.\n",
      "Nonunit_fsub_feature_tests=1506. Nonunit_bsub_feature_tests=528.\n",
      "Megabytes=0.81.\n",
      "User_CPU=0.02, System_CPU=0.01, Wall_clock=0.\n",
      "\n",
      "============================== end of statistics =====================\n",
      "\n",
      "============================== end of search =========================\n",
      "\n",
      "THEOREM PROVED\n",
      "\n",
      "THEOREM PROVED\n",
      "\n",
      "Exiting with 1 proof.\n",
      "\n",
      "------ process 370431 exit (max_proofs) ------\n",
      "\u0007\n",
      "Process 370431 exit (max_proofs) Mon Apr 29 20:47:34 2024\n"
     ]
    }
   ],
   "source": [
    "! prover9 -f /tmp/sqrt2.in"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "This problem is solved instantly. I thought the same problem given to vampire or eprover takes a while (10s for [vampire](https://vprover.github.io/usage.html), a couple minutes for [eprover](https://github.com/eprover/eprover)) but it turns out running them without their command line arguments significantly slows them down. I think this has tripped me up before. These seems like very bad defaults. I used the included `ladr_to_tptp` tool to extract the following [tptp](https://www.tptp.org/) problem."
   ]
  },
  {
   "cell_type": "code",
   "execution_count": null,
   "metadata": {},
   "outputs": [],
   "source": [
    "%%file /tmp/sqrt.p\n",
    "\n",
    "cnf(sos,axiom,m(one,A) = A).\n",
    "cnf(sos,axiom,m(A,one) = A).\n",
    "cnf(sos,axiom,m(A,m(B,C)) = m(m(A,B),C)).\n",
    "cnf(sos,axiom,m(A,B) = m(B,A)).\n",
    "cnf(sos,axiom,m(A,B) != m(A,C) | B = C).\n",
    "cnf(sos,axiom,~ d(A,B) | m(A,f(A,B)) = B).\n",
    "cnf(sos,axiom, m(A,B) != C | d(A,C)).\n",
    "cnf(sos,axiom,~ d(two,m(A,B)) | d(two,A) | d(two,B)).\n",
    "cnf(sos,axiom,m(a,a) = m(two,m(b,b))).\n",
    "cnf(sos,axiom,~ d(A,a) | ~ d(A,b) | A = one).\n",
    "cnf(sos,axiom,~ two = one)."
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "Both of these succeed basically immediately, but the output is quite long."
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 1,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "% Running in auto input_syntax mode. Trying TPTP\n",
      "% lrs+10_11_cond=on:drc=off:flr=on:fsr=off:gsp=on:gs=on:gsem=off:lma=on:msp=off:nm=4:nwc=1.5:nicw=on:sas=z3:sims=off:sp=scramble:stl=188_1169 on sqrt for (600ds/0Mi)\n",
      "perf_event_open failed (instruction limiting will be disabled): Permission denied\n",
      "(If you are seeing 'Permission denied' ask your admin to run 'sudo sysctl -w kernel.perf_event_paranoid=-1' for you.)\n",
      "% Refutation found. Thanks to Tanya!\n",
      "% SZS status Unsatisfiable for sqrt\n",
      "% SZS output start Proof for sqrt\n",
      "fof(f39584,plain,(\n",
      "  $false),\n",
      "  inference(avatar_smt_refutation,[],[f17,f22,f49,f84,f118,f123,f132,f137,f143,f186,f192,f212,f262,f268,f318,f322,f326,f332,f353,f363,f384,f393,f398,f400,f411,f419,f461,f704,f761,f766,f772,f981,f1016,f1022,f1284,f1409,f1414,f1416,f1442,f1447,f1452,f1516,f1569,f1807,f2104,f3775,f3966,f3971,f3977,f4241,f4400,f4406,f4415,f4437,f4442,f4666,f4671,f4677,f4683,f4950,f4963,f5012,f5016,f5020,f5062,f5069,f5217,f5222,f5227,f5242,f5247,f5254,f5259,f5320,f5539,f5544,f5549,f5554,f5559,f5564,f5569,f5574,f5579,f5586,f5591,f5596,f5689,f5695,f5960,f5967,f5976,f5983,f6011,f6026,f6035,f6042,f6048,f6054,f6060,f6066,f6072,f6078,f6084,f6090,f6096,f6102,f6108,f6206,f6211,f6216,f6222,f6229,f6234,f6235,f6245,f7019,f7058,f7063,f7068,f7165,f7166,f7171,f7172,f7177,f7182,f7187,f7194,f7199,f7200,f7208,f7660,f8108,f8114,f8120,f8124,f8130,f8135,f8141,f8146,f8165,f8170,f8175,f8177,f8191,f8196,f8219,f8222,f8234,f9517,f10565,f10647,f10753,f10759,f10819,f10826,f11051,f12292,f12299,f12305,f12428,f12436,f12560,f12750,f14032,f14316,f15133,f15139,f15145,f15151,f15157,f15163,f15164,f15169,f15174,f15183,f15190,f15205,f15210,f15216,f15227,f15233,f15242,f15291,f15297,f19429,f19489,f19498,f19591,f19599,f19668,f22729,f25452,f25458,f25464,f25538,f25718,f26169,f26173,f28151,f28165,f28811,f28822,f29812,f29824,f30299,f30816,f30951,f30975,f31082,f31241,f31255,f31260,f31294,f31482,f32339,f32354,f33397,f33412,f35130,f38938,f39203,f39210,f39215,f39221,f39239,f39245,f39254,f39262,f39265,f39270,f39278,f39290,f39297,f39302,f39310,f39311,f39312,f39313,f39315,f39355,f39365,f39377,f39398,f39410,f39415,f39436,f39449,f39462,f39467,f39472,f39477,f39512,f39525,f39530,f39544,f39550,f39570,f39575,f39578,f39583])).\n",
      "fof(f39583,plain,(\n",
      "  spl0_232 | ~spl0_118),\n",
      "  inference(avatar_split_clause,[],[f39577,f7158,f39580])).\n",
      "fof(f39580,plain,(\n",
      "  spl0_232 <=> b = m(two,f(two,b))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_232])])).\n",
      "fof(f7158,plain,(\n",
      "  spl0_118 <=> d(two,b)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_118])])).\n",
      "fof(f39577,plain,(\n",
      "  b = m(two,f(two,b)) | ~spl0_118),\n",
      "  inference(resolution,[],[f7160,f6])).\n",
      "fof(f6,axiom,(\n",
      "  ( ! [X0,X1] : (~d(X0,X1) | m(X0,f(X0,X1)) = X1) )),\n",
      "  file('/tmp/sqrt.p',unknown)).\n",
      "fof(f7160,plain,(\n",
      "  d(two,b) | ~spl0_118),\n",
      "  inference(avatar_component_clause,[],[f7158])).\n",
      "fof(f39578,plain,(\n",
      "  spl0_1 | ~spl0_17 | ~spl0_118),\n",
      "  inference(avatar_split_clause,[],[f39576,f7158,f315,f14])).\n",
      "fof(f14,plain,(\n",
      "  spl0_1 <=> one = two),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).\n",
      "fof(f315,plain,(\n",
      "  spl0_17 <=> d(two,a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_17])])).\n",
      "fof(f39576,plain,(\n",
      "  ~d(two,a) | one = two | ~spl0_118),\n",
      "  inference(resolution,[],[f7160,f10])).\n",
      "fof(f10,axiom,(\n",
      "  ( ! [X0] : (~d(X0,b) | ~d(X0,a) | one = X0) )),\n",
      "  file('/tmp/sqrt.p',unknown)).\n",
      "fof(f39575,plain,(\n",
      "  spl0_231 | ~spl0_206),\n",
      "  inference(avatar_split_clause,[],[f39568,f39200,f39572])).\n",
      "fof(f39572,plain,(\n",
      "  spl0_231 <=> m(b,b) = m(two,f(two,m(b,b)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_231])])).\n",
      "fof(f39200,plain,(\n",
      "  spl0_206 <=> d(two,m(b,b))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_206])])).\n",
      "fof(f39568,plain,(\n",
      "  m(b,b) = m(two,f(two,m(b,b))) | ~spl0_206),\n",
      "  inference(resolution,[],[f39202,f6])).\n",
      "fof(f39202,plain,(\n",
      "  d(two,m(b,b)) | ~spl0_206),\n",
      "  inference(avatar_component_clause,[],[f39200])).\n",
      "fof(f39570,plain,(\n",
      "  spl0_118 | ~spl0_206),\n",
      "  inference(avatar_split_clause,[],[f39569,f39200,f7158])).\n",
      "fof(f39569,plain,(\n",
      "  d(two,b) | ~spl0_206),\n",
      "  inference(duplicate_literal_removal,[],[f39567])).\n",
      "fof(f39567,plain,(\n",
      "  d(two,b) | d(two,b) | ~spl0_206),\n",
      "  inference(resolution,[],[f39202,f8])).\n",
      "fof(f8,axiom,(\n",
      "  ( ! [X0,X1] : (~d(two,m(X0,X1)) | d(two,X0) | d(two,X1)) )),\n",
      "  file('/tmp/sqrt.p',unknown)).\n",
      "fof(f39550,plain,(\n",
      "  spl0_230 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39545,f38935,f39547])).\n",
      "fof(f39547,plain,(\n",
      "  spl0_230 <=> a = f(m(one,f(two,a)),m(b,b))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_230])])).\n",
      "fof(f38935,plain,(\n",
      "  spl0_205 <=> m(b,b) = m(a,f(two,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_205])])).\n",
      "fof(f39545,plain,(\n",
      "  a = f(m(one,f(two,a)),m(b,b)) | ~spl0_205),\n",
      "  inference(forward_demodulation,[],[f39188,f4])).\n",
      "fof(f4,axiom,(\n",
      "  ( ! [X0,X1] : (m(X0,X1) = m(X1,X0)) )),\n",
      "  file('/tmp/sqrt.p',unknown)).\n",
      "fof(f39188,plain,(\n",
      "  a = f(m(f(two,a),one),m(b,b)) | ~spl0_205),\n",
      "  inference(superposition,[],[f19170,f38937])).\n",
      "fof(f38937,plain,(\n",
      "  m(b,b) = m(a,f(two,a)) | ~spl0_205),\n",
      "  inference(avatar_component_clause,[],[f38935])).\n",
      "fof(f19170,plain,(\n",
      "  ( ! [X0,X1] : (f(m(X0,one),m(X1,X0)) = X1) )),\n",
      "  inference(superposition,[],[f18968,f11212])).\n",
      "fof(f11212,plain,(\n",
      "  ( ! [X0,X1] : (m(X0,X1) = m(m(X1,one),X0)) )),\n",
      "  inference(superposition,[],[f10282,f4])).\n",
      "fof(f10282,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,X0) = m(X1,m(X0,one))) )),\n",
      "  inference(superposition,[],[f10249,f151])).\n",
      "fof(f151,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X0,m(X1,X2)) = m(X2,m(X0,X1))) )),\n",
      "  inference(superposition,[],[f3,f4])).\n",
      "fof(f3,axiom,(\n",
      "  ( ! [X2,X0,X1] : (m(X0,m(X1,X2)) = m(m(X0,X1),X2)) )),\n",
      "  file('/tmp/sqrt.p',unknown)).\n",
      "fof(f10249,plain,(\n",
      "  ( ! [X0,X1] : (m(X0,X1) = m(X1,m(one,X0))) )),\n",
      "  inference(equality_resolution,[],[f8255])).\n",
      "fof(f8255,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X1,m(one,X0)) != X2 | m(X0,X1) = X2) )),\n",
      "  inference(superposition,[],[f50,f151])).\n",
      "fof(f50,plain,(\n",
      "  ( ! [X0,X1] : (m(one,X1) != X0 | X0 = X1) )),\n",
      "  inference(superposition,[],[f5,f1])).\n",
      "fof(f1,axiom,(\n",
      "  ( ! [X0] : (m(one,X0) = X0) )),\n",
      "  file('/tmp/sqrt.p',unknown)).\n",
      "fof(f5,axiom,(\n",
      "  ( ! [X2,X0,X1] : (m(X0,X1) != m(X0,X2) | X1 = X2) )),\n",
      "  file('/tmp/sqrt.p',unknown)).\n",
      "fof(f18968,plain,(\n",
      "  ( ! [X0,X1] : (f(X0,m(X0,X1)) = X1) )),\n",
      "  inference(equality_resolution,[],[f639])).\n",
      "fof(f639,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X0,X1) != m(X0,X2) | f(X0,m(X0,X1)) = X2) )),\n",
      "  inference(superposition,[],[f5,f67])).\n",
      "fof(f67,plain,(\n",
      "  ( ! [X0,X1] : (m(X0,X1) = m(X0,f(X0,m(X0,X1)))) )),\n",
      "  inference(resolution,[],[f6,f12])).\n",
      "fof(f12,plain,(\n",
      "  ( ! [X0,X1] : (d(X0,m(X0,X1))) )),\n",
      "  inference(equality_resolution,[],[f7])).\n",
      "fof(f7,axiom,(\n",
      "  ( ! [X2,X0,X1] : (d(X0,X2) | m(X0,X1) != X2) )),\n",
      "  file('/tmp/sqrt.p',unknown)).\n",
      "fof(f39544,plain,(\n",
      "  spl0_229 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39539,f38935,f39541])).\n",
      "fof(f39541,plain,(\n",
      "  spl0_229 <=> f(two,a) = f(m(a,one),m(b,b))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_229])])).\n",
      "fof(f39539,plain,(\n",
      "  f(two,a) = f(m(a,one),m(b,b)) | ~spl0_205),\n",
      "  inference(forward_demodulation,[],[f39187,f4])).\n",
      "fof(f39187,plain,(\n",
      "  f(two,a) = f(m(one,a),m(b,b)) | ~spl0_205),\n",
      "  inference(superposition,[],[f19169,f38937])).\n",
      "fof(f19169,plain,(\n",
      "  ( ! [X0,X1] : (f(m(one,X0),m(X0,X1)) = X1) )),\n",
      "  inference(superposition,[],[f18968,f10290])).\n",
      "fof(f10290,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,X0) = m(m(one,X1),X0)) )),\n",
      "  inference(superposition,[],[f10249,f4])).\n",
      "fof(f39530,plain,(\n",
      "  spl0_228 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39173,f38935,f39527])).\n",
      "fof(f39527,plain,(\n",
      "  spl0_228 <=> a = f(f(two,a),m(b,b))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_228])])).\n",
      "fof(f39173,plain,(\n",
      "  a = f(f(two,a),m(b,b)) | ~spl0_205),\n",
      "  inference(superposition,[],[f19151,f38937])).\n",
      "fof(f19151,plain,(\n",
      "  ( ! [X0,X1] : (f(X0,m(X1,X0)) = X1) )),\n",
      "  inference(superposition,[],[f18968,f4])).\n",
      "fof(f39525,plain,(\n",
      "  spl0_227 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39172,f38935,f39522])).\n",
      "fof(f39522,plain,(\n",
      "  spl0_227 <=> f(two,a) = f(a,m(b,b))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_227])])).\n",
      "fof(f39172,plain,(\n",
      "  f(two,a) = f(a,m(b,b)) | ~spl0_205),\n",
      "  inference(superposition,[],[f18968,f38937])).\n",
      "fof(f39512,plain,(\n",
      "  ~spl0_211 | spl0_217 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39511,f38935,f39352,f39242])).\n",
      "fof(f39242,plain,(\n",
      "  spl0_211 <=> a = m(b,b)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_211])])).\n",
      "fof(f39352,plain,(\n",
      "  spl0_217 <=> one = f(a,m(b,b))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_217])])).\n",
      "fof(f39511,plain,(\n",
      "  one = f(a,m(b,b)) | a != m(b,b) | ~spl0_205),\n",
      "  inference(forward_demodulation,[],[f39510,f38937])).\n",
      "fof(f39510,plain,(\n",
      "  one = f(a,m(a,f(two,a))) | a != m(b,b) | ~spl0_205),\n",
      "  inference(forward_demodulation,[],[f39165,f4])).\n",
      "fof(f39165,plain,(\n",
      "  a != m(b,b) | one = f(a,m(f(two,a),a)) | ~spl0_205),\n",
      "  inference(superposition,[],[f11431,f38937])).\n",
      "fof(f11431,plain,(\n",
      "  ( ! [X0,X1] : (m(X0,X1) != X0 | one = f(X0,m(X1,X0))) )),\n",
      "  inference(forward_demodulation,[],[f11298,f10269])).\n",
      "fof(f10269,plain,(\n",
      "  ( ! [X0,X1] : (m(X0,X1) = m(X1,m(X0,one))) )),\n",
      "  inference(superposition,[],[f10249,f4])).\n",
      "fof(f11298,plain,(\n",
      "  ( ! [X0,X1] : (m(X0,X1) != X0 | one = f(X0,m(X0,m(X1,one)))) )),\n",
      "  inference(superposition,[],[f644,f10282])).\n",
      "fof(f644,plain,(\n",
      "  ( ! [X0,X1] : (m(X0,X1) != X0 | one = f(X0,m(X0,X1))) )),\n",
      "  inference(superposition,[],[f51,f67])).\n",
      "fof(f51,plain,(\n",
      "  ( ! [X0,X1] : (m(X0,X1) != X0 | one = X1) )),\n",
      "  inference(superposition,[],[f5,f2])).\n",
      "fof(f2,axiom,(\n",
      "  ( ! [X0] : (m(X0,one) = X0) )),\n",
      "  file('/tmp/sqrt.p',unknown)).\n",
      "fof(f39477,plain,(\n",
      "  spl0_226 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39139,f38935,f39474])).\n",
      "fof(f39474,plain,(\n",
      "  spl0_226 <=> d(m(a,one),m(b,b))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_226])])).\n",
      "fof(f39139,plain,(\n",
      "  d(m(a,one),m(b,b)) | ~spl0_205),\n",
      "  inference(superposition,[],[f10311,f38937])).\n",
      "fof(f10311,plain,(\n",
      "  ( ! [X0,X1] : (d(m(X1,one),m(X1,X0))) )),\n",
      "  inference(superposition,[],[f549,f10249])).\n",
      "fof(f549,plain,(\n",
      "  ( ! [X2,X0,X1] : (d(m(X0,X2),m(X1,m(X2,X0)))) )),\n",
      "  inference(forward_demodulation,[],[f542,f3])).\n",
      "fof(f542,plain,(\n",
      "  ( ! [X2,X0,X1] : (d(m(X0,X2),m(m(X1,X2),X0))) )),\n",
      "  inference(superposition,[],[f237,f4])).\n",
      "fof(f237,plain,(\n",
      "  ( ! [X2,X0,X1] : (d(m(X2,X0),m(X2,m(X1,X0)))) )),\n",
      "  inference(superposition,[],[f159,f4])).\n",
      "fof(f159,plain,(\n",
      "  ( ! [X2,X0,X1] : (d(m(X0,X1),m(X0,m(X1,X2)))) )),\n",
      "  inference(superposition,[],[f12,f3])).\n",
      "fof(f39472,plain,(\n",
      "  spl0_215 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39471,f38935,f39275])).\n",
      "fof(f39275,plain,(\n",
      "  spl0_215 <=> d(m(one,f(two,a)),m(b,b))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_215])])).\n",
      "fof(f39471,plain,(\n",
      "  d(m(one,f(two,a)),m(b,b)) | ~spl0_205),\n",
      "  inference(forward_demodulation,[],[f39137,f4])).\n",
      "fof(f39137,plain,(\n",
      "  d(m(f(two,a),one),m(b,b)) | ~spl0_205),\n",
      "  inference(superposition,[],[f10300,f38937])).\n",
      "fof(f10300,plain,(\n",
      "  ( ! [X0,X1] : (d(m(X0,one),m(X1,X0))) )),\n",
      "  inference(superposition,[],[f159,f10249])).\n",
      "fof(f39467,plain,(\n",
      "  spl0_24 | spl0_225 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39463,f38935,f39465,f386])).\n",
      "fof(f386,plain,(\n",
      "  spl0_24 <=> one = f(two,a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_24])])).\n",
      "fof(f39465,plain,(\n",
      "  spl0_225 <=> ! [X0] : m(a,X0) != m(b,m(b,X0))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_225])])).\n",
      "fof(f39463,plain,(\n",
      "  ( ! [X0] : (m(a,X0) != m(b,m(b,X0)) | one = f(two,a)) ) | ~spl0_205),\n",
      "  inference(forward_literal_rewriting,[],[f39134,f2902])).\n",
      "fof(f2902,plain,(\n",
      "  ( ! [X2,X3,X0,X1] : (m(X2,m(X1,X0)) != X3 | m(m(X1,X2),X0) = X3) )),\n",
      "  inference(superposition,[],[f966,f885])).\n",
      "fof(f885,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X2,m(X0,X1)) = m(X1,m(X0,X2))) )),\n",
      "  inference(superposition,[],[f145,f4])).\n",
      "fof(f145,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X0,m(X1,X2)) = m(m(X1,X0),X2)) )),\n",
      "  inference(superposition,[],[f3,f4])).\n",
      "fof(f966,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X1,X0) != X2 | m(X0,X1) = X2) )),\n",
      "  inference(forward_demodulation,[],[f923,f2])).\n",
      "fof(f923,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X1,m(X0,one)) != X2 | m(X0,X1) = X2) )),\n",
      "  inference(superposition,[],[f61,f145])).\n",
      "fof(f61,plain,(\n",
      "  ( ! [X0,X1] : (m(X0,one) != X1 | X0 = X1) )),\n",
      "  inference(superposition,[],[f50,f4])).\n",
      "fof(f39134,plain,(\n",
      "  ( ! [X0] : (m(m(b,b),X0) != m(a,X0) | one = f(two,a)) ) | ~spl0_205),\n",
      "  inference(superposition,[],[f9502,f38937])).\n",
      "fof(f9502,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X1,X0) != m(m(X1,X2),X0) | one = X2) )),\n",
      "  inference(superposition,[],[f897,f4])).\n",
      "fof(f897,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X0,X1) != m(X1,m(X0,X2)) | one = X2) )),\n",
      "  inference(superposition,[],[f51,f145])).\n",
      "fof(f39462,plain,(\n",
      "  ~spl0_44 | spl0_163 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39461,f38935,f15176,f1562])).\n",
      "fof(f1562,plain,(\n",
      "  spl0_44 <=> m(b,b) = f(two,a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_44])])).\n",
      "fof(f15176,plain,(\n",
      "  spl0_163 <=> one = a),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_163])])).\n",
      "fof(f39461,plain,(\n",
      "  one = a | m(b,b) != f(two,a) | ~spl0_205),\n",
      "  inference(forward_demodulation,[],[f39133,f18968])).\n",
      "fof(f39133,plain,(\n",
      "  m(b,b) != f(two,a) | one = f(f(two,a),m(f(two,a),a)) | ~spl0_205),\n",
      "  inference(superposition,[],[f9245,f38937])).\n",
      "fof(f9245,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,X0) != X0 | one = f(X0,m(X0,X1))) )),\n",
      "  inference(superposition,[],[f644,f4])).\n",
      "fof(f39449,plain,(\n",
      "  spl0_223 | spl0_224 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39441,f38935,f39446,f39443])).\n",
      "fof(f39443,plain,(\n",
      "  spl0_223 <=> ! [X0] : m(X0,m(b,b)) != X0),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_223])])).\n",
      "fof(f39446,plain,(\n",
      "  spl0_224 <=> one = m(a,f(two,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_224])])).\n",
      "fof(f39441,plain,(\n",
      "  ( ! [X0] : (one = m(a,f(two,a)) | m(X0,m(b,b)) != X0) ) | ~spl0_205),\n",
      "  inference(forward_literal_rewriting,[],[f39129,f966])).\n",
      "fof(f39129,plain,(\n",
      "  ( ! [X0] : (m(X0,m(b,b)) != X0 | one = m(f(two,a),a)) ) | ~spl0_205),\n",
      "  inference(superposition,[],[f6514,f38937])).\n",
      "fof(f6514,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X0,m(X2,X1)) != X0 | one = m(X1,X2)) )),\n",
      "  inference(superposition,[],[f51,f2380])).\n",
      "fof(f2380,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X1,m(X0,X2)) = m(X1,m(X2,X0))) )),\n",
      "  inference(superposition,[],[f883,f151])).\n",
      "fof(f883,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X0,m(X1,X2)) = m(X1,m(X0,X2))) )),\n",
      "  inference(superposition,[],[f145,f3])).\n",
      "fof(f39436,plain,(\n",
      "  spl0_24 | spl0_222 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39432,f38935,f39434,f386])).\n",
      "fof(f39434,plain,(\n",
      "  spl0_222 <=> ! [X0] : m(X0,a) != m(b,m(b,X0))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_222])])).\n",
      "fof(f39432,plain,(\n",
      "  ( ! [X0] : (m(X0,a) != m(b,m(b,X0)) | one = f(two,a)) ) | ~spl0_205),\n",
      "  inference(forward_literal_rewriting,[],[f39126,f2902])).\n",
      "fof(f39126,plain,(\n",
      "  ( ! [X0] : (m(m(b,b),X0) != m(X0,a) | one = f(two,a)) ) | ~spl0_205),\n",
      "  inference(superposition,[],[f5000,f38937])).\n",
      "fof(f5000,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X0,X1) != m(m(X1,X2),X0) | one = X2) )),\n",
      "  inference(superposition,[],[f161,f4])).\n",
      "fof(f161,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X0,m(X1,X2)) != m(X0,X1) | one = X2) )),\n",
      "  inference(superposition,[],[f51,f3])).\n",
      "fof(f39415,plain,(\n",
      "  spl0_163 | spl0_221 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39114,f38935,f39413,f15176])).\n",
      "fof(f39413,plain,(\n",
      "  spl0_221 <=> ! [X0] : m(X0,m(b,b)) != m(X0,f(two,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_221])])).\n",
      "fof(f39114,plain,(\n",
      "  ( ! [X0] : (m(X0,m(b,b)) != m(X0,f(two,a)) | one = a) ) | ~spl0_205),\n",
      "  inference(superposition,[],[f2464,f38937])).\n",
      "fof(f2464,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X1,X2) != m(X1,m(X0,X2)) | one = X0) )),\n",
      "  inference(superposition,[],[f65,f883])).\n",
      "fof(f65,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,X0) != X0 | one = X1) )),\n",
      "  inference(superposition,[],[f51,f4])).\n",
      "fof(f39410,plain,(\n",
      "  ~spl0_44 | spl0_163 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39409,f38935,f15176,f1562])).\n",
      "fof(f39409,plain,(\n",
      "  one = a | m(b,b) != f(two,a) | ~spl0_205),\n",
      "  inference(forward_demodulation,[],[f39111,f19151])).\n",
      "fof(f39111,plain,(\n",
      "  m(b,b) != f(two,a) | one = f(f(two,a),m(a,f(two,a))) | ~spl0_205),\n",
      "  inference(superposition,[],[f2253,f38937])).\n",
      "fof(f2253,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,X0) != X0 | one = f(X0,m(X1,X0))) )),\n",
      "  inference(superposition,[],[f51,f68])).\n",
      "fof(f68,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,X0) = m(X0,f(X0,m(X1,X0)))) )),\n",
      "  inference(resolution,[],[f6,f31])).\n",
      "fof(f31,plain,(\n",
      "  ( ! [X0,X1] : (d(X0,m(X1,X0))) )),\n",
      "  inference(superposition,[],[f12,f4])).\n",
      "fof(f39398,plain,(\n",
      "  spl0_163 | spl0_220 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39096,f38935,f39396,f15176])).\n",
      "fof(f39396,plain,(\n",
      "  spl0_220 <=> ! [X0] : m(X0,m(b,b)) != m(f(two,a),X0)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_220])])).\n",
      "fof(f39096,plain,(\n",
      "  ( ! [X0] : (m(X0,m(b,b)) != m(f(two,a),X0) | one = a) ) | ~spl0_205),\n",
      "  inference(superposition,[],[f1139,f38937])).\n",
      "fof(f1139,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X1,X2) != m(X2,m(X0,X1)) | one = X0) )),\n",
      "  inference(superposition,[],[f65,f151])).\n",
      "fof(f39377,plain,(\n",
      "  spl0_24 | spl0_219 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39074,f38935,f39375,f386])).\n",
      "fof(f39375,plain,(\n",
      "  spl0_219 <=> ! [X0] : m(a,X0) != m(X0,m(b,b))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_219])])).\n",
      "fof(f39074,plain,(\n",
      "  ( ! [X0] : (m(a,X0) != m(X0,m(b,b)) | one = f(two,a)) ) | ~spl0_205),\n",
      "  inference(superposition,[],[f897,f38937])).\n",
      "fof(f39365,plain,(\n",
      "  spl0_218 | ~spl0_44 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39067,f38935,f1562,f39362])).\n",
      "fof(f39362,plain,(\n",
      "  spl0_218 <=> a = f(f(two,a),f(two,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_218])])).\n",
      "fof(f39067,plain,(\n",
      "  m(b,b) != f(two,a) | a = f(f(two,a),f(two,a)) | ~spl0_205),\n",
      "  inference(superposition,[],[f833,f38937])).\n",
      "fof(f833,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,X0) != X0 | f(X0,X0) = X1) )),\n",
      "  inference(superposition,[],[f74,f4])).\n",
      "fof(f74,plain,(\n",
      "  ( ! [X0,X1] : (m(X0,X1) != X0 | f(X0,X0) = X1) )),\n",
      "  inference(superposition,[],[f5,f69])).\n",
      "fof(f69,plain,(\n",
      "  ( ! [X0] : (m(X0,f(X0,X0)) = X0) )),\n",
      "  inference(resolution,[],[f6,f24])).\n",
      "fof(f24,plain,(\n",
      "  ( ! [X0] : (d(X0,X0)) )),\n",
      "  inference(superposition,[],[f12,f2])).\n",
      "fof(f39355,plain,(\n",
      "  ~spl0_211 | spl0_217 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39350,f38935,f39352,f39242])).\n",
      "fof(f39350,plain,(\n",
      "  one = f(a,m(b,b)) | a != m(b,b) | ~spl0_205),\n",
      "  inference(forward_demodulation,[],[f39062,f38937])).\n",
      "fof(f39062,plain,(\n",
      "  a != m(b,b) | one = f(a,m(a,f(two,a))) | ~spl0_205),\n",
      "  inference(superposition,[],[f644,f38937])).\n",
      "fof(f39315,plain,(\n",
      "  ~spl0_211 | spl0_24 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39314,f38935,f386,f39242])).\n",
      "fof(f39314,plain,(\n",
      "  one = f(two,a) | a != m(b,b) | ~spl0_205),\n",
      "  inference(forward_demodulation,[],[f39014,f79])).\n",
      "fof(f79,plain,(\n",
      "  ( ! [X0] : (one = f(X0,X0)) )),\n",
      "  inference(trivial_inequality_removal,[],[f78])).\n",
      "fof(f78,plain,(\n",
      "  ( ! [X0] : (X0 != X0 | one = f(X0,X0)) )),\n",
      "  inference(superposition,[],[f51,f69])).\n",
      "fof(f39014,plain,(\n",
      "  a != m(b,b) | f(two,a) = f(a,a) | ~spl0_205),\n",
      "  inference(superposition,[],[f74,f38937])).\n",
      "fof(f39313,plain,(\n",
      "  spl0_163 | ~spl0_44 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39013,f38935,f1562,f15176])).\n",
      "fof(f39013,plain,(\n",
      "  m(b,b) != f(two,a) | one = a | ~spl0_205),\n",
      "  inference(superposition,[],[f65,f38937])).\n",
      "fof(f39312,plain,(\n",
      "  spl0_24 | ~spl0_211 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39011,f38935,f39242,f386])).\n",
      "fof(f39011,plain,(\n",
      "  a != m(b,b) | one = f(two,a) | ~spl0_205),\n",
      "  inference(superposition,[],[f51,f38937])).\n",
      "fof(f39311,plain,(\n",
      "  spl0_208 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39010,f38935,f39212])).\n",
      "fof(f39212,plain,(\n",
      "  spl0_208 <=> d(f(two,a),m(b,b))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_208])])).\n",
      "fof(f39010,plain,(\n",
      "  d(f(two,a),m(b,b)) | ~spl0_205),\n",
      "  inference(superposition,[],[f31,f38937])).\n",
      "fof(f39310,plain,(\n",
      "  spl0_207 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39009,f38935,f39207])).\n",
      "fof(f39207,plain,(\n",
      "  spl0_207 <=> d(a,m(b,b))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_207])])).\n",
      "fof(f39009,plain,(\n",
      "  d(a,m(b,b)) | ~spl0_205),\n",
      "  inference(superposition,[],[f12,f38937])).\n",
      "fof(f39302,plain,(\n",
      "  spl0_25 | ~spl0_211 | ~spl0_20 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f38997,f38935,f329,f39242,f390])).\n",
      "fof(f390,plain,(\n",
      "  spl0_25 <=> two = a),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_25])])).\n",
      "fof(f329,plain,(\n",
      "  spl0_20 <=> a = m(two,f(two,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_20])])).\n",
      "fof(f38997,plain,(\n",
      "  a != m(b,b) | two = a | (~spl0_20 | ~spl0_205)),\n",
      "  inference(superposition,[],[f822,f38937])).\n",
      "fof(f822,plain,(\n",
      "  ( ! [X0] : (a != m(X0,f(two,a)) | two = X0) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f469,f331])).\n",
      "fof(f331,plain,(\n",
      "  a = m(two,f(two,a)) | ~spl0_20),\n",
      "  inference(avatar_component_clause,[],[f329])).\n",
      "fof(f469,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X1,X0) != m(X2,X0) | X1 = X2) )),\n",
      "  inference(superposition,[],[f52,f4])).\n",
      "fof(f52,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X1,X0) != m(X0,X2) | X1 = X2) )),\n",
      "  inference(superposition,[],[f5,f4])).\n",
      "fof(f39297,plain,(\n",
      "  ~spl0_216 | ~spl0_195 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39292,f38935,f30945,f39294])).\n",
      "fof(f39294,plain,(\n",
      "  spl0_216 <=> f(two,a) = m(b,m(b,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_216])])).\n",
      "fof(f30945,plain,(\n",
      "  spl0_195 <=> ! [X0] : m(X0,m(a,a)) != X0),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_195])])).\n",
      "fof(f39292,plain,(\n",
      "  f(two,a) != m(b,m(b,a)) | (~spl0_195 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f39291,f4])).\n",
      "fof(f39291,plain,(\n",
      "  f(two,a) != m(b,m(a,b)) | (~spl0_195 | ~spl0_205)),\n",
      "  inference(forward_literal_rewriting,[],[f38993,f6853])).\n",
      "fof(f6853,plain,(\n",
      "  ( ! [X2,X3,X0,X1] : (m(X0,m(X2,X1)) != X3 | m(X2,m(X0,X1)) = X3) )),\n",
      "  inference(superposition,[],[f966,f2382])).\n",
      "fof(f2382,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(m(X1,X2),X0) = m(X1,m(X0,X2))) )),\n",
      "  inference(superposition,[],[f883,f4])).\n",
      "fof(f38993,plain,(\n",
      "  f(two,a) != m(a,m(b,b)) | (~spl0_195 | ~spl0_205)),\n",
      "  inference(superposition,[],[f31435,f38937])).\n",
      "fof(f31435,plain,(\n",
      "  ( ! [X0] : (m(a,m(a,X0)) != X0) ) | ~spl0_195),\n",
      "  inference(superposition,[],[f30946,f151])).\n",
      "fof(f30946,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) != X0) ) | ~spl0_195),\n",
      "  inference(avatar_component_clause,[],[f30945])).\n",
      "fof(f39290,plain,(\n",
      "  spl0_87 | ~spl0_20 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39289,f38935,f329,f5686])).\n",
      "fof(f5686,plain,(\n",
      "  spl0_87 <=> d(m(two,two),m(b,m(b,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_87])])).\n",
      "fof(f39289,plain,(\n",
      "  d(m(two,two),m(b,m(b,a))) | (~spl0_20 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f39288,f151])).\n",
      "fof(f39288,plain,(\n",
      "  d(m(two,two),m(b,m(a,b))) | (~spl0_20 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f38992,f151])).\n",
      "fof(f38992,plain,(\n",
      "  d(m(two,two),m(a,m(b,b))) | (~spl0_20 | ~spl0_205)),\n",
      "  inference(superposition,[],[f28264,f38937])).\n",
      "fof(f28264,plain,(\n",
      "  ( ! [X0] : (d(m(two,two),m(a,m(a,X0)))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f27826,f4])).\n",
      "fof(f27826,plain,(\n",
      "  ( ! [X0] : (d(m(two,two),m(a,m(X0,a)))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f3592,f1102])).\n",
      "fof(f1102,plain,(\n",
      "  ( ! [X0] : (m(X0,a) = m(two,m(f(two,a),X0))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f151,f331])).\n",
      "fof(f3592,plain,(\n",
      "  ( ! [X0,X1] : (d(m(X0,two),m(a,m(X0,X1)))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f3506,f151])).\n",
      "fof(f3506,plain,(\n",
      "  ( ! [X0,X1] : (d(m(X0,two),m(X0,m(X1,a)))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f1143,f331])).\n",
      "fof(f1143,plain,(\n",
      "  ( ! [X2,X3,X0,X1] : (d(m(X3,X0),m(X3,m(X2,m(X0,X1))))) )),\n",
      "  inference(superposition,[],[f159,f151])).\n",
      "fof(f39278,plain,(\n",
      "  spl0_215 | ~spl0_20 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f38986,f38935,f329,f39275])).\n",
      "fof(f38986,plain,(\n",
      "  d(m(one,f(two,a)),m(b,b)) | (~spl0_20 | ~spl0_205)),\n",
      "  inference(superposition,[],[f15299,f38937])).\n",
      "fof(f15299,plain,(\n",
      "  ( ! [X0] : (d(m(one,f(two,a)),m(a,X0))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f12893,f4])).\n",
      "fof(f12893,plain,(\n",
      "  ( ! [X0] : (d(m(one,f(two,a)),m(X0,a))) ) | ~spl0_20),\n",
      "  inference(forward_demodulation,[],[f12845,f4])).\n",
      "fof(f12845,plain,(\n",
      "  ( ! [X0] : (d(m(f(two,a),one),m(X0,a))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f10621,f331])).\n",
      "fof(f10621,plain,(\n",
      "  ( ! [X2,X0,X1] : (d(m(X2,one),m(X1,m(X0,X2)))) )),\n",
      "  inference(superposition,[],[f10300,f145])).\n",
      "fof(f39270,plain,(\n",
      "  spl0_214 | ~spl0_20 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f38980,f38935,f329,f39267])).\n",
      "fof(f39267,plain,(\n",
      "  spl0_214 <=> d(m(two,one),m(b,b))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_214])])).\n",
      "fof(f38980,plain,(\n",
      "  d(m(two,one),m(b,b)) | (~spl0_20 | ~spl0_205)),\n",
      "  inference(superposition,[],[f10499,f38937])).\n",
      "fof(f10499,plain,(\n",
      "  ( ! [X0] : (d(m(two,one),m(a,X0))) ) | ~spl0_20),\n",
      "  inference(forward_demodulation,[],[f10334,f4])).\n",
      "fof(f10334,plain,(\n",
      "  ( ! [X0] : (d(m(one,two),m(a,X0))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f3589,f10249])).\n",
      "fof(f3589,plain,(\n",
      "  ( ! [X0,X1] : (d(m(X0,two),m(X1,m(X0,a)))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f3506,f883])).\n",
      "fof(f39265,plain,(\n",
      "  ~spl0_213 | ~spl0_137 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39264,f38935,f9515,f39259])).\n",
      "fof(f39259,plain,(\n",
      "  spl0_213 <=> m(b,m(b,a)) = m(b,f(two,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_213])])).\n",
      "fof(f9515,plain,(\n",
      "  spl0_137 <=> ! [X0] : m(X0,m(a,a)) != m(b,X0)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_137])])).\n",
      "fof(f39264,plain,(\n",
      "  m(b,m(b,a)) != m(b,f(two,a)) | (~spl0_137 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f39263,f151])).\n",
      "fof(f39263,plain,(\n",
      "  m(b,m(a,b)) != m(b,f(two,a)) | (~spl0_137 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f38979,f151])).\n",
      "fof(f38979,plain,(\n",
      "  m(a,m(b,b)) != m(b,f(two,a)) | (~spl0_137 | ~spl0_205)),\n",
      "  inference(superposition,[],[f9521,f38937])).\n",
      "fof(f9521,plain,(\n",
      "  ( ! [X0] : (m(a,m(a,X0)) != m(b,X0)) ) | ~spl0_137),\n",
      "  inference(superposition,[],[f9516,f151])).\n",
      "fof(f9516,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) != m(b,X0)) ) | ~spl0_137),\n",
      "  inference(avatar_component_clause,[],[f9515])).\n",
      "fof(f39262,plain,(\n",
      "  ~spl0_213 | ~spl0_66 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39257,f38935,f5018,f39259])).\n",
      "fof(f5018,plain,(\n",
      "  spl0_66 <=> ! [X0] : m(X0,m(a,a)) != m(X0,b)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_66])])).\n",
      "fof(f39257,plain,(\n",
      "  m(b,m(b,a)) != m(b,f(two,a)) | (~spl0_66 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f39256,f151])).\n",
      "fof(f39256,plain,(\n",
      "  m(b,m(a,b)) != m(b,f(two,a)) | (~spl0_66 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f39255,f151])).\n",
      "fof(f39255,plain,(\n",
      "  m(a,m(b,b)) != m(b,f(two,a)) | (~spl0_66 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f38978,f4])).\n",
      "fof(f38978,plain,(\n",
      "  m(a,m(b,b)) != m(f(two,a),b) | (~spl0_66 | ~spl0_205)),\n",
      "  inference(superposition,[],[f9014,f38937])).\n",
      "fof(f9014,plain,(\n",
      "  ( ! [X0] : (m(a,m(a,X0)) != m(X0,b)) ) | ~spl0_66),\n",
      "  inference(superposition,[],[f5019,f151])).\n",
      "fof(f5019,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) != m(X0,b)) ) | ~spl0_66),\n",
      "  inference(avatar_component_clause,[],[f5018])).\n",
      "fof(f39254,plain,(\n",
      "  ~spl0_212 | ~spl0_20 | ~spl0_64 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39249,f38935,f5010,f329,f39251])).\n",
      "fof(f39251,plain,(\n",
      "  spl0_212 <=> a = m(b,m(b,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_212])])).\n",
      "fof(f5010,plain,(\n",
      "  spl0_64 <=> ! [X0] : m(X0,two) != m(X0,m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_64])])).\n",
      "fof(f39249,plain,(\n",
      "  a != m(b,m(b,a)) | (~spl0_20 | ~spl0_64 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f39248,f331])).\n",
      "fof(f39248,plain,(\n",
      "  m(two,f(two,a)) != m(b,m(b,a)) | (~spl0_64 | ~spl0_205)),\n",
      "  inference(forward_literal_rewriting,[],[f39247,f966])).\n",
      "fof(f39247,plain,(\n",
      "  m(f(two,a),two) != m(b,m(b,a)) | (~spl0_64 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f39246,f151])).\n",
      "fof(f39246,plain,(\n",
      "  m(f(two,a),two) != m(b,m(a,b)) | (~spl0_64 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f38977,f151])).\n",
      "fof(f38977,plain,(\n",
      "  m(f(two,a),two) != m(a,m(b,b)) | (~spl0_64 | ~spl0_205)),\n",
      "  inference(superposition,[],[f8955,f38937])).\n",
      "fof(f8955,plain,(\n",
      "  ( ! [X0] : (m(X0,two) != m(a,m(a,X0))) ) | ~spl0_64),\n",
      "  inference(superposition,[],[f5011,f151])).\n",
      "fof(f5011,plain,(\n",
      "  ( ! [X0] : (m(X0,two) != m(X0,m(a,a))) ) | ~spl0_64),\n",
      "  inference(avatar_component_clause,[],[f5010])).\n",
      "fof(f39245,plain,(\n",
      "  ~spl0_211 | ~spl0_20 | ~spl0_65 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39240,f38935,f5014,f329,f39242])).\n",
      "fof(f5014,plain,(\n",
      "  spl0_65 <=> ! [X0] : m(X0,two) != m(X0,a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_65])])).\n",
      "fof(f39240,plain,(\n",
      "  a != m(b,b) | (~spl0_20 | ~spl0_65 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f38976,f331])).\n",
      "fof(f38976,plain,(\n",
      "  m(b,b) != m(two,f(two,a)) | (~spl0_65 | ~spl0_205)),\n",
      "  inference(superposition,[],[f8370,f38937])).\n",
      "fof(f8370,plain,(\n",
      "  ( ! [X0] : (m(two,X0) != m(a,X0)) ) | ~spl0_65),\n",
      "  inference(superposition,[],[f7437,f4])).\n",
      "fof(f7437,plain,(\n",
      "  ( ! [X0] : (m(X0,two) != m(a,X0)) ) | ~spl0_65),\n",
      "  inference(superposition,[],[f5015,f4])).\n",
      "fof(f5015,plain,(\n",
      "  ( ! [X0] : (m(X0,two) != m(X0,a)) ) | ~spl0_65),\n",
      "  inference(avatar_component_clause,[],[f5014])).\n",
      "fof(f39239,plain,(\n",
      "  ~spl0_210 | ~spl0_20 | ~spl0_126 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39234,f38935,f7658,f329,f39236])).\n",
      "fof(f39236,plain,(\n",
      "  spl0_210 <=> m(b,a) = m(b,m(b,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_210])])).\n",
      "fof(f7658,plain,(\n",
      "  spl0_126 <=> ! [X0] : m(a,m(a,X0)) != m(b,m(X0,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_126])])).\n",
      "fof(f39234,plain,(\n",
      "  m(b,a) != m(b,m(b,a)) | (~spl0_20 | ~spl0_126 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f39233,f4])).\n",
      "fof(f39233,plain,(\n",
      "  m(b,m(a,b)) != m(b,a) | (~spl0_20 | ~spl0_126 | ~spl0_205)),\n",
      "  inference(forward_literal_rewriting,[],[f39232,f6853])).\n",
      "fof(f39232,plain,(\n",
      "  m(a,m(b,b)) != m(b,a) | (~spl0_20 | ~spl0_126 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f39231,f331])).\n",
      "fof(f39231,plain,(\n",
      "  m(a,m(b,b)) != m(b,m(two,f(two,a))) | (~spl0_126 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f38975,f4])).\n",
      "fof(f38975,plain,(\n",
      "  m(a,m(b,b)) != m(b,m(f(two,a),two)) | (~spl0_126 | ~spl0_205)),\n",
      "  inference(superposition,[],[f7659,f38937])).\n",
      "fof(f7659,plain,(\n",
      "  ( ! [X0] : (m(a,m(a,X0)) != m(b,m(X0,two))) ) | ~spl0_126),\n",
      "  inference(avatar_component_clause,[],[f7658])).\n",
      "fof(f39221,plain,(\n",
      "  spl0_209 | ~spl0_20 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f38951,f38935,f329,f39218])).\n",
      "fof(f39218,plain,(\n",
      "  spl0_209 <=> d(m(f(two,a),f(two,a)),m(b,b))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_209])])).\n",
      "fof(f38951,plain,(\n",
      "  d(m(f(two,a),f(two,a)),m(b,b)) | (~spl0_20 | ~spl0_205)),\n",
      "  inference(superposition,[],[f1040,f38937])).\n",
      "fof(f1040,plain,(\n",
      "  ( ! [X0] : (d(m(X0,f(two,a)),m(a,X0))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f541,f4])).\n",
      "fof(f541,plain,(\n",
      "  ( ! [X0] : (d(m(X0,f(two,a)),m(X0,a))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f237,f331])).\n",
      "fof(f39215,plain,(\n",
      "  spl0_208 | ~spl0_20 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f38945,f38935,f329,f39212])).\n",
      "fof(f38945,plain,(\n",
      "  d(f(two,a),m(b,b)) | (~spl0_20 | ~spl0_205)),\n",
      "  inference(superposition,[],[f440,f38937])).\n",
      "fof(f440,plain,(\n",
      "  ( ! [X0] : (d(f(two,a),m(a,X0))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f374,f4])).\n",
      "fof(f374,plain,(\n",
      "  ( ! [X0] : (d(f(two,a),m(X0,a))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f160,f331])).\n",
      "fof(f160,plain,(\n",
      "  ( ! [X2,X0,X1] : (d(X2,m(X0,m(X1,X2)))) )),\n",
      "  inference(superposition,[],[f31,f3])).\n",
      "fof(f39210,plain,(\n",
      "  spl0_207 | ~spl0_20 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f39205,f38935,f329,f39207])).\n",
      "fof(f39205,plain,(\n",
      "  d(a,m(b,b)) | (~spl0_20 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f39204,f331])).\n",
      "fof(f39204,plain,(\n",
      "  d(m(two,f(two,a)),m(b,b)) | (~spl0_20 | ~spl0_205)),\n",
      "  inference(forward_demodulation,[],[f38944,f4])).\n",
      "fof(f38944,plain,(\n",
      "  d(m(f(two,a),two),m(b,b)) | (~spl0_20 | ~spl0_205)),\n",
      "  inference(superposition,[],[f431,f38937])).\n",
      "fof(f431,plain,(\n",
      "  ( ! [X0] : (d(m(X0,two),m(a,X0))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f372,f4])).\n",
      "fof(f372,plain,(\n",
      "  ( ! [X0] : (d(m(X0,two),m(X0,a))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f159,f331])).\n",
      "fof(f39203,plain,(\n",
      "  spl0_206 | ~spl0_18 | ~spl0_205),\n",
      "  inference(avatar_split_clause,[],[f38943,f38935,f320,f39200])).\n",
      "fof(f320,plain,(\n",
      "  spl0_18 <=> ! [X0] : d(two,m(a,X0))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_18])])).\n",
      "fof(f38943,plain,(\n",
      "  d(two,m(b,b)) | (~spl0_18 | ~spl0_205)),\n",
      "  inference(superposition,[],[f321,f38937])).\n",
      "fof(f321,plain,(\n",
      "  ( ! [X0] : (d(two,m(a,X0))) ) | ~spl0_18),\n",
      "  inference(avatar_component_clause,[],[f320])).\n",
      "fof(f38938,plain,(\n",
      "  spl0_205 | ~spl0_20 | ~spl0_53),\n",
      "  inference(avatar_split_clause,[],[f38933,f4397,f329,f38935])).\n",
      "fof(f4397,plain,(\n",
      "  spl0_53 <=> m(b,b) = f(two,m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_53])])).\n",
      "fof(f38933,plain,(\n",
      "  m(b,b) = m(a,f(two,a)) | (~spl0_20 | ~spl0_53)),\n",
      "  inference(forward_literal_rewriting,[],[f38926,f966])).\n",
      "fof(f38926,plain,(\n",
      "  m(b,b) = m(f(two,a),a) | (~spl0_20 | ~spl0_53)),\n",
      "  inference(superposition,[],[f21460,f4399])).\n",
      "fof(f4399,plain,(\n",
      "  m(b,b) = f(two,m(a,a)) | ~spl0_53),\n",
      "  inference(avatar_component_clause,[],[f4397])).\n",
      "fof(f21460,plain,(\n",
      "  ( ! [X0] : (m(f(two,a),X0) = f(two,m(X0,a))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f19153,f331])).\n",
      "fof(f19153,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X1,X2) = f(X0,m(X2,m(X0,X1)))) )),\n",
      "  inference(superposition,[],[f18968,f151])).\n",
      "fof(f35130,plain,(\n",
      "  spl0_204 | ~spl0_20 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f35125,f1281,f329,f35127])).\n",
      "fof(f35127,plain,(\n",
      "  spl0_204 <=> d(m(b,f(two,a)),m(a,m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_204])])).\n",
      "fof(f1281,plain,(\n",
      "  spl0_36 <=> m(a,a) = m(b,m(b,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_36])])).\n",
      "fof(f35125,plain,(\n",
      "  d(m(b,f(two,a)),m(a,m(a,a))) | (~spl0_20 | ~spl0_36)),\n",
      "  inference(forward_demodulation,[],[f35082,f4])).\n",
      "fof(f35082,plain,(\n",
      "  d(m(f(two,a),b),m(a,m(a,a))) | (~spl0_20 | ~spl0_36)),\n",
      "  inference(superposition,[],[f26274,f331])).\n",
      "fof(f26274,plain,(\n",
      "  ( ! [X0,X1] : (d(m(X0,b),m(a,m(a,m(X1,X0))))) ) | ~spl0_36),\n",
      "  inference(forward_demodulation,[],[f26220,f3])).\n",
      "fof(f26220,plain,(\n",
      "  ( ! [X0,X1] : (d(m(X0,b),m(m(a,a),m(X1,X0)))) ) | ~spl0_36),\n",
      "  inference(superposition,[],[f3507,f885])).\n",
      "fof(f3507,plain,(\n",
      "  ( ! [X0,X1] : (d(m(X0,b),m(X0,m(X1,m(a,a))))) ) | ~spl0_36),\n",
      "  inference(superposition,[],[f1143,f1283])).\n",
      "fof(f1283,plain,(\n",
      "  m(a,a) = m(b,m(b,two)) | ~spl0_36),\n",
      "  inference(avatar_component_clause,[],[f1281])).\n",
      "fof(f33412,plain,(\n",
      "  spl0_3 | spl0_203 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f33084,f19,f33410,f42])).\n",
      "fof(f42,plain,(\n",
      "  spl0_3 <=> one = b),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).\n",
      "fof(f33410,plain,(\n",
      "  spl0_203 <=> ! [X0,X1] : m(X1,m(X0,m(a,a))) != m(X1,m(b,m(X0,two)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_203])])).\n",
      "fof(f19,plain,(\n",
      "  spl0_2 <=> m(a,a) = m(two,m(b,b))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).\n",
      "fof(f33084,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(X0,m(a,a))) != m(X1,m(b,m(X0,two))) | one = b) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f2464,f1292])).\n",
      "fof(f1292,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) = m(b,m(b,m(X0,two)))) ) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f1291,f151])).\n",
      "fof(f1291,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) = m(b,m(two,m(b,X0)))) ) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f1290,f4])).\n",
      "fof(f1290,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) = m(b,m(m(b,X0),two))) ) | ~spl0_2),\n",
      "  inference(forward_literal_rewriting,[],[f1289,f1001])).\n",
      "fof(f1001,plain,(\n",
      "  ( ! [X2,X3,X0,X1] : (m(X0,m(X1,X2)) != X3 | m(X2,m(X0,X1)) = X3) )),\n",
      "  inference(superposition,[],[f966,f3])).\n",
      "fof(f1289,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) = m(two,m(b,m(b,X0)))) ) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f1101,f3])).\n",
      "fof(f1101,plain,(\n",
      "  ( ! [X0] : (m(two,m(m(b,b),X0)) = m(X0,m(a,a))) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f151,f21])).\n",
      "fof(f21,plain,(\n",
      "  m(a,a) = m(two,m(b,b)) | ~spl0_2),\n",
      "  inference(avatar_component_clause,[],[f19])).\n",
      "fof(f33397,plain,(\n",
      "  spl0_3 | spl0_202 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f33393,f19,f33395,f42])).\n",
      "fof(f33395,plain,(\n",
      "  spl0_202 <=> ! [X0,X1] : m(X1,m(X0,m(a,a))) != m(two,m(X0,m(b,X1)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_202])])).\n",
      "fof(f33393,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(X0,m(a,a))) != m(two,m(X0,m(b,X1))) | one = b) ) | ~spl0_2),\n",
      "  inference(forward_literal_rewriting,[],[f33392,f2902])).\n",
      "fof(f33392,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(X0,m(a,a))) != m(m(X0,two),m(b,X1)) | one = b) ) | ~spl0_2),\n",
      "  inference(forward_literal_rewriting,[],[f33073,f2902])).\n",
      "fof(f33073,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(X0,m(a,a))) != m(m(b,m(X0,two)),X1) | one = b) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f1139,f1292])).\n",
      "fof(f32354,plain,(\n",
      "  spl0_1 | spl0_201 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f32031,f329,f32352,f14])).\n",
      "fof(f32352,plain,(\n",
      "  spl0_201 <=> ! [X0,X1] : m(X1,m(a,X0)) != m(X1,m(f(two,a),X0))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_201])])).\n",
      "fof(f32031,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(a,X0)) != m(X1,m(f(two,a),X0)) | one = two) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f2464,f7211])).\n",
      "fof(f7211,plain,(\n",
      "  ( ! [X0] : (m(a,X0) = m(two,m(f(two,a),X0))) ) | ~spl0_20),\n",
      "  inference(forward_literal_rewriting,[],[f882,f6853])).\n",
      "fof(f882,plain,(\n",
      "  ( ! [X0] : (m(a,X0) = m(f(two,a),m(two,X0))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f145,f331])).\n",
      "fof(f32339,plain,(\n",
      "  spl0_1 | spl0_200 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f32335,f329,f32337,f14])).\n",
      "fof(f32337,plain,(\n",
      "  spl0_200 <=> ! [X0,X1] : m(X1,m(a,X0)) != m(X0,m(f(two,a),X1))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_200])])).\n",
      "fof(f32335,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(a,X0)) != m(X0,m(f(two,a),X1)) | one = two) ) | ~spl0_20),\n",
      "  inference(forward_literal_rewriting,[],[f32020,f2902])).\n",
      "fof(f32020,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(a,X0)) != m(m(f(two,a),X0),X1) | one = two) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f1139,f7211])).\n",
      "fof(f31482,plain,(\n",
      "  ~spl0_67 | ~spl0_195),\n",
      "  inference(avatar_split_clause,[],[f31477,f30945,f5066])).\n",
      "fof(f5066,plain,(\n",
      "  spl0_67 <=> one = m(a,a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_67])])).\n",
      "fof(f31477,plain,(\n",
      "  one != m(a,a) | ~spl0_195),\n",
      "  inference(superposition,[],[f31434,f10286])).\n",
      "fof(f10286,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,X0) = m(X1,m(one,X0))) )),\n",
      "  inference(superposition,[],[f10249,f885])).\n",
      "fof(f31434,plain,(\n",
      "  ( ! [X0] : (m(a,m(X0,a)) != X0) ) | ~spl0_195),\n",
      "  inference(superposition,[],[f30946,f151])).\n",
      "fof(f31294,plain,(\n",
      "  spl0_197 | spl0_1 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f31293,f19,f14,f31080])).\n",
      "fof(f31080,plain,(\n",
      "  spl0_197 <=> ! [X0] : m(X0,m(a,a)) != m(X0,m(b,b))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_197])])).\n",
      "fof(f31293,plain,(\n",
      "  ( ! [X0] : (one = two | m(X0,m(a,a)) != m(X0,m(b,b))) ) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f30707,f18968])).\n",
      "fof(f30707,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) != m(X0,m(b,b)) | one = f(m(X0,m(b,b)),m(m(X0,m(b,b)),two))) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f9245,f1272])).\n",
      "fof(f1272,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) = m(two,m(X0,m(b,b)))) ) | ~spl0_2),\n",
      "  inference(forward_literal_rewriting,[],[f1083,f1000])).\n",
      "fof(f1000,plain,(\n",
      "  ( ! [X2,X3,X0,X1] : (m(X1,m(X0,X2)) != X3 | m(X2,m(X0,X1)) = X3) )),\n",
      "  inference(superposition,[],[f966,f145])).\n",
      "fof(f1083,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) = m(m(b,b),m(X0,two))) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f151,f21])).\n",
      "fof(f31260,plain,(\n",
      "  spl0_1 | spl0_199 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f30691,f19,f31258,f14])).\n",
      "fof(f31258,plain,(\n",
      "  spl0_199 <=> ! [X0,X1] : m(X1,m(X0,m(a,a))) != m(X1,m(X0,m(b,b)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_199])])).\n",
      "fof(f30691,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(X0,m(a,a))) != m(X1,m(X0,m(b,b))) | one = two) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f2464,f1272])).\n",
      "fof(f31255,plain,(\n",
      "  spl0_197 | spl0_1 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f31254,f19,f14,f31080])).\n",
      "fof(f31254,plain,(\n",
      "  ( ! [X0] : (one = two | m(X0,m(a,a)) != m(X0,m(b,b))) ) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f30688,f19151])).\n",
      "fof(f30688,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) != m(X0,m(b,b)) | one = f(m(X0,m(b,b)),m(two,m(X0,m(b,b))))) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f2253,f1272])).\n",
      "fof(f31241,plain,(\n",
      "  spl0_1 | spl0_198 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f31237,f19,f31239,f14])).\n",
      "fof(f31239,plain,(\n",
      "  spl0_198 <=> ! [X0,X1] : m(X1,m(X0,m(a,a))) != m(b,m(b,m(X0,X1)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_198])])).\n",
      "fof(f31237,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(X0,m(a,a))) != m(b,m(b,m(X0,X1))) | one = two) ) | ~spl0_2),\n",
      "  inference(forward_literal_rewriting,[],[f31236,f2902])).\n",
      "fof(f31236,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(X0,m(a,a))) != m(m(b,b),m(X0,X1)) | one = two) ) | ~spl0_2),\n",
      "  inference(forward_literal_rewriting,[],[f30680,f2902])).\n",
      "fof(f30680,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(X0,m(a,a))) != m(m(X0,m(b,b)),X1) | one = two) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f1139,f1272])).\n",
      "fof(f31082,plain,(\n",
      "  spl0_1 | spl0_197 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f30600,f19,f31080,f14])).\n",
      "fof(f30600,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) != m(X0,m(b,b)) | one = two) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f65,f1272])).\n",
      "fof(f30975,plain,(\n",
      "  spl0_195 | spl0_196 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f30974,f19,f30948,f30945])).\n",
      "fof(f30948,plain,(\n",
      "  spl0_196 <=> one = m(b,m(b,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_196])])).\n",
      "fof(f30974,plain,(\n",
      "  ( ! [X0] : (one = m(b,m(b,two)) | m(X0,m(a,a)) != X0) ) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f30973,f4])).\n",
      "fof(f30973,plain,(\n",
      "  ( ! [X0] : (one = m(b,m(two,b)) | m(X0,m(a,a)) != X0) ) | ~spl0_2),\n",
      "  inference(forward_literal_rewriting,[],[f30548,f6853])).\n",
      "fof(f30548,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) != X0 | one = m(two,m(b,b))) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f2461,f1272])).\n",
      "fof(f2461,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X1,m(X0,X2)) != X0 | one = m(X1,X2)) )),\n",
      "  inference(superposition,[],[f51,f883])).\n",
      "fof(f30951,plain,(\n",
      "  spl0_195 | spl0_196 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f30943,f19,f30948,f30945])).\n",
      "fof(f30943,plain,(\n",
      "  ( ! [X0] : (one = m(b,m(b,two)) | m(X0,m(a,a)) != X0) ) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f30942,f4])).\n",
      "fof(f30942,plain,(\n",
      "  ( ! [X0] : (one = m(b,m(two,b)) | m(X0,m(a,a)) != X0) ) | ~spl0_2),\n",
      "  inference(forward_literal_rewriting,[],[f30941,f6853])).\n",
      "fof(f30941,plain,(\n",
      "  ( ! [X0] : (one = m(two,m(b,b)) | m(X0,m(a,a)) != X0) ) | ~spl0_2),\n",
      "  inference(forward_literal_rewriting,[],[f30539,f966])).\n",
      "fof(f30539,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) != X0 | one = m(m(b,b),two)) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f1136,f1272])).\n",
      "fof(f1136,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X2,m(X0,X1)) != X0 | one = m(X1,X2)) )),\n",
      "  inference(superposition,[],[f51,f151])).\n",
      "fof(f30816,plain,(\n",
      "  spl0_194 | ~spl0_2 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f30811,f329,f19,f30813])).\n",
      "fof(f30813,plain,(\n",
      "  spl0_194 <=> m(b,m(b,a)) = m(a,m(a,f(two,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_194])])).\n",
      "fof(f30811,plain,(\n",
      "  m(b,m(b,a)) = m(a,m(a,f(two,a))) | (~spl0_2 | ~spl0_20)),\n",
      "  inference(forward_literal_rewriting,[],[f30810,f20757])).\n",
      "fof(f20757,plain,(\n",
      "  ( ! [X2,X3,X0,X1] : (m(m(X1,X0),X2) != X3 | m(X0,m(X1,X2)) = X3) )),\n",
      "  inference(forward_literal_rewriting,[],[f20756,f11307])).\n",
      "fof(f11307,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X0,X1) != X2 | m(m(X1,one),X0) = X2) )),\n",
      "  inference(superposition,[],[f966,f10282])).\n",
      "fof(f20756,plain,(\n",
      "  ( ! [X2,X3,X0,X1] : (m(X0,m(X1,X2)) = X3 | m(m(X2,one),m(X1,X0)) != X3) )),\n",
      "  inference(forward_demodulation,[],[f20215,f3])).\n",
      "fof(f20215,plain,(\n",
      "  ( ! [X2,X3,X0,X1] : (m(m(X2,one),m(X1,X0)) != X3 | m(m(X0,X1),X2) = X3) )),\n",
      "  inference(superposition,[],[f8256,f6443])).\n",
      "fof(f6443,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(m(X1,X2),X0) = m(X0,m(X2,X1))) )),\n",
      "  inference(superposition,[],[f2380,f4])).\n",
      "fof(f8256,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X0,m(X1,one)) != X2 | m(X0,X1) = X2) )),\n",
      "  inference(superposition,[],[f50,f151])).\n",
      "fof(f30810,plain,(\n",
      "  m(b,m(b,a)) = m(m(a,a),f(two,a)) | (~spl0_2 | ~spl0_20)),\n",
      "  inference(forward_literal_rewriting,[],[f30809,f966])).\n",
      "fof(f30809,plain,(\n",
      "  m(b,m(b,a)) = m(f(two,a),m(a,a)) | (~spl0_2 | ~spl0_20)),\n",
      "  inference(forward_demodulation,[],[f30808,f151])).\n",
      "fof(f30808,plain,(\n",
      "  m(b,m(a,b)) = m(f(two,a),m(a,a)) | (~spl0_2 | ~spl0_20)),\n",
      "  inference(forward_demodulation,[],[f30807,f151])).\n",
      "fof(f30807,plain,(\n",
      "  m(a,m(b,b)) = m(f(two,a),m(a,a)) | (~spl0_2 | ~spl0_20)),\n",
      "  inference(forward_demodulation,[],[f30454,f4])).\n",
      "fof(f30454,plain,(\n",
      "  m(m(b,b),a) = m(f(two,a),m(a,a)) | (~spl0_2 | ~spl0_20)),\n",
      "  inference(superposition,[],[f1272,f1102])).\n",
      "fof(f30299,plain,(\n",
      "  spl0_193 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f30260,f329,f30296])).\n",
      "fof(f30296,plain,(\n",
      "  spl0_193 <=> d(m(f(two,a),f(two,a)),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_193])])).\n",
      "fof(f30260,plain,(\n",
      "  d(m(f(two,a),f(two,a)),m(a,a)) | ~spl0_20),\n",
      "  inference(superposition,[],[f4721,f1102])).\n",
      "fof(f4721,plain,(\n",
      "  ( ! [X0,X1] : (d(m(f(two,a),X0),m(X1,m(X0,a)))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f1145,f331])).\n",
      "fof(f1145,plain,(\n",
      "  ( ! [X2,X3,X0,X1] : (d(m(X1,X2),m(X3,m(X2,m(X0,X1))))) )),\n",
      "  inference(superposition,[],[f160,f151])).\n",
      "fof(f29824,plain,(\n",
      "  spl0_1 | spl0_192 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f29549,f329,f29822,f14])).\n",
      "fof(f29822,plain,(\n",
      "  spl0_192 <=> ! [X0,X1] : m(X1,m(a,X0)) != m(X1,m(X0,f(two,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_192])])).\n",
      "fof(f29549,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(a,X0)) != m(X1,m(X0,f(two,a))) | one = two) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f2464,f3367])).\n",
      "fof(f3367,plain,(\n",
      "  ( ! [X0] : (m(a,X0) = m(two,m(X0,f(two,a)))) ) | ~spl0_20),\n",
      "  inference(forward_literal_rewriting,[],[f3203,f1000])).\n",
      "fof(f3203,plain,(\n",
      "  ( ! [X0] : (m(a,X0) = m(f(two,a),m(X0,two))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f1086,f331])).\n",
      "fof(f1086,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(X2,m(X0,X1)) = m(m(X1,X2),X0)) )),\n",
      "  inference(superposition,[],[f151,f4])).\n",
      "fof(f29812,plain,(\n",
      "  spl0_1 | spl0_191 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f29808,f329,f29810,f14])).\n",
      "fof(f29810,plain,(\n",
      "  spl0_191 <=> ! [X0,X1] : m(X1,m(a,X0)) != m(f(two,a),m(X0,X1))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_191])])).\n",
      "fof(f29808,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(a,X0)) != m(f(two,a),m(X0,X1)) | one = two) ) | ~spl0_20),\n",
      "  inference(forward_literal_rewriting,[],[f29538,f2902])).\n",
      "fof(f29538,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(a,X0)) != m(m(X0,f(two,a)),X1) | one = two) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f1139,f3367])).\n",
      "fof(f28822,plain,(\n",
      "  spl0_1 | spl0_190 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f28555,f329,f28820,f14])).\n",
      "fof(f28820,plain,(\n",
      "  spl0_190 <=> ! [X0,X1] : m(X1,m(X0,a)) != m(X1,m(X0,f(two,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_190])])).\n",
      "fof(f28555,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(X0,a)) != m(X1,m(X0,f(two,a))) | one = two) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f2464,f1273])).\n",
      "fof(f1273,plain,(\n",
      "  ( ! [X0] : (m(X0,a) = m(two,m(X0,f(two,a)))) ) | ~spl0_20),\n",
      "  inference(forward_literal_rewriting,[],[f1084,f1000])).\n",
      "fof(f1084,plain,(\n",
      "  ( ! [X0] : (m(X0,a) = m(f(two,a),m(X0,two))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f151,f331])).\n",
      "fof(f28811,plain,(\n",
      "  spl0_1 | spl0_189 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f28807,f329,f28809,f14])).\n",
      "fof(f28809,plain,(\n",
      "  spl0_189 <=> ! [X0,X1] : m(X1,m(X0,a)) != m(f(two,a),m(X0,X1))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_189])])).\n",
      "fof(f28807,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(X0,a)) != m(f(two,a),m(X0,X1)) | one = two) ) | ~spl0_20),\n",
      "  inference(forward_literal_rewriting,[],[f28545,f2902])).\n",
      "fof(f28545,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(X0,a)) != m(m(X0,f(two,a)),X1) | one = two) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f1139,f1273])).\n",
      "fof(f28165,plain,(\n",
      "  spl0_1 | spl0_188 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f27818,f329,f28163,f14])).\n",
      "fof(f28163,plain,(\n",
      "  spl0_188 <=> ! [X0,X1] : m(X1,m(X0,a)) != m(X1,m(f(two,a),X0))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_188])])).\n",
      "fof(f27818,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(X0,a)) != m(X1,m(f(two,a),X0)) | one = two) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f2464,f1102])).\n",
      "fof(f28151,plain,(\n",
      "  spl0_1 | spl0_187 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f28147,f329,f28149,f14])).\n",
      "fof(f28149,plain,(\n",
      "  spl0_187 <=> ! [X0,X1] : m(X1,m(X0,a)) != m(X0,m(f(two,a),X1))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_187])])).\n",
      "fof(f28147,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(X0,a)) != m(X0,m(f(two,a),X1)) | one = two) ) | ~spl0_20),\n",
      "  inference(forward_literal_rewriting,[],[f27808,f2902])).\n",
      "fof(f27808,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(X0,a)) != m(m(f(two,a),X0),X1) | one = two) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f1139,f1102])).\n",
      "fof(f26173,plain,(\n",
      "  spl0_3 | spl0_186 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f26137,f1281,f26171,f42])).\n",
      "fof(f26171,plain,(\n",
      "  spl0_186 <=> ! [X0] : m(X0,m(a,a)) != m(X0,m(b,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_186])])).\n",
      "fof(f26137,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) != m(X0,m(b,two)) | one = b) ) | ~spl0_36),\n",
      "  inference(superposition,[],[f2464,f1283])).\n",
      "fof(f26169,plain,(\n",
      "  spl0_3 | spl0_185 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f26136,f19,f26167,f42])).\n",
      "fof(f26167,plain,(\n",
      "  spl0_185 <=> ! [X0,X1] : m(X1,m(a,m(a,X0))) != m(X1,m(b,m(X0,two)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_185])])).\n",
      "fof(f26136,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(a,m(a,X0))) != m(X1,m(b,m(X0,two))) | one = b) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f2464,f7203])).\n",
      "fof(f7203,plain,(\n",
      "  ( ! [X0] : (m(a,m(a,X0)) = m(b,m(b,m(X0,two)))) ) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f7202,f151])).\n",
      "fof(f7202,plain,(\n",
      "  ( ! [X0] : (m(a,m(a,X0)) = m(b,m(two,m(b,X0)))) ) | ~spl0_2),\n",
      "  inference(forward_literal_rewriting,[],[f166,f6853])).\n",
      "fof(f166,plain,(\n",
      "  ( ! [X0] : (m(a,m(a,X0)) = m(two,m(b,m(b,X0)))) ) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f165,f3])).\n",
      "fof(f165,plain,(\n",
      "  ( ! [X0] : (m(two,m(m(b,b),X0)) = m(a,m(a,X0))) ) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f149,f3])).\n",
      "fof(f149,plain,(\n",
      "  ( ! [X0] : (m(two,m(m(b,b),X0)) = m(m(a,a),X0)) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f3,f21])).\n",
      "fof(f25718,plain,(\n",
      "  spl0_184 | ~spl0_20 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f25713,f1281,f329,f25715])).\n",
      "fof(f25715,plain,(\n",
      "  spl0_184 <=> d(m(b,a),m(a,m(a,f(two,a))))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_184])])).\n",
      "fof(f25713,plain,(\n",
      "  d(m(b,a),m(a,m(a,f(two,a)))) | (~spl0_20 | ~spl0_36)),\n",
      "  inference(forward_demodulation,[],[f25712,f3])).\n",
      "fof(f25712,plain,(\n",
      "  d(m(b,a),m(m(a,a),f(two,a))) | (~spl0_20 | ~spl0_36)),\n",
      "  inference(forward_demodulation,[],[f25671,f4])).\n",
      "fof(f25671,plain,(\n",
      "  d(m(b,a),m(f(two,a),m(a,a))) | (~spl0_20 | ~spl0_36)),\n",
      "  inference(superposition,[],[f1429,f331])).\n",
      "fof(f1429,plain,(\n",
      "  ( ! [X0] : (d(m(b,m(two,X0)),m(X0,m(a,a)))) ) | ~spl0_36),\n",
      "  inference(forward_demodulation,[],[f1386,f3])).\n",
      "fof(f1386,plain,(\n",
      "  ( ! [X0] : (d(m(m(b,two),X0),m(X0,m(a,a)))) ) | ~spl0_36),\n",
      "  inference(superposition,[],[f253,f1283])).\n",
      "fof(f253,plain,(\n",
      "  ( ! [X2,X0,X1] : (d(m(X0,X1),m(X1,m(X2,X0)))) )),\n",
      "  inference(forward_demodulation,[],[f242,f3])).\n",
      "fof(f242,plain,(\n",
      "  ( ! [X2,X0,X1] : (d(m(X0,X1),m(m(X1,X2),X0))) )),\n",
      "  inference(superposition,[],[f159,f4])).\n",
      "fof(f25538,plain,(\n",
      "  ~spl0_183 | ~spl0_20 | ~spl0_180),\n",
      "  inference(avatar_split_clause,[],[f25520,f25450,f329,f25535])).\n",
      "fof(f25535,plain,(\n",
      "  spl0_183 <=> a = m(two,a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_183])])).\n",
      "fof(f25450,plain,(\n",
      "  spl0_180 <=> ! [X0] : m(X0,a) != m(f(two,a),X0)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_180])])).\n",
      "fof(f25520,plain,(\n",
      "  a != m(two,a) | (~spl0_20 | ~spl0_180)),\n",
      "  inference(superposition,[],[f25475,f331])).\n",
      "fof(f25475,plain,(\n",
      "  ( ! [X0] : (m(X0,a) != m(X0,f(two,a))) ) | ~spl0_180),\n",
      "  inference(superposition,[],[f25451,f4])).\n",
      "fof(f25451,plain,(\n",
      "  ( ! [X0] : (m(X0,a) != m(f(two,a),X0)) ) | ~spl0_180),\n",
      "  inference(avatar_component_clause,[],[f25450])).\n",
      "fof(f25464,plain,(\n",
      "  spl0_3 | spl0_182 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f25460,f1281,f25462,f42])).\n",
      "fof(f25462,plain,(\n",
      "  spl0_182 <=> ! [X0] : m(X0,m(a,a)) != m(b,m(two,X0))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_182])])).\n",
      "fof(f25460,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) != m(b,m(two,X0)) | one = b) ) | ~spl0_36),\n",
      "  inference(forward_literal_rewriting,[],[f25459,f6853])).\n",
      "fof(f25459,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) != m(two,m(b,X0)) | one = b) ) | ~spl0_36),\n",
      "  inference(forward_literal_rewriting,[],[f25408,f2902])).\n",
      "fof(f25408,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) != m(m(b,two),X0) | one = b) ) | ~spl0_36),\n",
      "  inference(superposition,[],[f1139,f1283])).\n",
      "fof(f25458,plain,(\n",
      "  spl0_3 | spl0_181 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f25454,f19,f25456,f42])).\n",
      "fof(f25456,plain,(\n",
      "  spl0_181 <=> ! [X0,X1] : m(X1,m(a,m(a,X0))) != m(two,m(X0,m(b,X1)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_181])])).\n",
      "fof(f25454,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(a,m(a,X0))) != m(two,m(X0,m(b,X1))) | one = b) ) | ~spl0_2),\n",
      "  inference(forward_literal_rewriting,[],[f25453,f2902])).\n",
      "fof(f25453,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(a,m(a,X0))) != m(m(X0,two),m(b,X1)) | one = b) ) | ~spl0_2),\n",
      "  inference(forward_literal_rewriting,[],[f25407,f2902])).\n",
      "fof(f25407,plain,(\n",
      "  ( ! [X0,X1] : (m(X1,m(a,m(a,X0))) != m(m(b,m(X0,two)),X1) | one = b) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f1139,f7203])).\n",
      "fof(f25452,plain,(\n",
      "  spl0_1 | spl0_180 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f25406,f329,f25450,f14])).\n",
      "fof(f25406,plain,(\n",
      "  ( ! [X0] : (m(X0,a) != m(f(two,a),X0) | one = two) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f1139,f331])).\n",
      "fof(f22729,plain,(\n",
      "  spl0_179 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f22643,f1281,f22726])).\n",
      "fof(f22726,plain,(\n",
      "  spl0_179 <=> two = f(m(b,b),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_179])])).\n",
      "fof(f22643,plain,(\n",
      "  two = f(m(b,b),m(a,a)) | ~spl0_36),\n",
      "  inference(superposition,[],[f19167,f1283])).\n",
      "fof(f19167,plain,(\n",
      "  ( ! [X2,X0,X1] : (f(m(X0,X1),m(X1,m(X0,X2))) = X2) )),\n",
      "  inference(superposition,[],[f18968,f145])).\n",
      "fof(f19668,plain,(\n",
      "  spl0_178 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f19643,f329,f19665])).\n",
      "fof(f19665,plain,(\n",
      "  spl0_178 <=> two = f(m(one,f(two,a)),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_178])])).\n",
      "fof(f19643,plain,(\n",
      "  two = f(m(one,f(two,a)),a) | ~spl0_20),\n",
      "  inference(superposition,[],[f19451,f331])).\n",
      "fof(f19451,plain,(\n",
      "  ( ! [X0,X1] : (f(m(one,X1),m(X0,X1)) = X0) )),\n",
      "  inference(superposition,[],[f19151,f10286])).\n",
      "fof(f19599,plain,(\n",
      "  spl0_177 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f19594,f1281,f19596])).\n",
      "fof(f19596,plain,(\n",
      "  spl0_177 <=> m(b,two) = f(m(b,one),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_177])])).\n",
      "fof(f19594,plain,(\n",
      "  m(b,two) = f(m(b,one),m(a,a)) | ~spl0_36),\n",
      "  inference(forward_demodulation,[],[f19554,f4])).\n",
      "fof(f19554,plain,(\n",
      "  m(b,two) = f(m(one,b),m(a,a)) | ~spl0_36),\n",
      "  inference(superposition,[],[f19169,f1283])).\n",
      "fof(f19591,plain,(\n",
      "  spl0_176 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f19586,f329,f19588])).\n",
      "fof(f19588,plain,(\n",
      "  spl0_176 <=> f(two,a) = f(m(two,one),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_176])])).\n",
      "fof(f19586,plain,(\n",
      "  f(two,a) = f(m(two,one),a) | ~spl0_20),\n",
      "  inference(forward_demodulation,[],[f19552,f4])).\n",
      "fof(f19552,plain,(\n",
      "  f(two,a) = f(m(one,two),a) | ~spl0_20),\n",
      "  inference(superposition,[],[f19169,f331])).\n",
      "fof(f19498,plain,(\n",
      "  spl0_175 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f19462,f1281,f19495])).\n",
      "fof(f19495,plain,(\n",
      "  spl0_175 <=> b = f(m(b,two),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_175])])).\n",
      "fof(f19462,plain,(\n",
      "  b = f(m(b,two),m(a,a)) | ~spl0_36),\n",
      "  inference(superposition,[],[f19151,f1283])).\n",
      "fof(f19489,plain,(\n",
      "  spl0_174 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f19460,f329,f19486])).\n",
      "fof(f19486,plain,(\n",
      "  spl0_174 <=> two = f(f(two,a),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_174])])).\n",
      "fof(f19460,plain,(\n",
      "  two = f(f(two,a),a) | ~spl0_20),\n",
      "  inference(superposition,[],[f19151,f331])).\n",
      "fof(f19429,plain,(\n",
      "  spl0_173 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f19173,f1281,f19426])).\n",
      "fof(f19426,plain,(\n",
      "  spl0_173 <=> m(b,two) = f(b,m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_173])])).\n",
      "fof(f19173,plain,(\n",
      "  m(b,two) = f(b,m(a,a)) | ~spl0_36),\n",
      "  inference(superposition,[],[f18968,f1283])).\n",
      "fof(f15297,plain,(\n",
      "  spl0_172 | ~spl0_170),\n",
      "  inference(avatar_split_clause,[],[f15292,f15230,f15294])).\n",
      "fof(f15294,plain,(\n",
      "  spl0_172 <=> a = m(f(a,m(a,one)),f(f(a,m(a,one)),a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_172])])).\n",
      "fof(f15230,plain,(\n",
      "  spl0_170 <=> d(f(a,m(a,one)),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_170])])).\n",
      "fof(f15292,plain,(\n",
      "  a = m(f(a,m(a,one)),f(f(a,m(a,one)),a)) | ~spl0_170),\n",
      "  inference(resolution,[],[f15232,f6])).\n",
      "fof(f15232,plain,(\n",
      "  d(f(a,m(a,one)),a) | ~spl0_170),\n",
      "  inference(avatar_component_clause,[],[f15230])).\n",
      "fof(f15291,plain,(\n",
      "  ~spl0_163 | ~spl0_162 | spl0_164),\n",
      "  inference(avatar_split_clause,[],[f15290,f15180,f15171,f15176])).\n",
      "fof(f15171,plain,(\n",
      "  spl0_162 <=> one = f(m(a,one),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_162])])).\n",
      "fof(f15180,plain,(\n",
      "  spl0_164 <=> a = f(m(a,one),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_164])])).\n",
      "fof(f15290,plain,(\n",
      "  one != a | (~spl0_162 | spl0_164)),\n",
      "  inference(superposition,[],[f15182,f15173])).\n",
      "fof(f15173,plain,(\n",
      "  one = f(m(a,one),a) | ~spl0_162),\n",
      "  inference(avatar_component_clause,[],[f15171])).\n",
      "fof(f15182,plain,(\n",
      "  a != f(m(a,one),a) | spl0_164),\n",
      "  inference(avatar_component_clause,[],[f15180])).\n",
      "fof(f15242,plain,(\n",
      "  spl0_171 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15237,f12747,f15239])).\n",
      "fof(f15239,plain,(\n",
      "  spl0_171 <=> d(f(f(m(a,one),a),m(one,f(m(a,one),a))),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_171])])).\n",
      "fof(f12747,plain,(\n",
      "  spl0_152 <=> a = m(a,f(m(a,one),a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_152])])).\n",
      "fof(f15237,plain,(\n",
      "  d(f(f(m(a,one),a),m(one,f(m(a,one),a))),a) | ~spl0_152),\n",
      "  inference(forward_demodulation,[],[f15119,f4])).\n",
      "fof(f15119,plain,(\n",
      "  d(f(f(m(a,one),a),m(f(m(a,one),a),one)),a) | ~spl0_152),\n",
      "  inference(superposition,[],[f11234,f12749])).\n",
      "fof(f12749,plain,(\n",
      "  a = m(a,f(m(a,one),a)) | ~spl0_152),\n",
      "  inference(avatar_component_clause,[],[f12747])).\n",
      "fof(f11234,plain,(\n",
      "  ( ! [X0,X1] : (d(f(X1,m(X1,one)),m(X0,X1))) )),\n",
      "  inference(superposition,[],[f650,f10282])).\n",
      "fof(f650,plain,(\n",
      "  ( ! [X2,X0,X1] : (d(f(X0,m(X0,X1)),m(X2,m(X0,X1)))) )),\n",
      "  inference(superposition,[],[f160,f67])).\n",
      "fof(f15233,plain,(\n",
      "  spl0_170 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15228,f12747,f15230])).\n",
      "fof(f15228,plain,(\n",
      "  d(f(a,m(a,one)),a) | ~spl0_152),\n",
      "  inference(forward_demodulation,[],[f15111,f4])).\n",
      "fof(f15111,plain,(\n",
      "  d(f(a,m(one,a)),a) | ~spl0_152),\n",
      "  inference(superposition,[],[f10323,f12749])).\n",
      "fof(f10323,plain,(\n",
      "  ( ! [X0,X1] : (d(f(X1,m(one,X1)),m(X1,X0))) )),\n",
      "  inference(superposition,[],[f2266,f10249])).\n",
      "fof(f2266,plain,(\n",
      "  ( ! [X2,X0,X1] : (d(f(X0,m(X1,X0)),m(X2,m(X1,X0)))) )),\n",
      "  inference(superposition,[],[f160,f68])).\n",
      "fof(f15227,plain,(\n",
      "  spl0_169 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15222,f12747,f15224])).\n",
      "fof(f15224,plain,(\n",
      "  spl0_169 <=> d(m(one,f(m(a,one),a)),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_169])])).\n",
      "fof(f15222,plain,(\n",
      "  d(m(one,f(m(a,one),a)),a) | ~spl0_152),\n",
      "  inference(forward_demodulation,[],[f15108,f4])).\n",
      "fof(f15108,plain,(\n",
      "  d(m(f(m(a,one),a),one),a) | ~spl0_152),\n",
      "  inference(superposition,[],[f10300,f12749])).\n",
      "fof(f15216,plain,(\n",
      "  spl0_162 | spl0_168 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15089,f12747,f15214,f15171])).\n",
      "fof(f15214,plain,(\n",
      "  spl0_168 <=> ! [X0] : m(a,X0) != m(X0,a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_168])])).\n",
      "fof(f15089,plain,(\n",
      "  ( ! [X0] : (m(a,X0) != m(X0,a) | one = f(m(a,one),a)) ) | ~spl0_152),\n",
      "  inference(superposition,[],[f897,f12749])).\n",
      "fof(f15210,plain,(\n",
      "  spl0_167 | ~spl0_164 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15084,f12747,f15180,f15207])).\n",
      "fof(f15207,plain,(\n",
      "  spl0_167 <=> a = f(f(m(a,one),a),f(m(a,one),a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_167])])).\n",
      "fof(f15084,plain,(\n",
      "  a != f(m(a,one),a) | a = f(f(m(a,one),a),f(m(a,one),a)) | ~spl0_152),\n",
      "  inference(superposition,[],[f833,f12749])).\n",
      "fof(f15205,plain,(\n",
      "  spl0_166 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15083,f12747,f15202])).\n",
      "fof(f15202,plain,(\n",
      "  spl0_166 <=> d(f(f(m(a,one),a),a),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_166])])).\n",
      "fof(f15083,plain,(\n",
      "  d(f(f(m(a,one),a),a),a) | ~spl0_152),\n",
      "  inference(superposition,[],[f691,f12749])).\n",
      "fof(f691,plain,(\n",
      "  ( ! [X0,X1] : (d(f(X0,m(X1,X0)),m(X1,X0))) )),\n",
      "  inference(superposition,[],[f643,f4])).\n",
      "fof(f643,plain,(\n",
      "  ( ! [X0,X1] : (d(f(X0,m(X0,X1)),m(X0,X1))) )),\n",
      "  inference(superposition,[],[f31,f67])).\n",
      "fof(f15190,plain,(\n",
      "  spl0_165 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15047,f12747,f15187])).\n",
      "fof(f15187,plain,(\n",
      "  spl0_165 <=> a = m(f(m(a,one),a),f(f(m(a,one),a),a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_165])])).\n",
      "fof(f15047,plain,(\n",
      "  a = m(f(m(a,one),a),f(f(m(a,one),a),a)) | ~spl0_152),\n",
      "  inference(superposition,[],[f68,f12749])).\n",
      "fof(f15183,plain,(\n",
      "  spl0_163 | ~spl0_164 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15045,f12747,f15180,f15176])).\n",
      "fof(f15045,plain,(\n",
      "  a != f(m(a,one),a) | one = a | ~spl0_152),\n",
      "  inference(superposition,[],[f65,f12749])).\n",
      "fof(f15174,plain,(\n",
      "  spl0_162 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15127,f12747,f15171])).\n",
      "fof(f15127,plain,(\n",
      "  one = f(m(a,one),a) | ~spl0_152),\n",
      "  inference(trivial_inequality_removal,[],[f15043])).\n",
      "fof(f15043,plain,(\n",
      "  a != a | one = f(m(a,one),a) | ~spl0_152),\n",
      "  inference(superposition,[],[f51,f12749])).\n",
      "fof(f15169,plain,(\n",
      "  spl0_161 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15042,f12747,f15166])).\n",
      "fof(f15166,plain,(\n",
      "  spl0_161 <=> d(f(m(a,one),a),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_161])])).\n",
      "fof(f15042,plain,(\n",
      "  d(f(m(a,one),a),a) | ~spl0_152),\n",
      "  inference(superposition,[],[f31,f12749])).\n",
      "fof(f15164,plain,(\n",
      "  ~spl0_160 | ~spl0_137 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15034,f12747,f9515,f15160])).\n",
      "fof(f15160,plain,(\n",
      "  spl0_160 <=> m(a,a) = m(b,f(m(a,one),a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_160])])).\n",
      "fof(f15034,plain,(\n",
      "  m(a,a) != m(b,f(m(a,one),a)) | (~spl0_137 | ~spl0_152)),\n",
      "  inference(superposition,[],[f9521,f12749])).\n",
      "fof(f15163,plain,(\n",
      "  ~spl0_160 | ~spl0_66 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15158,f12747,f5018,f15160])).\n",
      "fof(f15158,plain,(\n",
      "  m(a,a) != m(b,f(m(a,one),a)) | (~spl0_66 | ~spl0_152)),\n",
      "  inference(forward_literal_rewriting,[],[f15033,f966])).\n",
      "fof(f15033,plain,(\n",
      "  m(a,a) != m(f(m(a,one),a),b) | (~spl0_66 | ~spl0_152)),\n",
      "  inference(superposition,[],[f9014,f12749])).\n",
      "fof(f15157,plain,(\n",
      "  ~spl0_159 | ~spl0_64 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15152,f12747,f5010,f15154])).\n",
      "fof(f15154,plain,(\n",
      "  spl0_159 <=> m(a,a) = m(two,f(m(a,one),a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_159])])).\n",
      "fof(f15152,plain,(\n",
      "  m(a,a) != m(two,f(m(a,one),a)) | (~spl0_64 | ~spl0_152)),\n",
      "  inference(forward_literal_rewriting,[],[f15032,f966])).\n",
      "fof(f15032,plain,(\n",
      "  m(a,a) != m(f(m(a,one),a),two) | (~spl0_64 | ~spl0_152)),\n",
      "  inference(superposition,[],[f8955,f12749])).\n",
      "fof(f15151,plain,(\n",
      "  ~spl0_158 | ~spl0_65 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15031,f12747,f5014,f15148])).\n",
      "fof(f15148,plain,(\n",
      "  spl0_158 <=> a = m(two,f(m(a,one),a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_158])])).\n",
      "fof(f15031,plain,(\n",
      "  a != m(two,f(m(a,one),a)) | (~spl0_65 | ~spl0_152)),\n",
      "  inference(superposition,[],[f8370,f12749])).\n",
      "fof(f15145,plain,(\n",
      "  spl0_157 | ~spl0_36 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15140,f12747,f1281,f15142])).\n",
      "fof(f15142,plain,(\n",
      "  spl0_157 <=> d(m(b,f(m(a,one),a)),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_157])])).\n",
      "fof(f15140,plain,(\n",
      "  d(m(b,f(m(a,one),a)),m(a,a)) | (~spl0_36 | ~spl0_152)),\n",
      "  inference(forward_demodulation,[],[f15022,f4])).\n",
      "fof(f15022,plain,(\n",
      "  d(m(f(m(a,one),a),b),m(a,a)) | (~spl0_36 | ~spl0_152)),\n",
      "  inference(superposition,[],[f1460,f12749])).\n",
      "fof(f1460,plain,(\n",
      "  ( ! [X0] : (d(m(X0,b),m(a,m(a,X0)))) ) | ~spl0_36),\n",
      "  inference(superposition,[],[f1375,f151])).\n",
      "fof(f1375,plain,(\n",
      "  ( ! [X0] : (d(m(X0,b),m(X0,m(a,a)))) ) | ~spl0_36),\n",
      "  inference(superposition,[],[f159,f1283])).\n",
      "fof(f15139,plain,(\n",
      "  spl0_156 | ~spl0_20 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15134,f12747,f329,f15136])).\n",
      "fof(f15136,plain,(\n",
      "  spl0_156 <=> d(m(f(two,a),f(m(a,one),a)),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_156])])).\n",
      "fof(f15134,plain,(\n",
      "  d(m(f(two,a),f(m(a,one),a)),a) | (~spl0_20 | ~spl0_152)),\n",
      "  inference(forward_demodulation,[],[f15018,f4])).\n",
      "fof(f15018,plain,(\n",
      "  d(m(f(m(a,one),a),f(two,a)),a) | (~spl0_20 | ~spl0_152)),\n",
      "  inference(superposition,[],[f1040,f12749])).\n",
      "fof(f15133,plain,(\n",
      "  spl0_155 | ~spl0_20 | ~spl0_152),\n",
      "  inference(avatar_split_clause,[],[f15128,f12747,f329,f15130])).\n",
      "fof(f15130,plain,(\n",
      "  spl0_155 <=> d(m(two,f(m(a,one),a)),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_155])])).\n",
      "fof(f15128,plain,(\n",
      "  d(m(two,f(m(a,one),a)),a) | (~spl0_20 | ~spl0_152)),\n",
      "  inference(forward_demodulation,[],[f15012,f4])).\n",
      "fof(f15012,plain,(\n",
      "  d(m(f(m(a,one),a),two),a) | (~spl0_20 | ~spl0_152)),\n",
      "  inference(superposition,[],[f431,f12749])).\n",
      "fof(f14316,plain,(\n",
      "  spl0_154 | ~spl0_147),\n",
      "  inference(avatar_split_clause,[],[f14311,f12296,f14313])).\n",
      "fof(f14313,plain,(\n",
      "  spl0_154 <=> m(a,a) = m(f(b,m(b,one)),f(f(b,m(b,one)),m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_154])])).\n",
      "fof(f12296,plain,(\n",
      "  spl0_147 <=> d(f(b,m(b,one)),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_147])])).\n",
      "fof(f14311,plain,(\n",
      "  m(a,a) = m(f(b,m(b,one)),f(f(b,m(b,one)),m(a,a))) | ~spl0_147),\n",
      "  inference(resolution,[],[f12298,f6])).\n",
      "fof(f12298,plain,(\n",
      "  d(f(b,m(b,one)),m(a,a)) | ~spl0_147),\n",
      "  inference(avatar_component_clause,[],[f12296])).\n",
      "fof(f14032,plain,(\n",
      "  spl0_153 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f14027,f329,f14029])).\n",
      "fof(f14029,plain,(\n",
      "  spl0_153 <=> d(f(f(two,a),m(one,f(two,a))),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_153])])).\n",
      "fof(f14027,plain,(\n",
      "  d(f(f(two,a),m(one,f(two,a))),a) | ~spl0_20),\n",
      "  inference(forward_demodulation,[],[f13998,f4])).\n",
      "fof(f13998,plain,(\n",
      "  d(f(f(two,a),m(f(two,a),one)),a) | ~spl0_20),\n",
      "  inference(superposition,[],[f11234,f331])).\n",
      "fof(f12750,plain,(\n",
      "  spl0_152 | ~spl0_151),\n",
      "  inference(avatar_split_clause,[],[f12745,f12557,f12747])).\n",
      "fof(f12557,plain,(\n",
      "  spl0_151 <=> d(m(a,one),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_151])])).\n",
      "fof(f12745,plain,(\n",
      "  a = m(a,f(m(a,one),a)) | ~spl0_151),\n",
      "  inference(forward_literal_rewriting,[],[f12744,f966])).\n",
      "fof(f12744,plain,(\n",
      "  a = m(f(m(a,one),a),a) | ~spl0_151),\n",
      "  inference(forward_literal_rewriting,[],[f12743,f11177])).\n",
      "fof(f11177,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(m(X1,one),X0) != X2 | m(X0,X1) = X2) )),\n",
      "  inference(superposition,[],[f8256,f4])).\n",
      "fof(f12743,plain,(\n",
      "  a = m(m(a,one),f(m(a,one),a)) | ~spl0_151),\n",
      "  inference(resolution,[],[f12559,f6])).\n",
      "fof(f12559,plain,(\n",
      "  d(m(a,one),a) | ~spl0_151),\n",
      "  inference(avatar_component_clause,[],[f12557])).\n",
      "fof(f12560,plain,(\n",
      "  spl0_151 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f12555,f329,f12557])).\n",
      "fof(f12555,plain,(\n",
      "  d(m(a,one),a) | ~spl0_20),\n",
      "  inference(forward_demodulation,[],[f12545,f4])).\n",
      "fof(f12545,plain,(\n",
      "  d(m(one,a),a) | ~spl0_20),\n",
      "  inference(superposition,[],[f10519,f331])).\n",
      "fof(f10519,plain,(\n",
      "  ( ! [X0,X1] : (d(m(one,m(X1,X0)),m(X1,X0))) )),\n",
      "  inference(forward_demodulation,[],[f10373,f3])).\n",
      "fof(f10373,plain,(\n",
      "  ( ! [X0,X1] : (d(m(m(one,X1),X0),m(X1,X0))) )),\n",
      "  inference(superposition,[],[f276,f10249])).\n",
      "fof(f276,plain,(\n",
      "  ( ! [X0,X1] : (d(m(X0,X1),m(X1,X0))) )),\n",
      "  inference(superposition,[],[f231,f2])).\n",
      "fof(f231,plain,(\n",
      "  ( ! [X2,X0,X1] : (d(m(X1,X0),m(X0,m(X1,X2)))) )),\n",
      "  inference(superposition,[],[f159,f4])).\n",
      "fof(f12436,plain,(\n",
      "  spl0_150 | ~spl0_149),\n",
      "  inference(avatar_split_clause,[],[f12431,f12425,f12433])).\n",
      "fof(f12433,plain,(\n",
      "  spl0_150 <=> a = f(m(one,one),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_150])])).\n",
      "fof(f12425,plain,(\n",
      "  spl0_149 <=> d(m(one,one),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_149])])).\n",
      "fof(f12431,plain,(\n",
      "  a = f(m(one,one),a) | ~spl0_149),\n",
      "  inference(forward_literal_rewriting,[],[f12430,f50])).\n",
      "fof(f12430,plain,(\n",
      "  a = m(one,f(m(one,one),a)) | ~spl0_149),\n",
      "  inference(forward_literal_rewriting,[],[f12429,f10243])).\n",
      "fof(f10243,plain,(\n",
      "  ( ! [X2,X0,X1] : (m(m(one,X1),X0) != X2 | m(X1,X0) = X2) )),\n",
      "  inference(superposition,[],[f8255,f4])).\n",
      "fof(f12429,plain,(\n",
      "  a = m(m(one,one),f(m(one,one),a)) | ~spl0_149),\n",
      "  inference(resolution,[],[f12427,f6])).\n",
      "fof(f12427,plain,(\n",
      "  d(m(one,one),a) | ~spl0_149),\n",
      "  inference(avatar_component_clause,[],[f12425])).\n",
      "fof(f12428,plain,(\n",
      "  spl0_149 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f12419,f329,f12425])).\n",
      "fof(f12419,plain,(\n",
      "  d(m(one,one),a) | ~spl0_20),\n",
      "  inference(superposition,[],[f12357,f331])).\n",
      "fof(f12357,plain,(\n",
      "  ( ! [X0,X1] : (d(m(one,one),m(X0,X1))) )),\n",
      "  inference(superposition,[],[f10361,f10286])).\n",
      "fof(f10361,plain,(\n",
      "  ( ! [X2,X0,X1] : (d(m(one,X1),m(X2,m(X1,X0)))) )),\n",
      "  inference(superposition,[],[f160,f10249])).\n",
      "fof(f12305,plain,(\n",
      "  spl0_148 | ~spl0_146),\n",
      "  inference(avatar_split_clause,[],[f12300,f12289,f12302])).\n",
      "fof(f12302,plain,(\n",
      "  spl0_148 <=> a = m(f(two,m(two,one)),f(f(two,m(two,one)),a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_148])])).\n",
      "fof(f12289,plain,(\n",
      "  spl0_146 <=> d(f(two,m(two,one)),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_146])])).\n",
      "fof(f12300,plain,(\n",
      "  a = m(f(two,m(two,one)),f(f(two,m(two,one)),a)) | ~spl0_146),\n",
      "  inference(resolution,[],[f12291,f6])).\n",
      "fof(f12291,plain,(\n",
      "  d(f(two,m(two,one)),a) | ~spl0_146),\n",
      "  inference(avatar_component_clause,[],[f12289])).\n",
      "fof(f12299,plain,(\n",
      "  spl0_147 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f12294,f1281,f12296])).\n",
      "fof(f12294,plain,(\n",
      "  d(f(b,m(b,one)),m(a,a)) | ~spl0_36),\n",
      "  inference(forward_demodulation,[],[f12254,f4])).\n",
      "fof(f12254,plain,(\n",
      "  d(f(b,m(one,b)),m(a,a)) | ~spl0_36),\n",
      "  inference(superposition,[],[f10323,f1283])).\n",
      "fof(f12292,plain,(\n",
      "  spl0_146 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f12287,f329,f12289])).\n",
      "fof(f12287,plain,(\n",
      "  d(f(two,m(two,one)),a) | ~spl0_20),\n",
      "  inference(forward_demodulation,[],[f12252,f4])).\n",
      "fof(f12252,plain,(\n",
      "  d(f(two,m(one,two)),a) | ~spl0_20),\n",
      "  inference(superposition,[],[f10323,f331])).\n",
      "fof(f11051,plain,(\n",
      "  spl0_144 | spl0_145),\n",
      "  inference(avatar_split_clause,[],[f10878,f11049,f11045])).\n",
      "fof(f11045,plain,(\n",
      "  spl0_144 <=> d(two,one)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_144])])).\n",
      "fof(f11049,plain,(\n",
      "  spl0_145 <=> ! [X0,X1] : (~d(two,m(X1,X0)) | d(two,m(X0,X1)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_145])])).\n",
      "fof(f10878,plain,(\n",
      "  ( ! [X0,X1] : (~d(two,m(X1,X0)) | d(two,m(X0,X1)) | d(two,one)) )),\n",
      "  inference(superposition,[],[f306,f10269])).\n",
      "fof(f306,plain,(\n",
      "  ( ! [X2,X0,X1] : (~d(two,m(X0,m(X1,X2))) | d(two,m(X0,X1)) | d(two,X2)) )),\n",
      "  inference(superposition,[],[f8,f3])).\n",
      "fof(f10826,plain,(\n",
      "  spl0_143 | ~spl0_139),\n",
      "  inference(avatar_split_clause,[],[f10821,f10644,f10823])).\n",
      "fof(f10823,plain,(\n",
      "  spl0_143 <=> a = m(f(two,a),f(m(one,f(two,a)),a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_143])])).\n",
      "fof(f10644,plain,(\n",
      "  spl0_139 <=> d(m(one,f(two,a)),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_139])])).\n",
      "fof(f10821,plain,(\n",
      "  a = m(f(two,a),f(m(one,f(two,a)),a)) | ~spl0_139),\n",
      "  inference(forward_literal_rewriting,[],[f10820,f10243])).\n",
      "fof(f10820,plain,(\n",
      "  a = m(m(one,f(two,a)),f(m(one,f(two,a)),a)) | ~spl0_139),\n",
      "  inference(resolution,[],[f10646,f6])).\n",
      "fof(f10646,plain,(\n",
      "  d(m(one,f(two,a)),a) | ~spl0_139),\n",
      "  inference(avatar_component_clause,[],[f10644])).\n",
      "fof(f10819,plain,(\n",
      "  spl0_142 | ~spl0_138),\n",
      "  inference(avatar_split_clause,[],[f10814,f10562,f10816])).\n",
      "fof(f10816,plain,(\n",
      "  spl0_142 <=> m(a,a) = m(m(b,one),f(m(b,one),m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_142])])).\n",
      "fof(f10562,plain,(\n",
      "  spl0_138 <=> d(m(b,one),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_138])])).\n",
      "fof(f10814,plain,(\n",
      "  m(a,a) = m(m(b,one),f(m(b,one),m(a,a))) | ~spl0_138),\n",
      "  inference(resolution,[],[f10564,f6])).\n",
      "fof(f10564,plain,(\n",
      "  d(m(b,one),m(a,a)) | ~spl0_138),\n",
      "  inference(avatar_component_clause,[],[f10562])).\n",
      "fof(f10759,plain,(\n",
      "  spl0_141 | ~spl0_140),\n",
      "  inference(avatar_split_clause,[],[f10754,f10750,f10756])).\n",
      "fof(f10756,plain,(\n",
      "  spl0_141 <=> a = m(m(two,one),f(m(two,one),a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_141])])).\n",
      "fof(f10750,plain,(\n",
      "  spl0_140 <=> d(m(two,one),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_140])])).\n",
      "fof(f10754,plain,(\n",
      "  a = m(m(two,one),f(m(two,one),a)) | ~spl0_140),\n",
      "  inference(resolution,[],[f10752,f6])).\n",
      "fof(f10752,plain,(\n",
      "  d(m(two,one),a) | ~spl0_140),\n",
      "  inference(avatar_component_clause,[],[f10750])).\n",
      "fof(f10753,plain,(\n",
      "  spl0_140 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f10734,f329,f10750])).\n",
      "fof(f10734,plain,(\n",
      "  d(m(two,one),a) | ~spl0_20),\n",
      "  inference(superposition,[],[f10311,f331])).\n",
      "fof(f10647,plain,(\n",
      "  spl0_139 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f10642,f329,f10644])).\n",
      "fof(f10642,plain,(\n",
      "  d(m(one,f(two,a)),a) | ~spl0_20),\n",
      "  inference(forward_demodulation,[],[f10623,f4])).\n",
      "fof(f10623,plain,(\n",
      "  d(m(f(two,a),one),a) | ~spl0_20),\n",
      "  inference(superposition,[],[f10300,f331])).\n",
      "fof(f10565,plain,(\n",
      "  spl0_138 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f10560,f1281,f10562])).\n",
      "fof(f10560,plain,(\n",
      "  d(m(b,one),m(a,a)) | ~spl0_36),\n",
      "  inference(forward_demodulation,[],[f10432,f4])).\n",
      "fof(f10432,plain,(\n",
      "  d(m(one,b),m(a,a)) | ~spl0_36),\n",
      "  inference(superposition,[],[f1461,f10249])).\n",
      "fof(f1461,plain,(\n",
      "  ( ! [X0] : (d(m(X0,b),m(a,m(X0,a)))) ) | ~spl0_36),\n",
      "  inference(superposition,[],[f1375,f151])).\n",
      "fof(f9517,plain,(\n",
      "  spl0_37 | spl0_137 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f9493,f1281,f9515,f1402])).\n",
      "fof(f1402,plain,(\n",
      "  spl0_37 <=> one = m(b,two)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_37])])).\n",
      "fof(f9493,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) != m(b,X0) | one = m(b,two)) ) | ~spl0_36),\n",
      "  inference(superposition,[],[f897,f1283])).\n",
      "fof(f8234,plain,(\n",
      "  spl0_7 | ~spl0_29 | ~spl0_53),\n",
      "  inference(avatar_split_clause,[],[f8233,f4397,f701,f120])).\n",
      "fof(f120,plain,(\n",
      "  spl0_7 <=> d(m(b,b),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_7])])).\n",
      "fof(f701,plain,(\n",
      "  spl0_29 <=> d(f(two,m(a,a)),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_29])])).\n",
      "fof(f8233,plain,(\n",
      "  d(m(b,b),m(a,a)) | (~spl0_29 | ~spl0_53)),\n",
      "  inference(backward_demodulation,[],[f703,f4399])).\n",
      "fof(f703,plain,(\n",
      "  d(f(two,m(a,a)),m(a,a)) | ~spl0_29),\n",
      "  inference(avatar_component_clause,[],[f701])).\n",
      "fof(f8222,plain,(\n",
      "  ~spl0_120 | ~spl0_3 | spl0_37),\n",
      "  inference(avatar_split_clause,[],[f8221,f1402,f42,f7168])).\n",
      "fof(f7168,plain,(\n",
      "  spl0_120 <=> two = b),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_120])])).\n",
      "fof(f8221,plain,(\n",
      "  two != b | (~spl0_3 | spl0_37)),\n",
      "  inference(forward_literal_rewriting,[],[f7750,f7745])).\n",
      "fof(f7745,plain,(\n",
      "  ( ! [X0,X1] : (m(b,X0) = X1 | X0 != X1) ) | ~spl0_3),\n",
      "  inference(backward_demodulation,[],[f995,f44])).\n",
      "fof(f44,plain,(\n",
      "  one = b | ~spl0_3),\n",
      "  inference(avatar_component_clause,[],[f42])).\n",
      "fof(f995,plain,(\n",
      "  ( ! [X0,X1] : (X0 != X1 | m(one,X0) = X1) )),\n",
      "  inference(superposition,[],[f966,f2])).\n",
      "fof(f7750,plain,(\n",
      "  b != m(b,two) | (~spl0_3 | spl0_37)),\n",
      "  inference(backward_demodulation,[],[f1403,f44])).\n",
      "fof(f1403,plain,(\n",
      "  one != m(b,two) | spl0_37),\n",
      "  inference(avatar_component_clause,[],[f1402])).\n",
      "fof(f8219,plain,(\n",
      "  ~spl0_136 | ~spl0_3 | spl0_24),\n",
      "  inference(avatar_split_clause,[],[f7738,f386,f42,f8216])).\n",
      "fof(f8216,plain,(\n",
      "  spl0_136 <=> b = f(two,a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_136])])).\n",
      "fof(f7738,plain,(\n",
      "  b != f(two,a) | (~spl0_3 | spl0_24)),\n",
      "  inference(backward_demodulation,[],[f387,f44])).\n",
      "fof(f387,plain,(\n",
      "  one != f(two,a) | spl0_24),\n",
      "  inference(avatar_component_clause,[],[f386])).\n",
      "fof(f8196,plain,(\n",
      "  ~spl0_120 | ~spl0_3 | spl0_110),\n",
      "  inference(avatar_split_clause,[],[f8195,f6219,f42,f7168])).\n",
      "fof(f6219,plain,(\n",
      "  spl0_110 <=> two = f(two,two)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_110])])).\n",
      "fof(f8195,plain,(\n",
      "  two != b | (~spl0_3 | spl0_110)),\n",
      "  inference(backward_demodulation,[],[f6221,f7732])).\n",
      "fof(f7732,plain,(\n",
      "  ( ! [X0] : (b = f(X0,X0)) ) | ~spl0_3),\n",
      "  inference(backward_demodulation,[],[f79,f44])).\n",
      "fof(f6221,plain,(\n",
      "  two != f(two,two) | spl0_110),\n",
      "  inference(avatar_component_clause,[],[f6219])).\n",
      "fof(f8191,plain,(\n",
      "  ~spl0_120 | spl0_1 | ~spl0_3),\n",
      "  inference(avatar_split_clause,[],[f7725,f42,f14,f7168])).\n",
      "fof(f7725,plain,(\n",
      "  two != b | (spl0_1 | ~spl0_3)),\n",
      "  inference(backward_demodulation,[],[f16,f44])).\n",
      "fof(f16,plain,(\n",
      "  one != two | spl0_1),\n",
      "  inference(avatar_component_clause,[],[f14])).\n",
      "fof(f8177,plain,(\n",
      "  ~spl0_120 | ~spl0_3 | spl0_114),\n",
      "  inference(avatar_split_clause,[],[f7946,f7016,f42,f7168])).\n",
      "fof(f7016,plain,(\n",
      "  spl0_114 <=> two = m(b,b)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_114])])).\n",
      "fof(f7946,plain,(\n",
      "  two != b | (~spl0_3 | spl0_114)),\n",
      "  inference(backward_demodulation,[],[f7017,f7722])).\n",
      "fof(f7722,plain,(\n",
      "  ( ! [X0] : (m(b,X0) = X0) ) | ~spl0_3),\n",
      "  inference(backward_demodulation,[],[f1,f44])).\n",
      "fof(f7017,plain,(\n",
      "  two != m(b,b) | spl0_114),\n",
      "  inference(avatar_component_clause,[],[f7016])).\n",
      "fof(f8175,plain,(\n",
      "  spl0_135 | ~spl0_3 | ~spl0_60),\n",
      "  inference(avatar_split_clause,[],[f7943,f4674,f42,f8172])).\n",
      "fof(f8172,plain,(\n",
      "  spl0_135 <=> m(a,a) = m(f(b,b),f(f(b,b),m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_135])])).\n",
      "fof(f4674,plain,(\n",
      "  spl0_60 <=> m(a,a) = m(f(b,m(b,b)),f(f(b,m(b,b)),m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_60])])).\n",
      "fof(f7943,plain,(\n",
      "  m(a,a) = m(f(b,b),f(f(b,b),m(a,a))) | (~spl0_3 | ~spl0_60)),\n",
      "  inference(backward_demodulation,[],[f4676,f7722])).\n",
      "fof(f4676,plain,(\n",
      "  m(a,a) = m(f(b,m(b,b)),f(f(b,m(b,b)),m(a,a))) | ~spl0_60),\n",
      "  inference(avatar_component_clause,[],[f4674])).\n",
      "fof(f8170,plain,(\n",
      "  spl0_134 | ~spl0_3 | ~spl0_58),\n",
      "  inference(avatar_split_clause,[],[f7942,f4663,f42,f8167])).\n",
      "fof(f8167,plain,(\n",
      "  spl0_134 <=> d(f(b,b),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_134])])).\n",
      "fof(f4663,plain,(\n",
      "  spl0_58 <=> d(f(b,m(b,b)),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_58])])).\n",
      "fof(f7942,plain,(\n",
      "  d(f(b,b),m(a,a)) | (~spl0_3 | ~spl0_58)),\n",
      "  inference(backward_demodulation,[],[f4665,f7722])).\n",
      "fof(f4665,plain,(\n",
      "  d(f(b,m(b,b)),m(a,a)) | ~spl0_58),\n",
      "  inference(avatar_component_clause,[],[f4663])).\n",
      "fof(f8165,plain,(\n",
      "  spl0_133 | ~spl0_3 | ~spl0_53),\n",
      "  inference(avatar_split_clause,[],[f7941,f4397,f42,f8162])).\n",
      "fof(f8162,plain,(\n",
      "  spl0_133 <=> b = f(two,m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_133])])).\n",
      "fof(f7941,plain,(\n",
      "  b = f(two,m(a,a)) | (~spl0_3 | ~spl0_53)),\n",
      "  inference(backward_demodulation,[],[f4399,f7722])).\n",
      "fof(f8146,plain,(\n",
      "  spl0_132 | ~spl0_3 | ~spl0_63),\n",
      "  inference(avatar_split_clause,[],[f7910,f4960,f42,f8143])).\n",
      "fof(f8143,plain,(\n",
      "  spl0_132 <=> m(a,a) = m(f(two,two),f(f(two,two),m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_132])])).\n",
      "fof(f4960,plain,(\n",
      "  spl0_63 <=> m(a,a) = m(f(two,m(b,two)),f(f(two,m(b,two)),m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_63])])).\n",
      "fof(f7910,plain,(\n",
      "  m(a,a) = m(f(two,two),f(f(two,two),m(a,a))) | (~spl0_3 | ~spl0_63)),\n",
      "  inference(backward_demodulation,[],[f4962,f7722])).\n",
      "fof(f4962,plain,(\n",
      "  m(a,a) = m(f(two,m(b,two)),f(f(two,m(b,two)),m(a,a))) | ~spl0_63),\n",
      "  inference(avatar_component_clause,[],[f4960])).\n",
      "fof(f8141,plain,(\n",
      "  spl0_131 | ~spl0_3 | ~spl0_62),\n",
      "  inference(avatar_split_clause,[],[f7909,f4947,f42,f8138])).\n",
      "fof(f8138,plain,(\n",
      "  spl0_131 <=> d(f(two,two),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_131])])).\n",
      "fof(f4947,plain,(\n",
      "  spl0_62 <=> d(f(two,m(b,two)),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_62])])).\n",
      "fof(f7909,plain,(\n",
      "  d(f(two,two),m(a,a)) | (~spl0_3 | ~spl0_62)),\n",
      "  inference(backward_demodulation,[],[f4949,f7722])).\n",
      "fof(f4949,plain,(\n",
      "  d(f(two,m(b,two)),m(a,a)) | ~spl0_62),\n",
      "  inference(avatar_component_clause,[],[f4947])).\n",
      "fof(f8135,plain,(\n",
      "  spl0_130 | ~spl0_3 | ~spl0_61),\n",
      "  inference(avatar_split_clause,[],[f7907,f4680,f42,f8132])).\n",
      "fof(f8132,plain,(\n",
      "  spl0_130 <=> m(a,a) = m(f(b,two),f(f(b,two),m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_130])])).\n",
      "fof(f4680,plain,(\n",
      "  spl0_61 <=> m(a,a) = m(f(b,m(b,two)),f(f(b,m(b,two)),m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_61])])).\n",
      "fof(f7907,plain,(\n",
      "  m(a,a) = m(f(b,two),f(f(b,two),m(a,a))) | (~spl0_3 | ~spl0_61)),\n",
      "  inference(backward_demodulation,[],[f4682,f7722])).\n",
      "fof(f4682,plain,(\n",
      "  m(a,a) = m(f(b,m(b,two)),f(f(b,m(b,two)),m(a,a))) | ~spl0_61),\n",
      "  inference(avatar_component_clause,[],[f4680])).\n",
      "fof(f8130,plain,(\n",
      "  spl0_129 | ~spl0_3 | ~spl0_59),\n",
      "  inference(avatar_split_clause,[],[f7906,f4668,f42,f8127])).\n",
      "fof(f8127,plain,(\n",
      "  spl0_129 <=> d(f(b,two),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_129])])).\n",
      "fof(f4668,plain,(\n",
      "  spl0_59 <=> d(f(b,m(b,two)),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_59])])).\n",
      "fof(f7906,plain,(\n",
      "  d(f(b,two),m(a,a)) | (~spl0_3 | ~spl0_59)),\n",
      "  inference(backward_demodulation,[],[f4670,f7722])).\n",
      "fof(f4670,plain,(\n",
      "  d(f(b,m(b,two)),m(a,a)) | ~spl0_59),\n",
      "  inference(avatar_component_clause,[],[f4668])).\n",
      "fof(f8124,plain,(\n",
      "  spl0_29 | ~spl0_3 | ~spl0_41),\n",
      "  inference(avatar_split_clause,[],[f7895,f1444,f42,f701])).\n",
      "fof(f1444,plain,(\n",
      "  spl0_41 <=> d(f(m(b,two),m(a,a)),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_41])])).\n",
      "fof(f7895,plain,(\n",
      "  d(f(two,m(a,a)),m(a,a)) | (~spl0_3 | ~spl0_41)),\n",
      "  inference(backward_demodulation,[],[f1446,f7722])).\n",
      "fof(f1446,plain,(\n",
      "  d(f(m(b,two),m(a,a)),m(a,a)) | ~spl0_41),\n",
      "  inference(avatar_component_clause,[],[f1444])).\n",
      "fof(f8120,plain,(\n",
      "  spl0_128 | ~spl0_3 | ~spl0_46),\n",
      "  inference(avatar_split_clause,[],[f8115,f1804,f42,f8117])).\n",
      "fof(f8117,plain,(\n",
      "  spl0_128 <=> a = m(m(a,a),f(m(a,a),a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_128])])).\n",
      "fof(f1804,plain,(\n",
      "  spl0_46 <=> m(b,m(b,a)) = m(m(a,a),f(m(a,a),m(b,m(b,a))))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_46])])).\n",
      "fof(f8115,plain,(\n",
      "  a = m(m(a,a),f(m(a,a),a)) | (~spl0_3 | ~spl0_46)),\n",
      "  inference(forward_demodulation,[],[f7880,f7722])).\n",
      "fof(f7880,plain,(\n",
      "  m(b,a) = m(m(a,a),f(m(a,a),m(b,a))) | (~spl0_3 | ~spl0_46)),\n",
      "  inference(backward_demodulation,[],[f1806,f7722])).\n",
      "fof(f1806,plain,(\n",
      "  m(b,m(b,a)) = m(m(a,a),f(m(a,a),m(b,m(b,a)))) | ~spl0_46),\n",
      "  inference(avatar_component_clause,[],[f1804])).\n",
      "fof(f8114,plain,(\n",
      "  spl0_127 | ~spl0_3 | ~spl0_28),\n",
      "  inference(avatar_split_clause,[],[f8109,f458,f42,f8111])).\n",
      "fof(f8111,plain,(\n",
      "  spl0_127 <=> d(m(a,a),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_127])])).\n",
      "fof(f458,plain,(\n",
      "  spl0_28 <=> d(m(a,a),m(b,m(b,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_28])])).\n",
      "fof(f8109,plain,(\n",
      "  d(m(a,a),a) | (~spl0_3 | ~spl0_28)),\n",
      "  inference(forward_demodulation,[],[f7879,f7722])).\n",
      "fof(f7879,plain,(\n",
      "  d(m(a,a),m(b,a)) | (~spl0_3 | ~spl0_28)),\n",
      "  inference(backward_demodulation,[],[f460,f7722])).\n",
      "fof(f460,plain,(\n",
      "  d(m(a,a),m(b,m(b,a))) | ~spl0_28),\n",
      "  inference(avatar_component_clause,[],[f458])).\n",
      "fof(f8108,plain,(\n",
      "  spl0_9 | ~spl0_3 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f8107,f1281,f42,f129])).\n",
      "fof(f129,plain,(\n",
      "  spl0_9 <=> two = m(a,a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_9])])).\n",
      "fof(f8107,plain,(\n",
      "  two = m(a,a) | (~spl0_3 | ~spl0_36)),\n",
      "  inference(forward_demodulation,[],[f7878,f7722])).\n",
      "fof(f7878,plain,(\n",
      "  m(a,a) = m(b,two) | (~spl0_3 | ~spl0_36)),\n",
      "  inference(backward_demodulation,[],[f1283,f7722])).\n",
      "fof(f7660,plain,(\n",
      "  spl0_3 | spl0_126 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f7595,f19,f7658,f42])).\n",
      "fof(f7595,plain,(\n",
      "  ( ! [X0] : (m(a,m(a,X0)) != m(b,m(X0,two)) | one = b) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f65,f7203])).\n",
      "fof(f7208,plain,(\n",
      "  ~spl0_25 | spl0_24 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f5682,f329,f386,f390])).\n",
      "fof(f5682,plain,(\n",
      "  one = f(two,a) | two != a | ~spl0_20),\n",
      "  inference(forward_demodulation,[],[f839,f79])).\n",
      "fof(f839,plain,(\n",
      "  two != a | f(two,a) = f(two,two) | ~spl0_20),\n",
      "  inference(superposition,[],[f74,f331])).\n",
      "fof(f7200,plain,(\n",
      "  spl0_122 | ~spl0_120 | ~spl0_114),\n",
      "  inference(avatar_split_clause,[],[f7113,f7016,f7168,f7179])).\n",
      "fof(f7179,plain,(\n",
      "  spl0_122 <=> b = f(b,b)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_122])])).\n",
      "fof(f7113,plain,(\n",
      "  two != b | b = f(b,b) | ~spl0_114),\n",
      "  inference(superposition,[],[f833,f7018])).\n",
      "fof(f7018,plain,(\n",
      "  two = m(b,b) | ~spl0_114),\n",
      "  inference(avatar_component_clause,[],[f7016])).\n",
      "fof(f7199,plain,(\n",
      "  spl0_125 | ~spl0_114),\n",
      "  inference(avatar_split_clause,[],[f7110,f7016,f7196])).\n",
      "fof(f7196,plain,(\n",
      "  spl0_125 <=> d(f(b,two),two)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_125])])).\n",
      "fof(f7110,plain,(\n",
      "  d(f(b,two),two) | ~spl0_114),\n",
      "  inference(superposition,[],[f643,f7018])).\n",
      "fof(f7194,plain,(\n",
      "  spl0_118 | spl0_124 | ~spl0_114),\n",
      "  inference(avatar_split_clause,[],[f7101,f7016,f7192,f7158])).\n",
      "fof(f7192,plain,(\n",
      "  spl0_124 <=> ! [X0] : (~d(two,m(X0,two)) | d(two,m(X0,b)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_124])])).\n",
      "fof(f7101,plain,(\n",
      "  ( ! [X0] : (~d(two,m(X0,two)) | d(two,m(X0,b)) | d(two,b)) ) | ~spl0_114),\n",
      "  inference(superposition,[],[f306,f7018])).\n",
      "fof(f7187,plain,(\n",
      "  spl0_3 | spl0_123 | ~spl0_114),\n",
      "  inference(avatar_split_clause,[],[f7089,f7016,f7185,f42])).\n",
      "fof(f7185,plain,(\n",
      "  spl0_123 <=> ! [X0] : m(X0,two) != m(X0,b)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_123])])).\n",
      "fof(f7089,plain,(\n",
      "  ( ! [X0] : (m(X0,two) != m(X0,b) | one = b) ) | ~spl0_114),\n",
      "  inference(superposition,[],[f161,f7018])).\n",
      "fof(f7182,plain,(\n",
      "  spl0_122 | ~spl0_120 | ~spl0_114),\n",
      "  inference(avatar_split_clause,[],[f7081,f7016,f7168,f7179])).\n",
      "fof(f7081,plain,(\n",
      "  two != b | b = f(b,b) | ~spl0_114),\n",
      "  inference(superposition,[],[f74,f7018])).\n",
      "fof(f7177,plain,(\n",
      "  spl0_121 | ~spl0_114),\n",
      "  inference(avatar_split_clause,[],[f7079,f7016,f7174])).\n",
      "fof(f7174,plain,(\n",
      "  spl0_121 <=> two = m(b,f(b,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_121])])).\n",
      "fof(f7079,plain,(\n",
      "  two = m(b,f(b,two)) | ~spl0_114),\n",
      "  inference(superposition,[],[f67,f7018])).\n",
      "fof(f7172,plain,(\n",
      "  spl0_3 | ~spl0_120 | ~spl0_114),\n",
      "  inference(avatar_split_clause,[],[f7078,f7016,f7168,f42])).\n",
      "fof(f7078,plain,(\n",
      "  two != b | one = b | ~spl0_114),\n",
      "  inference(superposition,[],[f65,f7018])).\n",
      "fof(f7171,plain,(\n",
      "  spl0_3 | ~spl0_120 | ~spl0_114),\n",
      "  inference(avatar_split_clause,[],[f7075,f7016,f7168,f42])).\n",
      "fof(f7075,plain,(\n",
      "  two != b | one = b | ~spl0_114),\n",
      "  inference(superposition,[],[f51,f7018])).\n",
      "fof(f7166,plain,(\n",
      "  spl0_107 | ~spl0_114),\n",
      "  inference(avatar_split_clause,[],[f7073,f7016,f6203])).\n",
      "fof(f6203,plain,(\n",
      "  spl0_107 <=> d(b,two)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_107])])).\n",
      "fof(f7073,plain,(\n",
      "  d(b,two) | ~spl0_114),\n",
      "  inference(superposition,[],[f12,f7018])).\n",
      "fof(f7165,plain,(\n",
      "  spl0_118 | ~spl0_119 | ~spl0_114),\n",
      "  inference(avatar_split_clause,[],[f7128,f7016,f7162,f7158])).\n",
      "fof(f7162,plain,(\n",
      "  spl0_119 <=> d(two,two)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_119])])).\n",
      "fof(f7128,plain,(\n",
      "  ~d(two,two) | d(two,b) | ~spl0_114),\n",
      "  inference(duplicate_literal_removal,[],[f7072])).\n",
      "fof(f7072,plain,(\n",
      "  ~d(two,two) | d(two,b) | d(two,b) | ~spl0_114),\n",
      "  inference(superposition,[],[f8,f7018])).\n",
      "fof(f7068,plain,(\n",
      "  spl0_117 | ~spl0_98 | ~spl0_114),\n",
      "  inference(avatar_split_clause,[],[f7042,f7016,f6057,f7065])).\n",
      "fof(f7065,plain,(\n",
      "  spl0_117 <=> two = f(two,m(two,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_117])])).\n",
      "fof(f6057,plain,(\n",
      "  spl0_98 <=> m(b,b) = f(two,m(two,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_98])])).\n",
      "fof(f7042,plain,(\n",
      "  two = f(two,m(two,two)) | (~spl0_98 | ~spl0_114)),\n",
      "  inference(backward_demodulation,[],[f6059,f7018])).\n",
      "fof(f6059,plain,(\n",
      "  m(b,b) = f(two,m(two,two)) | ~spl0_98),\n",
      "  inference(avatar_component_clause,[],[f6057])).\n",
      "fof(f7063,plain,(\n",
      "  spl0_116 | ~spl0_97 | ~spl0_114),\n",
      "  inference(avatar_split_clause,[],[f7041,f7016,f6051,f7060])).\n",
      "fof(f7060,plain,(\n",
      "  spl0_116 <=> d(f(b,two),m(two,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_116])])).\n",
      "fof(f6051,plain,(\n",
      "  spl0_97 <=> d(f(b,m(b,b)),m(two,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_97])])).\n",
      "fof(f7041,plain,(\n",
      "  d(f(b,two),m(two,two)) | (~spl0_97 | ~spl0_114)),\n",
      "  inference(backward_demodulation,[],[f6053,f7018])).\n",
      "fof(f6053,plain,(\n",
      "  d(f(b,m(b,b)),m(two,two)) | ~spl0_97),\n",
      "  inference(avatar_component_clause,[],[f6051])).\n",
      "fof(f7058,plain,(\n",
      "  spl0_115 | ~spl0_96 | ~spl0_114),\n",
      "  inference(avatar_split_clause,[],[f7040,f7016,f6045,f7055])).\n",
      "fof(f7055,plain,(\n",
      "  spl0_115 <=> m(two,two) = m(f(b,two),f(f(b,two),m(two,two)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_115])])).\n",
      "fof(f6045,plain,(\n",
      "  spl0_96 <=> m(two,two) = m(f(b,m(b,b)),f(f(b,m(b,b)),m(two,two)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_96])])).\n",
      "fof(f7040,plain,(\n",
      "  m(two,two) = m(f(b,two),f(f(b,two),m(two,two))) | (~spl0_96 | ~spl0_114)),\n",
      "  inference(backward_demodulation,[],[f6047,f7018])).\n",
      "fof(f6047,plain,(\n",
      "  m(two,two) = m(f(b,m(b,b)),f(f(b,m(b,b)),m(two,two))) | ~spl0_96),\n",
      "  inference(avatar_component_clause,[],[f6045])).\n",
      "fof(f7019,plain,(\n",
      "  spl0_114 | ~spl0_2 | ~spl0_25),\n",
      "  inference(avatar_split_clause,[],[f7014,f390,f19,f7016])).\n",
      "fof(f7014,plain,(\n",
      "  two = m(b,b) | (~spl0_2 | ~spl0_25)),\n",
      "  inference(equality_resolution,[],[f6113])).\n",
      "fof(f6113,plain,(\n",
      "  ( ! [X0] : (m(two,X0) != m(two,two) | m(b,b) = X0) ) | (~spl0_2 | ~spl0_25)),\n",
      "  inference(backward_demodulation,[],[f108,f391])).\n",
      "fof(f391,plain,(\n",
      "  two = a | ~spl0_25),\n",
      "  inference(avatar_component_clause,[],[f390])).\n",
      "fof(f108,plain,(\n",
      "  ( ! [X0] : (m(a,a) != m(two,X0) | m(b,b) = X0) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f5,f21])).\n",
      "fof(f6245,plain,(\n",
      "  spl0_113 | ~spl0_25 | ~spl0_99),\n",
      "  inference(avatar_split_clause,[],[f6240,f6063,f390,f6242])).\n",
      "fof(f6242,plain,(\n",
      "  spl0_113 <=> m(b,m(b,two)) = m(m(two,two),f(m(two,two),m(b,m(b,two))))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_113])])).\n",
      "fof(f6063,plain,(\n",
      "  spl0_99 <=> m(b,m(b,a)) = m(m(two,two),f(m(two,two),m(b,m(b,a))))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_99])])).\n",
      "fof(f6240,plain,(\n",
      "  m(b,m(b,two)) = m(m(two,two),f(m(two,two),m(b,m(b,two)))) | (~spl0_25 | ~spl0_99)),\n",
      "  inference(forward_demodulation,[],[f6065,f391])).\n",
      "fof(f6065,plain,(\n",
      "  m(b,m(b,a)) = m(m(two,two),f(m(two,two),m(b,m(b,a)))) | ~spl0_99),\n",
      "  inference(avatar_component_clause,[],[f6063])).\n",
      "fof(f6235,plain,(\n",
      "  ~spl0_108 | ~spl0_25 | spl0_45),\n",
      "  inference(avatar_split_clause,[],[f6137,f1566,f390,f6208])).\n",
      "fof(f6208,plain,(\n",
      "  spl0_108 <=> two = m(two,two)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_108])])).\n",
      "fof(f1566,plain,(\n",
      "  spl0_45 <=> a = m(a,a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_45])])).\n",
      "fof(f6137,plain,(\n",
      "  two != m(two,two) | (~spl0_25 | spl0_45)),\n",
      "  inference(backward_demodulation,[],[f1568,f391])).\n",
      "fof(f1568,plain,(\n",
      "  a != m(a,a) | spl0_45),\n",
      "  inference(avatar_component_clause,[],[f1566])).\n",
      "fof(f6234,plain,(\n",
      "  ~spl0_112 | ~spl0_25 | spl0_39),\n",
      "  inference(avatar_split_clause,[],[f6136,f1411,f390,f6231])).\n",
      "fof(f6231,plain,(\n",
      "  spl0_112 <=> m(b,two) = m(two,two)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_112])])).\n",
      "fof(f1411,plain,(\n",
      "  spl0_39 <=> m(a,a) = m(b,two)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_39])])).\n",
      "fof(f6136,plain,(\n",
      "  m(b,two) != m(two,two) | (~spl0_25 | spl0_39)),\n",
      "  inference(backward_demodulation,[],[f1413,f391])).\n",
      "fof(f1413,plain,(\n",
      "  m(a,a) != m(b,two) | spl0_39),\n",
      "  inference(avatar_component_clause,[],[f1411])).\n",
      "fof(f6229,plain,(\n",
      "  ~spl0_111 | ~spl0_25 | spl0_38),\n",
      "  inference(avatar_split_clause,[],[f6135,f1406,f390,f6226])).\n",
      "fof(f6226,plain,(\n",
      "  spl0_111 <=> b = m(two,two)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_111])])).\n",
      "fof(f1406,plain,(\n",
      "  spl0_38 <=> m(a,a) = b),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_38])])).\n",
      "fof(f6135,plain,(\n",
      "  b != m(two,two) | (~spl0_25 | spl0_38)),\n",
      "  inference(backward_demodulation,[],[f1408,f391])).\n",
      "fof(f1408,plain,(\n",
      "  m(a,a) != b | spl0_38),\n",
      "  inference(avatar_component_clause,[],[f1406])).\n",
      "fof(f6222,plain,(\n",
      "  ~spl0_110 | ~spl0_25 | spl0_26),\n",
      "  inference(avatar_split_clause,[],[f6122,f395,f390,f6219])).\n",
      "fof(f395,plain,(\n",
      "  spl0_26 <=> a = f(two,a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_26])])).\n",
      "fof(f6122,plain,(\n",
      "  two != f(two,two) | (~spl0_25 | spl0_26)),\n",
      "  inference(backward_demodulation,[],[f397,f391])).\n",
      "fof(f397,plain,(\n",
      "  a != f(two,a) | spl0_26),\n",
      "  inference(avatar_component_clause,[],[f395])).\n",
      "fof(f6216,plain,(\n",
      "  ~spl0_109 | spl0_10 | ~spl0_25),\n",
      "  inference(avatar_split_clause,[],[f6115,f390,f134,f6213])).\n",
      "fof(f6213,plain,(\n",
      "  spl0_109 <=> m(b,b) = m(two,two)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_109])])).\n",
      "fof(f134,plain,(\n",
      "  spl0_10 <=> m(a,a) = m(b,b)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_10])])).\n",
      "fof(f6115,plain,(\n",
      "  m(b,b) != m(two,two) | (spl0_10 | ~spl0_25)),\n",
      "  inference(backward_demodulation,[],[f136,f391])).\n",
      "fof(f136,plain,(\n",
      "  m(a,a) != m(b,b) | spl0_10),\n",
      "  inference(avatar_component_clause,[],[f134])).\n",
      "fof(f6211,plain,(\n",
      "  ~spl0_108 | spl0_9 | ~spl0_25),\n",
      "  inference(avatar_split_clause,[],[f6114,f390,f129,f6208])).\n",
      "fof(f6114,plain,(\n",
      "  two != m(two,two) | (spl0_9 | ~spl0_25)),\n",
      "  inference(backward_demodulation,[],[f131,f391])).\n",
      "fof(f131,plain,(\n",
      "  two != m(a,a) | spl0_9),\n",
      "  inference(avatar_component_clause,[],[f129])).\n",
      "fof(f6206,plain,(\n",
      "  ~spl0_107 | spl0_4 | ~spl0_25),\n",
      "  inference(avatar_split_clause,[],[f6112,f390,f46,f6203])).\n",
      "fof(f46,plain,(\n",
      "  spl0_4 <=> d(b,a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_4])])).\n",
      "fof(f6112,plain,(\n",
      "  ~d(b,two) | (spl0_4 | ~spl0_25)),\n",
      "  inference(backward_demodulation,[],[f48,f391])).\n",
      "fof(f48,plain,(\n",
      "  ~d(b,a) | spl0_4),\n",
      "  inference(avatar_component_clause,[],[f46])).\n",
      "fof(f6108,plain,(\n",
      "  spl0_106 | ~spl0_20 | ~spl0_24 | ~spl0_76),\n",
      "  inference(avatar_split_clause,[],[f6103,f5541,f386,f329,f6105])).\n",
      "fof(f6105,plain,(\n",
      "  spl0_106 <=> d(m(b,b),m(two,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_106])])).\n",
      "fof(f5541,plain,(\n",
      "  spl0_76 <=> d(m(b,b),m(two,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_76])])).\n",
      "fof(f6103,plain,(\n",
      "  d(m(b,b),m(two,two)) | (~spl0_20 | ~spl0_24 | ~spl0_76)),\n",
      "  inference(forward_demodulation,[],[f5543,f5625])).\n",
      "fof(f5625,plain,(\n",
      "  ( ! [X0] : (m(two,X0) = m(X0,a)) ) | (~spl0_20 | ~spl0_24)),\n",
      "  inference(forward_demodulation,[],[f5288,f1])).\n",
      "fof(f5288,plain,(\n",
      "  ( ! [X0] : (m(X0,a) = m(two,m(one,X0))) ) | (~spl0_20 | ~spl0_24)),\n",
      "  inference(backward_demodulation,[],[f1102,f388])).\n",
      "fof(f388,plain,(\n",
      "  one = f(two,a) | ~spl0_24),\n",
      "  inference(avatar_component_clause,[],[f386])).\n",
      "fof(f5543,plain,(\n",
      "  d(m(b,b),m(two,a)) | ~spl0_76),\n",
      "  inference(avatar_component_clause,[],[f5541])).\n",
      "fof(f6102,plain,(\n",
      "  spl0_105 | ~spl0_20 | ~spl0_24 | ~spl0_77),\n",
      "  inference(avatar_split_clause,[],[f6097,f5546,f386,f329,f6099])).\n",
      "fof(f6099,plain,(\n",
      "  spl0_105 <=> d(b,m(two,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_105])])).\n",
      "fof(f5546,plain,(\n",
      "  spl0_77 <=> d(b,m(two,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_77])])).\n",
      "fof(f6097,plain,(\n",
      "  d(b,m(two,two)) | (~spl0_20 | ~spl0_24 | ~spl0_77)),\n",
      "  inference(forward_demodulation,[],[f5548,f5625])).\n",
      "fof(f5548,plain,(\n",
      "  d(b,m(two,a)) | ~spl0_77),\n",
      "  inference(avatar_component_clause,[],[f5546])).\n",
      "fof(f6096,plain,(\n",
      "  spl0_104 | ~spl0_20 | ~spl0_24 | ~spl0_78),\n",
      "  inference(avatar_split_clause,[],[f6091,f5551,f386,f329,f6093])).\n",
      "fof(f6093,plain,(\n",
      "  spl0_104 <=> m(two,two) = m(b,f(b,m(two,two)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_104])])).\n",
      "fof(f5551,plain,(\n",
      "  spl0_78 <=> m(two,a) = m(b,f(b,m(two,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_78])])).\n",
      "fof(f6091,plain,(\n",
      "  m(two,two) = m(b,f(b,m(two,two))) | (~spl0_20 | ~spl0_24 | ~spl0_78)),\n",
      "  inference(forward_demodulation,[],[f5553,f5625])).\n",
      "fof(f5553,plain,(\n",
      "  m(two,a) = m(b,f(b,m(two,a))) | ~spl0_78),\n",
      "  inference(avatar_component_clause,[],[f5551])).\n",
      "fof(f6090,plain,(\n",
      "  spl0_103 | ~spl0_20 | ~spl0_24 | ~spl0_79),\n",
      "  inference(avatar_split_clause,[],[f6085,f5556,f386,f329,f6087])).\n",
      "fof(f6087,plain,(\n",
      "  spl0_103 <=> m(two,two) = m(m(b,b),f(m(b,b),m(two,two)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_103])])).\n",
      "fof(f5556,plain,(\n",
      "  spl0_79 <=> m(two,a) = m(m(b,b),f(m(b,b),m(two,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_79])])).\n",
      "fof(f6085,plain,(\n",
      "  m(two,two) = m(m(b,b),f(m(b,b),m(two,two))) | (~spl0_20 | ~spl0_24 | ~spl0_79)),\n",
      "  inference(forward_demodulation,[],[f5558,f5625])).\n",
      "fof(f5558,plain,(\n",
      "  m(two,a) = m(m(b,b),f(m(b,b),m(two,a))) | ~spl0_79),\n",
      "  inference(avatar_component_clause,[],[f5556])).\n",
      "fof(f6084,plain,(\n",
      "  spl0_102 | ~spl0_20 | ~spl0_24 | ~spl0_80),\n",
      "  inference(avatar_split_clause,[],[f6079,f5561,f386,f329,f6081])).\n",
      "fof(f6081,plain,(\n",
      "  spl0_102 <=> d(f(m(b,b),m(two,two)),m(two,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_102])])).\n",
      "fof(f5561,plain,(\n",
      "  spl0_80 <=> d(f(m(b,b),m(two,a)),m(two,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_80])])).\n",
      "fof(f6079,plain,(\n",
      "  d(f(m(b,b),m(two,two)),m(two,two)) | (~spl0_20 | ~spl0_24 | ~spl0_80)),\n",
      "  inference(forward_demodulation,[],[f5563,f5625])).\n",
      "fof(f5563,plain,(\n",
      "  d(f(m(b,b),m(two,a)),m(two,a)) | ~spl0_80),\n",
      "  inference(avatar_component_clause,[],[f5561])).\n",
      "fof(f6078,plain,(\n",
      "  spl0_101 | ~spl0_20 | ~spl0_24 | ~spl0_81),\n",
      "  inference(avatar_split_clause,[],[f6073,f5566,f386,f329,f6075])).\n",
      "fof(f6075,plain,(\n",
      "  spl0_101 <=> d(f(b,m(two,two)),m(two,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_101])])).\n",
      "fof(f5566,plain,(\n",
      "  spl0_81 <=> d(f(b,m(two,a)),m(two,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_81])])).\n",
      "fof(f6073,plain,(\n",
      "  d(f(b,m(two,two)),m(two,two)) | (~spl0_20 | ~spl0_24 | ~spl0_81)),\n",
      "  inference(forward_demodulation,[],[f5568,f5625])).\n",
      "fof(f5568,plain,(\n",
      "  d(f(b,m(two,a)),m(two,a)) | ~spl0_81),\n",
      "  inference(avatar_component_clause,[],[f5566])).\n",
      "fof(f6072,plain,(\n",
      "  spl0_100 | ~spl0_20 | ~spl0_24 | ~spl0_82),\n",
      "  inference(avatar_split_clause,[],[f6067,f5571,f386,f329,f6069])).\n",
      "fof(f6069,plain,(\n",
      "  spl0_100 <=> m(two,two) = m(f(b,m(two,two)),f(f(b,m(two,two)),m(two,two)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_100])])).\n",
      "fof(f5571,plain,(\n",
      "  spl0_82 <=> m(two,a) = m(f(b,m(two,a)),f(f(b,m(two,a)),m(two,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_82])])).\n",
      "fof(f6067,plain,(\n",
      "  m(two,two) = m(f(b,m(two,two)),f(f(b,m(two,two)),m(two,two))) | (~spl0_20 | ~spl0_24 | ~spl0_82)),\n",
      "  inference(forward_demodulation,[],[f5573,f5625])).\n",
      "fof(f5573,plain,(\n",
      "  m(two,a) = m(f(b,m(two,a)),f(f(b,m(two,a)),m(two,a))) | ~spl0_82),\n",
      "  inference(avatar_component_clause,[],[f5571])).\n",
      "fof(f6066,plain,(\n",
      "  spl0_99 | ~spl0_20 | ~spl0_24 | ~spl0_83),\n",
      "  inference(avatar_split_clause,[],[f6061,f5576,f386,f329,f6063])).\n",
      "fof(f5576,plain,(\n",
      "  spl0_83 <=> m(b,m(b,a)) = m(m(two,a),f(m(two,a),m(b,m(b,a))))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_83])])).\n",
      "fof(f6061,plain,(\n",
      "  m(b,m(b,a)) = m(m(two,two),f(m(two,two),m(b,m(b,a)))) | (~spl0_20 | ~spl0_24 | ~spl0_83)),\n",
      "  inference(forward_demodulation,[],[f5578,f5625])).\n",
      "fof(f5578,plain,(\n",
      "  m(b,m(b,a)) = m(m(two,a),f(m(two,a),m(b,m(b,a)))) | ~spl0_83),\n",
      "  inference(avatar_component_clause,[],[f5576])).\n",
      "fof(f6060,plain,(\n",
      "  spl0_98 | ~spl0_20 | ~spl0_24 | ~spl0_84),\n",
      "  inference(avatar_split_clause,[],[f6055,f5583,f386,f329,f6057])).\n",
      "fof(f5583,plain,(\n",
      "  spl0_84 <=> m(b,b) = f(two,m(two,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_84])])).\n",
      "fof(f6055,plain,(\n",
      "  m(b,b) = f(two,m(two,two)) | (~spl0_20 | ~spl0_24 | ~spl0_84)),\n",
      "  inference(forward_demodulation,[],[f5585,f5625])).\n",
      "fof(f5585,plain,(\n",
      "  m(b,b) = f(two,m(two,a)) | ~spl0_84),\n",
      "  inference(avatar_component_clause,[],[f5583])).\n",
      "fof(f6054,plain,(\n",
      "  spl0_97 | ~spl0_20 | ~spl0_24 | ~spl0_85),\n",
      "  inference(avatar_split_clause,[],[f6049,f5588,f386,f329,f6051])).\n",
      "fof(f5588,plain,(\n",
      "  spl0_85 <=> d(f(b,m(b,b)),m(two,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_85])])).\n",
      "fof(f6049,plain,(\n",
      "  d(f(b,m(b,b)),m(two,two)) | (~spl0_20 | ~spl0_24 | ~spl0_85)),\n",
      "  inference(forward_demodulation,[],[f5590,f5625])).\n",
      "fof(f5590,plain,(\n",
      "  d(f(b,m(b,b)),m(two,a)) | ~spl0_85),\n",
      "  inference(avatar_component_clause,[],[f5588])).\n",
      "fof(f6048,plain,(\n",
      "  spl0_96 | ~spl0_20 | ~spl0_24 | ~spl0_86),\n",
      "  inference(avatar_split_clause,[],[f6043,f5593,f386,f329,f6045])).\n",
      "fof(f5593,plain,(\n",
      "  spl0_86 <=> m(two,a) = m(f(b,m(b,b)),f(f(b,m(b,b)),m(two,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_86])])).\n",
      "fof(f6043,plain,(\n",
      "  m(two,two) = m(f(b,m(b,b)),f(f(b,m(b,b)),m(two,two))) | (~spl0_20 | ~spl0_24 | ~spl0_86)),\n",
      "  inference(forward_demodulation,[],[f5595,f5625])).\n",
      "fof(f5595,plain,(\n",
      "  m(two,a) = m(f(b,m(b,b)),f(f(b,m(b,b)),m(two,a))) | ~spl0_86),\n",
      "  inference(avatar_component_clause,[],[f5593])).\n",
      "fof(f6042,plain,(\n",
      "  spl0_95 | ~spl0_15 | ~spl0_20 | ~spl0_24),\n",
      "  inference(avatar_split_clause,[],[f6037,f386,f329,f259,f6039])).\n",
      "fof(f6039,plain,(\n",
      "  spl0_95 <=> d(m(b,two),m(two,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_95])])).\n",
      "fof(f259,plain,(\n",
      "  spl0_15 <=> d(m(b,two),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_15])])).\n",
      "fof(f6037,plain,(\n",
      "  d(m(b,two),m(two,two)) | (~spl0_15 | ~spl0_20 | ~spl0_24)),\n",
      "  inference(forward_demodulation,[],[f6036,f5625])).\n",
      "fof(f6036,plain,(\n",
      "  d(m(b,two),m(two,a)) | (~spl0_15 | ~spl0_20 | ~spl0_24)),\n",
      "  inference(forward_demodulation,[],[f261,f5625])).\n",
      "fof(f261,plain,(\n",
      "  d(m(b,two),m(a,a)) | ~spl0_15),\n",
      "  inference(avatar_component_clause,[],[f259])).\n",
      "fof(f6035,plain,(\n",
      "  spl0_94 | ~spl0_16 | ~spl0_20 | ~spl0_24),\n",
      "  inference(avatar_split_clause,[],[f6030,f386,f329,f265,f6032])).\n",
      "fof(f6032,plain,(\n",
      "  spl0_94 <=> m(two,two) = m(m(b,two),f(m(b,two),m(two,two)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_94])])).\n",
      "fof(f265,plain,(\n",
      "  spl0_16 <=> m(a,a) = m(m(b,two),f(m(b,two),m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_16])])).\n",
      "fof(f6030,plain,(\n",
      "  m(two,two) = m(m(b,two),f(m(b,two),m(two,two))) | (~spl0_16 | ~spl0_20 | ~spl0_24)),\n",
      "  inference(forward_demodulation,[],[f6029,f5625])).\n",
      "fof(f6029,plain,(\n",
      "  m(two,a) = m(m(b,two),f(m(b,two),m(two,a))) | (~spl0_16 | ~spl0_20 | ~spl0_24)),\n",
      "  inference(forward_demodulation,[],[f267,f5625])).\n",
      "fof(f267,plain,(\n",
      "  m(a,a) = m(m(b,two),f(m(b,two),m(a,a))) | ~spl0_16),\n",
      "  inference(avatar_component_clause,[],[f265])).\n",
      "fof(f6026,plain,(\n",
      "  spl0_88 | ~spl0_20 | ~spl0_24 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f6025,f1281,f386,f329,f5692])).\n",
      "fof(f5692,plain,(\n",
      "  spl0_88 <=> m(b,m(b,two)) = m(two,two)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_88])])).\n",
      "fof(f6025,plain,(\n",
      "  m(b,m(b,two)) = m(two,two) | (~spl0_20 | ~spl0_24 | ~spl0_36)),\n",
      "  inference(forward_demodulation,[],[f6024,f5625])).\n",
      "fof(f6024,plain,(\n",
      "  m(b,m(b,two)) = m(two,a) | (~spl0_20 | ~spl0_24 | ~spl0_36)),\n",
      "  inference(forward_demodulation,[],[f1283,f5625])).\n",
      "fof(f6011,plain,(\n",
      "  spl0_93 | ~spl0_20 | ~spl0_24 | ~spl0_41),\n",
      "  inference(avatar_split_clause,[],[f6006,f1444,f386,f329,f6008])).\n",
      "fof(f6008,plain,(\n",
      "  spl0_93 <=> d(f(m(b,two),m(two,two)),m(two,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_93])])).\n",
      "fof(f6006,plain,(\n",
      "  d(f(m(b,two),m(two,two)),m(two,two)) | (~spl0_20 | ~spl0_24 | ~spl0_41)),\n",
      "  inference(forward_demodulation,[],[f6005,f5625])).\n",
      "fof(f6005,plain,(\n",
      "  d(f(m(b,two),m(two,a)),m(two,a)) | (~spl0_20 | ~spl0_24 | ~spl0_41)),\n",
      "  inference(forward_demodulation,[],[f1446,f5625])).\n",
      "fof(f5983,plain,(\n",
      "  spl0_92 | ~spl0_20 | ~spl0_24 | ~spl0_59),\n",
      "  inference(avatar_split_clause,[],[f5978,f4668,f386,f329,f5980])).\n",
      "fof(f5980,plain,(\n",
      "  spl0_92 <=> d(f(b,m(b,two)),m(two,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_92])])).\n",
      "fof(f5978,plain,(\n",
      "  d(f(b,m(b,two)),m(two,two)) | (~spl0_20 | ~spl0_24 | ~spl0_59)),\n",
      "  inference(forward_demodulation,[],[f5977,f5625])).\n",
      "fof(f5977,plain,(\n",
      "  d(f(b,m(b,two)),m(two,a)) | (~spl0_20 | ~spl0_24 | ~spl0_59)),\n",
      "  inference(forward_demodulation,[],[f4670,f5625])).\n",
      "fof(f5976,plain,(\n",
      "  spl0_91 | ~spl0_20 | ~spl0_24 | ~spl0_61),\n",
      "  inference(avatar_split_clause,[],[f5971,f4680,f386,f329,f5973])).\n",
      "fof(f5973,plain,(\n",
      "  spl0_91 <=> m(two,two) = m(f(b,m(b,two)),f(f(b,m(b,two)),m(two,two)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_91])])).\n",
      "fof(f5971,plain,(\n",
      "  m(two,two) = m(f(b,m(b,two)),f(f(b,m(b,two)),m(two,two))) | (~spl0_20 | ~spl0_24 | ~spl0_61)),\n",
      "  inference(forward_demodulation,[],[f5970,f5625])).\n",
      "fof(f5970,plain,(\n",
      "  m(two,a) = m(f(b,m(b,two)),f(f(b,m(b,two)),m(two,a))) | (~spl0_20 | ~spl0_24 | ~spl0_61)),\n",
      "  inference(forward_demodulation,[],[f4682,f5625])).\n",
      "fof(f5967,plain,(\n",
      "  spl0_90 | ~spl0_20 | ~spl0_24 | ~spl0_62),\n",
      "  inference(avatar_split_clause,[],[f5962,f4947,f386,f329,f5964])).\n",
      "fof(f5964,plain,(\n",
      "  spl0_90 <=> d(f(two,m(b,two)),m(two,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_90])])).\n",
      "fof(f5962,plain,(\n",
      "  d(f(two,m(b,two)),m(two,two)) | (~spl0_20 | ~spl0_24 | ~spl0_62)),\n",
      "  inference(forward_demodulation,[],[f5961,f5625])).\n",
      "fof(f5961,plain,(\n",
      "  d(f(two,m(b,two)),m(two,a)) | (~spl0_20 | ~spl0_24 | ~spl0_62)),\n",
      "  inference(forward_demodulation,[],[f4949,f5625])).\n",
      "fof(f5960,plain,(\n",
      "  spl0_89 | ~spl0_20 | ~spl0_24 | ~spl0_63),\n",
      "  inference(avatar_split_clause,[],[f5955,f4960,f386,f329,f5957])).\n",
      "fof(f5957,plain,(\n",
      "  spl0_89 <=> m(two,two) = m(f(two,m(b,two)),f(f(two,m(b,two)),m(two,two)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_89])])).\n",
      "fof(f5955,plain,(\n",
      "  m(two,two) = m(f(two,m(b,two)),f(f(two,m(b,two)),m(two,two))) | (~spl0_20 | ~spl0_24 | ~spl0_63)),\n",
      "  inference(forward_demodulation,[],[f5954,f5625])).\n",
      "fof(f5954,plain,(\n",
      "  m(two,a) = m(f(two,m(b,two)),f(f(two,m(b,two)),m(two,a))) | (~spl0_20 | ~spl0_24 | ~spl0_63)),\n",
      "  inference(forward_demodulation,[],[f4962,f5625])).\n",
      "fof(f5695,plain,(\n",
      "  spl0_88 | ~spl0_2 | ~spl0_20 | ~spl0_24),\n",
      "  inference(avatar_split_clause,[],[f5690,f386,f329,f19,f5692])).\n",
      "fof(f5690,plain,(\n",
      "  m(b,m(b,two)) = m(two,two) | (~spl0_2 | ~spl0_20 | ~spl0_24)),\n",
      "  inference(forward_demodulation,[],[f5532,f5625])).\n",
      "fof(f5532,plain,(\n",
      "  m(b,m(b,two)) = m(two,a) | (~spl0_2 | ~spl0_20 | ~spl0_24)),\n",
      "  inference(forward_demodulation,[],[f5531,f4])).\n",
      "fof(f5531,plain,(\n",
      "  m(b,m(two,b)) = m(two,a) | (~spl0_2 | ~spl0_20 | ~spl0_24)),\n",
      "  inference(forward_literal_rewriting,[],[f5450,f3266])).\n",
      "fof(f3266,plain,(\n",
      "  ( ! [X2,X3,X0,X1] : (m(X1,m(X2,X0)) != X3 | m(X2,m(X0,X1)) = X3) )),\n",
      "  inference(superposition,[],[f966,f1086])).\n",
      "fof(f5450,plain,(\n",
      "  m(two,m(b,b)) = m(two,a) | (~spl0_2 | ~spl0_20 | ~spl0_24)),\n",
      "  inference(backward_demodulation,[],[f21,f5321])).\n",
      "fof(f5321,plain,(\n",
      "  ( ! [X0] : (m(two,X0) = m(a,X0)) ) | (~spl0_20 | ~spl0_24)),\n",
      "  inference(forward_demodulation,[],[f5261,f1])).\n",
      "fof(f5261,plain,(\n",
      "  ( ! [X0] : (m(a,X0) = m(two,m(one,X0))) ) | (~spl0_20 | ~spl0_24)),\n",
      "  inference(backward_demodulation,[],[f364,f388])).\n",
      "fof(f364,plain,(\n",
      "  ( ! [X0] : (m(a,X0) = m(two,m(f(two,a),X0))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f3,f331])).\n",
      "fof(f5689,plain,(\n",
      "  spl0_87 | ~spl0_20 | ~spl0_24 | ~spl0_28),\n",
      "  inference(avatar_split_clause,[],[f5684,f458,f386,f329,f5686])).\n",
      "fof(f5684,plain,(\n",
      "  d(m(two,two),m(b,m(b,a))) | (~spl0_20 | ~spl0_24 | ~spl0_28)),\n",
      "  inference(forward_demodulation,[],[f5456,f5625])).\n",
      "fof(f5456,plain,(\n",
      "  d(m(two,a),m(b,m(b,a))) | (~spl0_20 | ~spl0_24 | ~spl0_28)),\n",
      "  inference(backward_demodulation,[],[f460,f5321])).\n",
      "fof(f5596,plain,(\n",
      "  spl0_86 | ~spl0_20 | ~spl0_24 | ~spl0_60),\n",
      "  inference(avatar_split_clause,[],[f5479,f4674,f386,f329,f5593])).\n",
      "fof(f5479,plain,(\n",
      "  m(two,a) = m(f(b,m(b,b)),f(f(b,m(b,b)),m(two,a))) | (~spl0_20 | ~spl0_24 | ~spl0_60)),\n",
      "  inference(backward_demodulation,[],[f4676,f5321])).\n",
      "fof(f5591,plain,(\n",
      "  spl0_85 | ~spl0_20 | ~spl0_24 | ~spl0_58),\n",
      "  inference(avatar_split_clause,[],[f5478,f4663,f386,f329,f5588])).\n",
      "fof(f5478,plain,(\n",
      "  d(f(b,m(b,b)),m(two,a)) | (~spl0_20 | ~spl0_24 | ~spl0_58)),\n",
      "  inference(backward_demodulation,[],[f4665,f5321])).\n",
      "fof(f5586,plain,(\n",
      "  spl0_84 | ~spl0_20 | ~spl0_24 | ~spl0_53),\n",
      "  inference(avatar_split_clause,[],[f5477,f4397,f386,f329,f5583])).\n",
      "fof(f5477,plain,(\n",
      "  m(b,b) = f(two,m(two,a)) | (~spl0_20 | ~spl0_24 | ~spl0_53)),\n",
      "  inference(backward_demodulation,[],[f4399,f5321])).\n",
      "fof(f5579,plain,(\n",
      "  spl0_83 | ~spl0_20 | ~spl0_24 | ~spl0_46),\n",
      "  inference(avatar_split_clause,[],[f5466,f1804,f386,f329,f5576])).\n",
      "fof(f5466,plain,(\n",
      "  m(b,m(b,a)) = m(m(two,a),f(m(two,a),m(b,m(b,a)))) | (~spl0_20 | ~spl0_24 | ~spl0_46)),\n",
      "  inference(backward_demodulation,[],[f1806,f5321])).\n",
      "fof(f5574,plain,(\n",
      "  spl0_82 | ~spl0_20 | ~spl0_24 | ~spl0_43),\n",
      "  inference(avatar_split_clause,[],[f5460,f1513,f386,f329,f5571])).\n",
      "fof(f1513,plain,(\n",
      "  spl0_43 <=> m(a,a) = m(f(b,m(a,a)),f(f(b,m(a,a)),m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_43])])).\n",
      "fof(f5460,plain,(\n",
      "  m(two,a) = m(f(b,m(two,a)),f(f(b,m(two,a)),m(two,a))) | (~spl0_20 | ~spl0_24 | ~spl0_43)),\n",
      "  inference(backward_demodulation,[],[f1515,f5321])).\n",
      "fof(f1515,plain,(\n",
      "  m(a,a) = m(f(b,m(a,a)),f(f(b,m(a,a)),m(a,a))) | ~spl0_43),\n",
      "  inference(avatar_component_clause,[],[f1513])).\n",
      "fof(f5569,plain,(\n",
      "  spl0_81 | ~spl0_20 | ~spl0_24 | ~spl0_40),\n",
      "  inference(avatar_split_clause,[],[f5459,f1439,f386,f329,f5566])).\n",
      "fof(f1439,plain,(\n",
      "  spl0_40 <=> d(f(b,m(a,a)),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_40])])).\n",
      "fof(f5459,plain,(\n",
      "  d(f(b,m(two,a)),m(two,a)) | (~spl0_20 | ~spl0_24 | ~spl0_40)),\n",
      "  inference(backward_demodulation,[],[f1441,f5321])).\n",
      "fof(f1441,plain,(\n",
      "  d(f(b,m(a,a)),m(a,a)) | ~spl0_40),\n",
      "  inference(avatar_component_clause,[],[f1439])).\n",
      "fof(f5564,plain,(\n",
      "  spl0_80 | ~spl0_20 | ~spl0_24 | ~spl0_30),\n",
      "  inference(avatar_split_clause,[],[f5457,f758,f386,f329,f5561])).\n",
      "fof(f758,plain,(\n",
      "  spl0_30 <=> d(f(m(b,b),m(a,a)),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_30])])).\n",
      "fof(f5457,plain,(\n",
      "  d(f(m(b,b),m(two,a)),m(two,a)) | (~spl0_20 | ~spl0_24 | ~spl0_30)),\n",
      "  inference(backward_demodulation,[],[f760,f5321])).\n",
      "fof(f760,plain,(\n",
      "  d(f(m(b,b),m(a,a)),m(a,a)) | ~spl0_30),\n",
      "  inference(avatar_component_clause,[],[f758])).\n",
      "fof(f5559,plain,(\n",
      "  spl0_79 | ~spl0_14 | ~spl0_20 | ~spl0_24),\n",
      "  inference(avatar_split_clause,[],[f5455,f386,f329,f209,f5556])).\n",
      "fof(f209,plain,(\n",
      "  spl0_14 <=> m(a,a) = m(m(b,b),f(m(b,b),m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_14])])).\n",
      "fof(f5455,plain,(\n",
      "  m(two,a) = m(m(b,b),f(m(b,b),m(two,a))) | (~spl0_14 | ~spl0_20 | ~spl0_24)),\n",
      "  inference(backward_demodulation,[],[f211,f5321])).\n",
      "fof(f211,plain,(\n",
      "  m(a,a) = m(m(b,b),f(m(b,b),m(a,a))) | ~spl0_14),\n",
      "  inference(avatar_component_clause,[],[f209])).\n",
      "fof(f5554,plain,(\n",
      "  spl0_78 | ~spl0_13 | ~spl0_20 | ~spl0_24),\n",
      "  inference(avatar_split_clause,[],[f5454,f386,f329,f189,f5551])).\n",
      "fof(f189,plain,(\n",
      "  spl0_13 <=> m(a,a) = m(b,f(b,m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_13])])).\n",
      "fof(f5454,plain,(\n",
      "  m(two,a) = m(b,f(b,m(two,a))) | (~spl0_13 | ~spl0_20 | ~spl0_24)),\n",
      "  inference(backward_demodulation,[],[f191,f5321])).\n",
      "fof(f191,plain,(\n",
      "  m(a,a) = m(b,f(b,m(a,a))) | ~spl0_13),\n",
      "  inference(avatar_component_clause,[],[f189])).\n",
      "fof(f5549,plain,(\n",
      "  spl0_77 | ~spl0_12 | ~spl0_20 | ~spl0_24),\n",
      "  inference(avatar_split_clause,[],[f5453,f386,f329,f183,f5546])).\n",
      "fof(f183,plain,(\n",
      "  spl0_12 <=> d(b,m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_12])])).\n",
      "fof(f5453,plain,(\n",
      "  d(b,m(two,a)) | (~spl0_12 | ~spl0_20 | ~spl0_24)),\n",
      "  inference(backward_demodulation,[],[f185,f5321])).\n",
      "fof(f185,plain,(\n",
      "  d(b,m(a,a)) | ~spl0_12),\n",
      "  inference(avatar_component_clause,[],[f183])).\n",
      "fof(f5544,plain,(\n",
      "  spl0_76 | ~spl0_7 | ~spl0_20 | ~spl0_24),\n",
      "  inference(avatar_split_clause,[],[f5452,f386,f329,f120,f5541])).\n",
      "fof(f5452,plain,(\n",
      "  d(m(b,b),m(two,a)) | (~spl0_7 | ~spl0_20 | ~spl0_24)),\n",
      "  inference(backward_demodulation,[],[f122,f5321])).\n",
      "fof(f122,plain,(\n",
      "  d(m(b,b),m(a,a)) | ~spl0_7),\n",
      "  inference(avatar_component_clause,[],[f120])).\n",
      "fof(f5539,plain,(\n",
      "  spl0_75 | ~spl0_2 | ~spl0_20 | ~spl0_24 | ~spl0_37),\n",
      "  inference(avatar_split_clause,[],[f5534,f1402,f386,f329,f19,f5536])).\n",
      "fof(f5536,plain,(\n",
      "  spl0_75 <=> b = m(two,a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_75])])).\n",
      "fof(f5534,plain,(\n",
      "  b = m(two,a) | (~spl0_2 | ~spl0_20 | ~spl0_24 | ~spl0_37)),\n",
      "  inference(forward_demodulation,[],[f5533,f2])).\n",
      "fof(f5533,plain,(\n",
      "  m(two,a) = m(b,one) | (~spl0_2 | ~spl0_20 | ~spl0_24 | ~spl0_37)),\n",
      "  inference(forward_demodulation,[],[f5532,f1404])).\n",
      "fof(f1404,plain,(\n",
      "  one = m(b,two) | ~spl0_37),\n",
      "  inference(avatar_component_clause,[],[f1402])).\n",
      "fof(f5320,plain,(\n",
      "  spl0_25 | ~spl0_20 | ~spl0_24),\n",
      "  inference(avatar_split_clause,[],[f5319,f386,f329,f390])).\n",
      "fof(f5319,plain,(\n",
      "  two = a | (~spl0_20 | ~spl0_24)),\n",
      "  inference(forward_demodulation,[],[f5260,f2])).\n",
      "fof(f5260,plain,(\n",
      "  a = m(two,one) | (~spl0_20 | ~spl0_24)),\n",
      "  inference(backward_demodulation,[],[f331,f388])).\n",
      "fof(f5259,plain,(\n",
      "  spl0_74 | ~spl0_37 | ~spl0_63),\n",
      "  inference(avatar_split_clause,[],[f5057,f4960,f1402,f5256])).\n",
      "fof(f5256,plain,(\n",
      "  spl0_74 <=> m(a,a) = m(f(two,one),f(f(two,one),m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_74])])).\n",
      "fof(f5057,plain,(\n",
      "  m(a,a) = m(f(two,one),f(f(two,one),m(a,a))) | (~spl0_37 | ~spl0_63)),\n",
      "  inference(backward_demodulation,[],[f4962,f1404])).\n",
      "fof(f5254,plain,(\n",
      "  spl0_73 | ~spl0_37 | ~spl0_62),\n",
      "  inference(avatar_split_clause,[],[f5056,f4947,f1402,f5251])).\n",
      "fof(f5251,plain,(\n",
      "  spl0_73 <=> d(f(two,one),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_73])])).\n",
      "fof(f5056,plain,(\n",
      "  d(f(two,one),m(a,a)) | (~spl0_37 | ~spl0_62)),\n",
      "  inference(backward_demodulation,[],[f4949,f1404])).\n",
      "fof(f5247,plain,(\n",
      "  spl0_72 | ~spl0_37 | ~spl0_61),\n",
      "  inference(avatar_split_clause,[],[f5054,f4680,f1402,f5244])).\n",
      "fof(f5244,plain,(\n",
      "  spl0_72 <=> m(a,a) = m(f(b,one),f(f(b,one),m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_72])])).\n",
      "fof(f5054,plain,(\n",
      "  m(a,a) = m(f(b,one),f(f(b,one),m(a,a))) | (~spl0_37 | ~spl0_61)),\n",
      "  inference(backward_demodulation,[],[f4682,f1404])).\n",
      "fof(f5242,plain,(\n",
      "  spl0_71 | ~spl0_37 | ~spl0_59),\n",
      "  inference(avatar_split_clause,[],[f5053,f4668,f1402,f5239])).\n",
      "fof(f5239,plain,(\n",
      "  spl0_71 <=> d(f(b,one),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_71])])).\n",
      "fof(f5053,plain,(\n",
      "  d(f(b,one),m(a,a)) | (~spl0_37 | ~spl0_59)),\n",
      "  inference(backward_demodulation,[],[f4670,f1404])).\n",
      "fof(f5227,plain,(\n",
      "  spl0_70 | ~spl0_36 | ~spl0_37 | ~spl0_48),\n",
      "  inference(avatar_split_clause,[],[f5159,f3772,f1402,f1281,f5224])).\n",
      "fof(f5224,plain,(\n",
      "  spl0_70 <=> d(m(b,m(two,two)),m(b,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_70])])).\n",
      "fof(f3772,plain,(\n",
      "  spl0_48 <=> d(m(b,m(two,two)),m(a,m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_48])])).\n",
      "fof(f5159,plain,(\n",
      "  d(m(b,m(two,two)),m(b,a)) | (~spl0_36 | ~spl0_37 | ~spl0_48)),\n",
      "  inference(backward_demodulation,[],[f3774,f5070])).\n",
      "fof(f5070,plain,(\n",
      "  ( ! [X0] : (m(a,m(a,X0)) = m(b,X0)) ) | (~spl0_36 | ~spl0_37)),\n",
      "  inference(forward_demodulation,[],[f5038,f2])).\n",
      "fof(f5038,plain,(\n",
      "  ( ! [X0] : (m(a,m(a,X0)) = m(b,m(X0,one))) ) | (~spl0_36 | ~spl0_37)),\n",
      "  inference(backward_demodulation,[],[f1418,f1404])).\n",
      "fof(f1418,plain,(\n",
      "  ( ! [X0] : (m(a,m(a,X0)) = m(b,m(X0,m(b,two)))) ) | ~spl0_36),\n",
      "  inference(forward_literal_rewriting,[],[f1417,f1001])).\n",
      "fof(f1417,plain,(\n",
      "  ( ! [X0] : (m(a,m(a,X0)) = m(m(b,two),m(b,X0))) ) | ~spl0_36),\n",
      "  inference(forward_demodulation,[],[f1372,f3])).\n",
      "fof(f1372,plain,(\n",
      "  ( ! [X0] : (m(m(a,a),X0) = m(m(b,two),m(b,X0))) ) | ~spl0_36),\n",
      "  inference(superposition,[],[f145,f1283])).\n",
      "fof(f3774,plain,(\n",
      "  d(m(b,m(two,two)),m(a,m(a,a))) | ~spl0_48),\n",
      "  inference(avatar_component_clause,[],[f3772])).\n",
      "fof(f5222,plain,(\n",
      "  spl0_69 | ~spl0_36 | ~spl0_37 | ~spl0_49),\n",
      "  inference(avatar_split_clause,[],[f5158,f3963,f1402,f1281,f5219])).\n",
      "fof(f5219,plain,(\n",
      "  spl0_69 <=> d(m(two,two),m(b,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_69])])).\n",
      "fof(f3963,plain,(\n",
      "  spl0_49 <=> d(m(two,two),m(a,m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_49])])).\n",
      "fof(f5158,plain,(\n",
      "  d(m(two,two),m(b,a)) | (~spl0_36 | ~spl0_37 | ~spl0_49)),\n",
      "  inference(backward_demodulation,[],[f3965,f5070])).\n",
      "fof(f3965,plain,(\n",
      "  d(m(two,two),m(a,m(a,a))) | ~spl0_49),\n",
      "  inference(avatar_component_clause,[],[f3963])).\n",
      "fof(f5217,plain,(\n",
      "  spl0_68 | ~spl0_36 | ~spl0_37 | ~spl0_52),\n",
      "  inference(avatar_split_clause,[],[f5157,f4238,f1402,f1281,f5214])).\n",
      "fof(f5214,plain,(\n",
      "  spl0_68 <=> m(b,a) = m(m(two,two),f(m(two,two),m(b,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_68])])).\n",
      "fof(f4238,plain,(\n",
      "  spl0_52 <=> m(a,m(a,a)) = m(m(two,two),f(m(two,two),m(a,m(a,a))))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_52])])).\n",
      "fof(f5157,plain,(\n",
      "  m(b,a) = m(m(two,two),f(m(two,two),m(b,a))) | (~spl0_36 | ~spl0_37 | ~spl0_52)),\n",
      "  inference(backward_demodulation,[],[f4240,f5070])).\n",
      "fof(f4240,plain,(\n",
      "  m(a,m(a,a)) = m(m(two,two),f(m(two,two),m(a,m(a,a)))) | ~spl0_52),\n",
      "  inference(avatar_component_clause,[],[f4238])).\n",
      "fof(f5069,plain,(\n",
      "  ~spl0_67 | ~spl0_37 | spl0_39),\n",
      "  inference(avatar_split_clause,[],[f5037,f1411,f1402,f5066])).\n",
      "fof(f5037,plain,(\n",
      "  one != m(a,a) | (~spl0_37 | spl0_39)),\n",
      "  inference(backward_demodulation,[],[f1413,f1404])).\n",
      "fof(f5062,plain,(\n",
      "  spl0_38 | ~spl0_36 | ~spl0_37),\n",
      "  inference(avatar_split_clause,[],[f5061,f1402,f1281,f1406])).\n",
      "fof(f5061,plain,(\n",
      "  m(a,a) = b | (~spl0_36 | ~spl0_37)),\n",
      "  inference(forward_demodulation,[],[f5030,f2])).\n",
      "fof(f5030,plain,(\n",
      "  m(a,a) = m(b,one) | (~spl0_36 | ~spl0_37)),\n",
      "  inference(backward_demodulation,[],[f1283,f1404])).\n",
      "fof(f5020,plain,(\n",
      "  spl0_37 | spl0_66 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f4993,f1281,f5018,f1402])).\n",
      "fof(f4993,plain,(\n",
      "  ( ! [X0] : (m(X0,m(a,a)) != m(X0,b) | one = m(b,two)) ) | ~spl0_36),\n",
      "  inference(superposition,[],[f161,f1283])).\n",
      "fof(f5016,plain,(\n",
      "  spl0_24 | spl0_65 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f4992,f329,f5014,f386])).\n",
      "fof(f4992,plain,(\n",
      "  ( ! [X0] : (m(X0,two) != m(X0,a) | one = f(two,a)) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f161,f331])).\n",
      "fof(f5012,plain,(\n",
      "  spl0_8 | spl0_64 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f4990,f19,f5010,f125])).\n",
      "fof(f125,plain,(\n",
      "  spl0_8 <=> one = m(b,b)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_8])])).\n",
      "fof(f4990,plain,(\n",
      "  ( ! [X0] : (m(X0,two) != m(X0,m(a,a)) | one = m(b,b)) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f161,f21])).\n",
      "fof(f4963,plain,(\n",
      "  spl0_63 | ~spl0_62),\n",
      "  inference(avatar_split_clause,[],[f4958,f4947,f4960])).\n",
      "fof(f4958,plain,(\n",
      "  m(a,a) = m(f(two,m(b,two)),f(f(two,m(b,two)),m(a,a))) | ~spl0_62),\n",
      "  inference(resolution,[],[f4949,f6])).\n",
      "fof(f4950,plain,(\n",
      "  spl0_62 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f4940,f1281,f4947])).\n",
      "fof(f4940,plain,(\n",
      "  d(f(two,m(b,two)),m(a,a)) | ~spl0_36),\n",
      "  inference(superposition,[],[f2266,f1283])).\n",
      "fof(f4683,plain,(\n",
      "  spl0_61 | ~spl0_59),\n",
      "  inference(avatar_split_clause,[],[f4678,f4668,f4680])).\n",
      "fof(f4678,plain,(\n",
      "  m(a,a) = m(f(b,m(b,two)),f(f(b,m(b,two)),m(a,a))) | ~spl0_59),\n",
      "  inference(resolution,[],[f4670,f6])).\n",
      "fof(f4677,plain,(\n",
      "  spl0_60 | ~spl0_58),\n",
      "  inference(avatar_split_clause,[],[f4672,f4663,f4674])).\n",
      "fof(f4672,plain,(\n",
      "  m(a,a) = m(f(b,m(b,b)),f(f(b,m(b,b)),m(a,a))) | ~spl0_58),\n",
      "  inference(resolution,[],[f4665,f6])).\n",
      "fof(f4671,plain,(\n",
      "  spl0_59 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f4655,f1281,f4668])).\n",
      "fof(f4655,plain,(\n",
      "  d(f(b,m(b,two)),m(a,a)) | ~spl0_36),\n",
      "  inference(superposition,[],[f650,f1283])).\n",
      "fof(f4666,plain,(\n",
      "  spl0_58 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f4653,f19,f4663])).\n",
      "fof(f4653,plain,(\n",
      "  d(f(b,m(b,b)),m(a,a)) | ~spl0_2),\n",
      "  inference(superposition,[],[f650,f21])).\n",
      "fof(f4442,plain,(\n",
      "  spl0_57 | ~spl0_55 | ~spl0_11),\n",
      "  inference(avatar_split_clause,[],[f4382,f140,f4412,f4439])).\n",
      "fof(f4439,plain,(\n",
      "  spl0_57 <=> two = f(f(two,m(a,a)),f(two,m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_57])])).\n",
      "fof(f4412,plain,(\n",
      "  spl0_55 <=> m(a,a) = f(two,m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_55])])).\n",
      "fof(f140,plain,(\n",
      "  spl0_11 <=> m(a,a) = m(two,f(two,m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_11])])).\n",
      "fof(f4382,plain,(\n",
      "  m(a,a) != f(two,m(a,a)) | two = f(f(two,m(a,a)),f(two,m(a,a))) | ~spl0_11),\n",
      "  inference(superposition,[],[f833,f142])).\n",
      "fof(f142,plain,(\n",
      "  m(a,a) = m(two,f(two,m(a,a))) | ~spl0_11),\n",
      "  inference(avatar_component_clause,[],[f140])).\n",
      "fof(f4437,plain,(\n",
      "  spl0_56 | ~spl0_11),\n",
      "  inference(avatar_split_clause,[],[f4381,f140,f4434])).\n",
      "fof(f4434,plain,(\n",
      "  spl0_56 <=> d(f(f(two,m(a,a)),m(a,a)),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_56])])).\n",
      "fof(f4381,plain,(\n",
      "  d(f(f(two,m(a,a)),m(a,a)),m(a,a)) | ~spl0_11),\n",
      "  inference(superposition,[],[f691,f142])).\n",
      "fof(f4415,plain,(\n",
      "  spl0_1 | ~spl0_55 | ~spl0_11),\n",
      "  inference(avatar_split_clause,[],[f4350,f140,f4412,f14])).\n",
      "fof(f4350,plain,(\n",
      "  m(a,a) != f(two,m(a,a)) | one = two | ~spl0_11),\n",
      "  inference(superposition,[],[f65,f142])).\n",
      "fof(f4406,plain,(\n",
      "  spl0_54 | ~spl0_11 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f4401,f329,f140,f4403])).\n",
      "fof(f4403,plain,(\n",
      "  spl0_54 <=> d(m(a,a),m(a,f(two,m(a,a))))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_54])])).\n",
      "fof(f4401,plain,(\n",
      "  d(m(a,a),m(a,f(two,m(a,a)))) | (~spl0_11 | ~spl0_20)),\n",
      "  inference(forward_demodulation,[],[f4332,f4])).\n",
      "fof(f4332,plain,(\n",
      "  d(m(a,a),m(f(two,m(a,a)),a)) | (~spl0_11 | ~spl0_20)),\n",
      "  inference(superposition,[],[f376,f142])).\n",
      "fof(f376,plain,(\n",
      "  ( ! [X0] : (d(m(two,X0),m(X0,a))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f231,f331])).\n",
      "fof(f4400,plain,(\n",
      "  spl0_53 | ~spl0_2 | ~spl0_11),\n",
      "  inference(avatar_split_clause,[],[f4395,f140,f19,f4397])).\n",
      "fof(f4395,plain,(\n",
      "  m(b,b) = f(two,m(a,a)) | (~spl0_2 | ~spl0_11)),\n",
      "  inference(trivial_inequality_removal,[],[f4330])).\n",
      "fof(f4330,plain,(\n",
      "  m(a,a) != m(a,a) | m(b,b) = f(two,m(a,a)) | (~spl0_2 | ~spl0_11)),\n",
      "  inference(superposition,[],[f108,f142])).\n",
      "fof(f4241,plain,(\n",
      "  spl0_52 | ~spl0_49),\n",
      "  inference(avatar_split_clause,[],[f4236,f3963,f4238])).\n",
      "fof(f4236,plain,(\n",
      "  m(a,m(a,a)) = m(m(two,two),f(m(two,two),m(a,m(a,a)))) | ~spl0_49),\n",
      "  inference(resolution,[],[f3965,f6])).\n",
      "fof(f3977,plain,(\n",
      "  spl0_51 | ~spl0_50),\n",
      "  inference(avatar_split_clause,[],[f3972,f3968,f3974])).\n",
      "fof(f3974,plain,(\n",
      "  spl0_51 <=> m(a,a) = m(m(two,two),f(m(two,two),m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_51])])).\n",
      "fof(f3968,plain,(\n",
      "  spl0_50 <=> d(m(two,two),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_50])])).\n",
      "fof(f3972,plain,(\n",
      "  m(a,a) = m(m(two,two),f(m(two,two),m(a,a))) | ~spl0_50),\n",
      "  inference(resolution,[],[f3970,f6])).\n",
      "fof(f3970,plain,(\n",
      "  d(m(two,two),m(a,a)) | ~spl0_50),\n",
      "  inference(avatar_component_clause,[],[f3968])).\n",
      "fof(f3971,plain,(\n",
      "  spl0_50 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f3946,f329,f3968])).\n",
      "fof(f3946,plain,(\n",
      "  d(m(two,two),m(a,a)) | ~spl0_20),\n",
      "  inference(superposition,[],[f3592,f331])).\n",
      "fof(f3966,plain,(\n",
      "  spl0_49 | ~spl0_2 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f3944,f329,f19,f3963])).\n",
      "fof(f3944,plain,(\n",
      "  d(m(two,two),m(a,m(a,a))) | (~spl0_2 | ~spl0_20)),\n",
      "  inference(superposition,[],[f3592,f21])).\n",
      "fof(f3775,plain,(\n",
      "  spl0_48 | ~spl0_20 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f3770,f1281,f329,f3772])).\n",
      "fof(f3770,plain,(\n",
      "  d(m(b,m(two,two)),m(a,m(a,a))) | (~spl0_20 | ~spl0_36)),\n",
      "  inference(forward_demodulation,[],[f3769,f151])).\n",
      "fof(f3769,plain,(\n",
      "  d(m(two,m(b,two)),m(a,m(a,a))) | (~spl0_20 | ~spl0_36)),\n",
      "  inference(forward_demodulation,[],[f3741,f4])).\n",
      "fof(f3741,plain,(\n",
      "  d(m(m(b,two),two),m(a,m(a,a))) | (~spl0_20 | ~spl0_36)),\n",
      "  inference(superposition,[],[f3587,f1283])).\n",
      "fof(f3587,plain,(\n",
      "  ( ! [X0,X1] : (d(m(X0,two),m(a,m(X1,X0)))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f3506,f885])).\n",
      "fof(f2104,plain,(\n",
      "  spl0_1 | spl0_47 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f2011,f19,f2102,f14])).\n",
      "fof(f2102,plain,(\n",
      "  spl0_47 <=> ! [X0] : m(a,m(a,X0)) != m(b,m(b,X0))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_47])])).\n",
      "fof(f2011,plain,(\n",
      "  ( ! [X0] : (m(a,m(a,X0)) != m(b,m(b,X0)) | one = two) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f65,f166])).\n",
      "fof(f1807,plain,(\n",
      "  spl0_46 | ~spl0_28),\n",
      "  inference(avatar_split_clause,[],[f1802,f458,f1804])).\n",
      "fof(f1802,plain,(\n",
      "  m(b,m(b,a)) = m(m(a,a),f(m(a,a),m(b,m(b,a)))) | ~spl0_28),\n",
      "  inference(resolution,[],[f460,f6])).\n",
      "fof(f1569,plain,(\n",
      "  spl0_44 | ~spl0_45 | ~spl0_2 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f1555,f329,f19,f1566,f1562])).\n",
      "fof(f1555,plain,(\n",
      "  a != m(a,a) | m(b,b) = f(two,a) | (~spl0_2 | ~spl0_20)),\n",
      "  inference(superposition,[],[f108,f331])).\n",
      "fof(f1516,plain,(\n",
      "  spl0_43 | ~spl0_40),\n",
      "  inference(avatar_split_clause,[],[f1511,f1439,f1513])).\n",
      "fof(f1511,plain,(\n",
      "  m(a,a) = m(f(b,m(a,a)),f(f(b,m(a,a)),m(a,a))) | ~spl0_40),\n",
      "  inference(resolution,[],[f1441,f6])).\n",
      "fof(f1452,plain,(\n",
      "  spl0_42 | ~spl0_39 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f1395,f1281,f1411,f1449])).\n",
      "fof(f1449,plain,(\n",
      "  spl0_42 <=> b = f(m(b,two),m(b,two))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_42])])).\n",
      "fof(f1395,plain,(\n",
      "  m(a,a) != m(b,two) | b = f(m(b,two),m(b,two)) | ~spl0_36),\n",
      "  inference(superposition,[],[f833,f1283])).\n",
      "fof(f1447,plain,(\n",
      "  spl0_41 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f1394,f1281,f1444])).\n",
      "fof(f1394,plain,(\n",
      "  d(f(m(b,two),m(a,a)),m(a,a)) | ~spl0_36),\n",
      "  inference(superposition,[],[f691,f1283])).\n",
      "fof(f1442,plain,(\n",
      "  spl0_40 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f1393,f1281,f1439])).\n",
      "fof(f1393,plain,(\n",
      "  d(f(b,m(a,a)),m(a,a)) | ~spl0_36),\n",
      "  inference(superposition,[],[f643,f1283])).\n",
      "fof(f1416,plain,(\n",
      "  ~spl0_38 | spl0_37 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f1415,f1281,f1402,f1406])).\n",
      "fof(f1415,plain,(\n",
      "  one = m(b,two) | m(a,a) != b | ~spl0_36),\n",
      "  inference(forward_demodulation,[],[f1371,f79])).\n",
      "fof(f1371,plain,(\n",
      "  m(a,a) != b | m(b,two) = f(b,b) | ~spl0_36),\n",
      "  inference(superposition,[],[f74,f1283])).\n",
      "fof(f1414,plain,(\n",
      "  spl0_3 | ~spl0_39 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f1369,f1281,f1411,f42])).\n",
      "fof(f1369,plain,(\n",
      "  m(a,a) != m(b,two) | one = b | ~spl0_36),\n",
      "  inference(superposition,[],[f65,f1283])).\n",
      "fof(f1409,plain,(\n",
      "  spl0_37 | ~spl0_38 | ~spl0_36),\n",
      "  inference(avatar_split_clause,[],[f1367,f1281,f1406,f1402])).\n",
      "fof(f1367,plain,(\n",
      "  m(a,a) != b | one = m(b,two) | ~spl0_36),\n",
      "  inference(superposition,[],[f51,f1283])).\n",
      "fof(f1284,plain,(\n",
      "  spl0_36 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f1279,f19,f1281])).\n",
      "fof(f1279,plain,(\n",
      "  m(a,a) = m(b,m(b,two)) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f1091,f4])).\n",
      "fof(f1091,plain,(\n",
      "  m(a,a) = m(b,m(two,b)) | ~spl0_2),\n",
      "  inference(superposition,[],[f151,f21])).\n",
      "fof(f1022,plain,(\n",
      "  spl0_35 | ~spl0_22),\n",
      "  inference(avatar_split_clause,[],[f1017,f360,f1019])).\n",
      "fof(f1019,plain,(\n",
      "  spl0_35 <=> m(a,a) = m(m(b,m(b,two)),f(m(b,m(b,two)),m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_35])])).\n",
      "fof(f360,plain,(\n",
      "  spl0_22 <=> d(m(b,m(b,two)),m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_22])])).\n",
      "fof(f1017,plain,(\n",
      "  m(a,a) = m(m(b,m(b,two)),f(m(b,m(b,two)),m(a,a))) | ~spl0_22),\n",
      "  inference(resolution,[],[f362,f6])).\n",
      "fof(f362,plain,(\n",
      "  d(m(b,m(b,two)),m(a,a)) | ~spl0_22),\n",
      "  inference(avatar_component_clause,[],[f360])).\n",
      "fof(f1016,plain,(\n",
      "  spl0_34 | ~spl0_21),\n",
      "  inference(avatar_split_clause,[],[f1011,f350,f1013])).\n",
      "fof(f1013,plain,(\n",
      "  spl0_34 <=> m(b,m(b,two)) = m(m(a,a),f(m(a,a),m(b,m(b,two))))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_34])])).\n",
      "fof(f350,plain,(\n",
      "  spl0_21 <=> d(m(a,a),m(b,m(b,two)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_21])])).\n",
      "fof(f1011,plain,(\n",
      "  m(b,m(b,two)) = m(m(a,a),f(m(a,a),m(b,m(b,two)))) | ~spl0_21),\n",
      "  inference(resolution,[],[f352,f6])).\n",
      "fof(f352,plain,(\n",
      "  d(m(a,a),m(b,m(b,two))) | ~spl0_21),\n",
      "  inference(avatar_component_clause,[],[f350])).\n",
      "fof(f981,plain,(\n",
      "  spl0_33 | ~spl0_29),\n",
      "  inference(avatar_split_clause,[],[f976,f701,f978])).\n",
      "fof(f978,plain,(\n",
      "  spl0_33 <=> m(a,a) = m(f(two,m(a,a)),f(f(two,m(a,a)),m(a,a)))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_33])])).\n",
      "fof(f976,plain,(\n",
      "  m(a,a) = m(f(two,m(a,a)),f(f(two,m(a,a)),m(a,a))) | ~spl0_29),\n",
      "  inference(resolution,[],[f703,f6])).\n",
      "fof(f772,plain,(\n",
      "  spl0_32 | ~spl0_31),\n",
      "  inference(avatar_split_clause,[],[f767,f763,f769])).\n",
      "fof(f769,plain,(\n",
      "  spl0_32 <=> a = m(f(f(two,a),a),f(f(f(two,a),a),a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_32])])).\n",
      "fof(f763,plain,(\n",
      "  spl0_31 <=> d(f(f(two,a),a),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_31])])).\n",
      "fof(f767,plain,(\n",
      "  a = m(f(f(two,a),a),f(f(f(two,a),a),a)) | ~spl0_31),\n",
      "  inference(resolution,[],[f765,f6])).\n",
      "fof(f765,plain,(\n",
      "  d(f(f(two,a),a),a) | ~spl0_31),\n",
      "  inference(avatar_component_clause,[],[f763])).\n",
      "fof(f766,plain,(\n",
      "  spl0_31 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f754,f329,f763])).\n",
      "fof(f754,plain,(\n",
      "  d(f(f(two,a),a),a) | ~spl0_20),\n",
      "  inference(superposition,[],[f691,f331])).\n",
      "fof(f761,plain,(\n",
      "  spl0_30 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f753,f19,f758])).\n",
      "fof(f753,plain,(\n",
      "  d(f(m(b,b),m(a,a)),m(a,a)) | ~spl0_2),\n",
      "  inference(superposition,[],[f691,f21])).\n",
      "fof(f704,plain,(\n",
      "  spl0_29 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f696,f19,f701])).\n",
      "fof(f696,plain,(\n",
      "  d(f(two,m(a,a)),m(a,a)) | ~spl0_2),\n",
      "  inference(superposition,[],[f643,f21])).\n",
      "fof(f461,plain,(\n",
      "  spl0_28 | ~spl0_2 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f456,f329,f19,f458])).\n",
      "fof(f456,plain,(\n",
      "  d(m(a,a),m(b,m(b,a))) | (~spl0_2 | ~spl0_20)),\n",
      "  inference(forward_demodulation,[],[f455,f151])).\n",
      "fof(f455,plain,(\n",
      "  d(m(a,a),m(b,m(a,b))) | (~spl0_2 | ~spl0_20)),\n",
      "  inference(forward_demodulation,[],[f454,f151])).\n",
      "fof(f454,plain,(\n",
      "  d(m(a,a),m(a,m(b,b))) | (~spl0_2 | ~spl0_20)),\n",
      "  inference(forward_demodulation,[],[f445,f4])).\n",
      "fof(f445,plain,(\n",
      "  d(m(a,a),m(m(b,b),a)) | (~spl0_2 | ~spl0_20)),\n",
      "  inference(superposition,[],[f376,f21])).\n",
      "fof(f419,plain,(\n",
      "  spl0_18 | ~spl0_19),\n",
      "  inference(avatar_split_clause,[],[f414,f324,f320])).\n",
      "fof(f324,plain,(\n",
      "  spl0_19 <=> ! [X0] : d(two,m(X0,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_19])])).\n",
      "fof(f414,plain,(\n",
      "  ( ! [X0] : (d(two,m(a,X0))) ) | ~spl0_19),\n",
      "  inference(superposition,[],[f325,f4])).\n",
      "fof(f325,plain,(\n",
      "  ( ! [X0] : (d(two,m(X0,a))) ) | ~spl0_19),\n",
      "  inference(avatar_component_clause,[],[f324])).\n",
      "fof(f411,plain,(\n",
      "  spl0_27 | ~spl0_23),\n",
      "  inference(avatar_split_clause,[],[f406,f381,f408])).\n",
      "fof(f408,plain,(\n",
      "  spl0_27 <=> a = m(f(two,a),f(f(two,a),a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_27])])).\n",
      "fof(f381,plain,(\n",
      "  spl0_23 <=> d(f(two,a),a)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_23])])).\n",
      "fof(f406,plain,(\n",
      "  a = m(f(two,a),f(f(two,a),a)) | ~spl0_23),\n",
      "  inference(resolution,[],[f383,f6])).\n",
      "fof(f383,plain,(\n",
      "  d(f(two,a),a) | ~spl0_23),\n",
      "  inference(avatar_component_clause,[],[f381])).\n",
      "fof(f400,plain,(\n",
      "  spl0_19 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f375,f329,f324])).\n",
      "fof(f375,plain,(\n",
      "  ( ! [X0] : (d(two,m(X0,a))) ) | ~spl0_20),\n",
      "  inference(superposition,[],[f170,f331])).\n",
      "fof(f170,plain,(\n",
      "  ( ! [X2,X0,X1] : (d(X1,m(X2,m(X1,X0)))) )),\n",
      "  inference(superposition,[],[f160,f4])).\n",
      "fof(f398,plain,(\n",
      "  spl0_1 | ~spl0_26 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f371,f329,f395,f14])).\n",
      "fof(f371,plain,(\n",
      "  a != f(two,a) | one = two | ~spl0_20),\n",
      "  inference(superposition,[],[f65,f331])).\n",
      "fof(f393,plain,(\n",
      "  spl0_24 | ~spl0_25 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f370,f329,f390,f386])).\n",
      "fof(f370,plain,(\n",
      "  two != a | one = f(two,a) | ~spl0_20),\n",
      "  inference(superposition,[],[f51,f331])).\n",
      "fof(f384,plain,(\n",
      "  spl0_23 | ~spl0_20),\n",
      "  inference(avatar_split_clause,[],[f369,f329,f381])).\n",
      "fof(f369,plain,(\n",
      "  d(f(two,a),a) | ~spl0_20),\n",
      "  inference(superposition,[],[f31,f331])).\n",
      "fof(f363,plain,(\n",
      "  spl0_22 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f358,f19,f360])).\n",
      "fof(f358,plain,(\n",
      "  d(m(b,m(b,two)),m(a,a)) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f357,f151])).\n",
      "fof(f357,plain,(\n",
      "  d(m(b,m(two,b)),m(a,a)) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f356,f151])).\n",
      "fof(f356,plain,(\n",
      "  d(m(two,m(b,b)),m(a,a)) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f345,f4])).\n",
      "fof(f345,plain,(\n",
      "  d(m(m(b,b),two),m(a,a)) | ~spl0_2),\n",
      "  inference(superposition,[],[f276,f21])).\n",
      "fof(f353,plain,(\n",
      "  spl0_21 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f348,f19,f350])).\n",
      "fof(f348,plain,(\n",
      "  d(m(a,a),m(b,m(b,two))) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f347,f151])).\n",
      "fof(f347,plain,(\n",
      "  d(m(a,a),m(b,m(two,b))) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f346,f151])).\n",
      "fof(f346,plain,(\n",
      "  d(m(a,a),m(two,m(b,b))) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f339,f4])).\n",
      "fof(f339,plain,(\n",
      "  d(m(a,a),m(m(b,b),two)) | ~spl0_2),\n",
      "  inference(superposition,[],[f276,f21])).\n",
      "fof(f332,plain,(\n",
      "  spl0_20 | ~spl0_17),\n",
      "  inference(avatar_split_clause,[],[f327,f315,f329])).\n",
      "fof(f327,plain,(\n",
      "  a = m(two,f(two,a)) | ~spl0_17),\n",
      "  inference(resolution,[],[f317,f6])).\n",
      "fof(f317,plain,(\n",
      "  d(two,a) | ~spl0_17),\n",
      "  inference(avatar_component_clause,[],[f315])).\n",
      "fof(f326,plain,(\n",
      "  spl0_19 | spl0_17 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f296,f19,f315,f324])).\n",
      "fof(f296,plain,(\n",
      "  ( ! [X0] : (d(two,a) | d(two,m(X0,a))) ) | ~spl0_2),\n",
      "  inference(resolution,[],[f8,f222])).\n",
      "fof(f222,plain,(\n",
      "  ( ! [X0] : (d(two,m(a,m(X0,a)))) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f218,f4])).\n",
      "fof(f218,plain,(\n",
      "  ( ! [X0] : (d(two,m(a,m(a,X0)))) ) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f214,f3])).\n",
      "fof(f214,plain,(\n",
      "  ( ! [X0] : (d(two,m(m(a,a),X0))) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f199,f4])).\n",
      "fof(f199,plain,(\n",
      "  ( ! [X0] : (d(two,m(X0,m(a,a)))) ) | ~spl0_2),\n",
      "  inference(superposition,[],[f170,f21])).\n",
      "fof(f322,plain,(\n",
      "  spl0_18 | spl0_17 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f295,f19,f315,f320])).\n",
      "fof(f295,plain,(\n",
      "  ( ! [X0] : (d(two,a) | d(two,m(a,X0))) ) | ~spl0_2),\n",
      "  inference(resolution,[],[f8,f218])).\n",
      "fof(f318,plain,(\n",
      "  spl0_17 | ~spl0_6),\n",
      "  inference(avatar_split_clause,[],[f313,f115,f315])).\n",
      "fof(f115,plain,(\n",
      "  spl0_6 <=> d(two,m(a,a))),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_6])])).\n",
      "fof(f313,plain,(\n",
      "  d(two,a) | ~spl0_6),\n",
      "  inference(duplicate_literal_removal,[],[f294])).\n",
      "fof(f294,plain,(\n",
      "  d(two,a) | d(two,a) | ~spl0_6),\n",
      "  inference(resolution,[],[f8,f117])).\n",
      "fof(f117,plain,(\n",
      "  d(two,m(a,a)) | ~spl0_6),\n",
      "  inference(avatar_component_clause,[],[f115])).\n",
      "fof(f268,plain,(\n",
      "  spl0_16 | ~spl0_15),\n",
      "  inference(avatar_split_clause,[],[f263,f259,f265])).\n",
      "fof(f263,plain,(\n",
      "  m(a,a) = m(m(b,two),f(m(b,two),m(a,a))) | ~spl0_15),\n",
      "  inference(resolution,[],[f261,f6])).\n",
      "fof(f262,plain,(\n",
      "  spl0_15 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f257,f19,f259])).\n",
      "fof(f257,plain,(\n",
      "  d(m(b,two),m(a,a)) | ~spl0_2),\n",
      "  inference(forward_demodulation,[],[f246,f4])).\n",
      "fof(f246,plain,(\n",
      "  d(m(two,b),m(a,a)) | ~spl0_2),\n",
      "  inference(superposition,[],[f159,f21])).\n",
      "fof(f212,plain,(\n",
      "  spl0_14 | ~spl0_7),\n",
      "  inference(avatar_split_clause,[],[f207,f120,f209])).\n",
      "fof(f207,plain,(\n",
      "  m(a,a) = m(m(b,b),f(m(b,b),m(a,a))) | ~spl0_7),\n",
      "  inference(resolution,[],[f122,f6])).\n",
      "fof(f192,plain,(\n",
      "  spl0_13 | ~spl0_12),\n",
      "  inference(avatar_split_clause,[],[f187,f183,f189])).\n",
      "fof(f187,plain,(\n",
      "  m(a,a) = m(b,f(b,m(a,a))) | ~spl0_12),\n",
      "  inference(resolution,[],[f185,f6])).\n",
      "fof(f186,plain,(\n",
      "  spl0_12 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f179,f19,f183])).\n",
      "fof(f179,plain,(\n",
      "  d(b,m(a,a)) | ~spl0_2),\n",
      "  inference(superposition,[],[f160,f21])).\n",
      "fof(f143,plain,(\n",
      "  spl0_11 | ~spl0_6),\n",
      "  inference(avatar_split_clause,[],[f138,f115,f140])).\n",
      "fof(f138,plain,(\n",
      "  m(a,a) = m(two,f(two,m(a,a))) | ~spl0_6),\n",
      "  inference(resolution,[],[f117,f6])).\n",
      "fof(f137,plain,(\n",
      "  spl0_1 | ~spl0_10 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f113,f19,f134,f14])).\n",
      "fof(f113,plain,(\n",
      "  m(a,a) != m(b,b) | one = two | ~spl0_2),\n",
      "  inference(superposition,[],[f65,f21])).\n",
      "fof(f132,plain,(\n",
      "  spl0_8 | ~spl0_9 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f112,f19,f129,f125])).\n",
      "fof(f112,plain,(\n",
      "  two != m(a,a) | one = m(b,b) | ~spl0_2),\n",
      "  inference(superposition,[],[f51,f21])).\n",
      "fof(f123,plain,(\n",
      "  spl0_7 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f111,f19,f120])).\n",
      "fof(f111,plain,(\n",
      "  d(m(b,b),m(a,a)) | ~spl0_2),\n",
      "  inference(superposition,[],[f31,f21])).\n",
      "fof(f118,plain,(\n",
      "  spl0_6 | ~spl0_2),\n",
      "  inference(avatar_split_clause,[],[f110,f19,f115])).\n",
      "fof(f110,plain,(\n",
      "  d(two,m(a,a)) | ~spl0_2),\n",
      "  inference(superposition,[],[f12,f21])).\n",
      "fof(f84,plain,(\n",
      "  spl0_5),\n",
      "  inference(avatar_split_clause,[],[f71,f81])).\n",
      "fof(f81,plain,(\n",
      "  spl0_5 <=> one = f(one,one)),\n",
      "  introduced(avatar_definition,[new_symbols(naming,[spl0_5])])).\n",
      "fof(f71,plain,(\n",
      "  one = f(one,one)),\n",
      "  inference(superposition,[],[f69,f1])).\n",
      "fof(f49,plain,(\n",
      "  spl0_3 | ~spl0_4),\n",
      "  inference(avatar_split_clause,[],[f39,f46,f42])).\n",
      "fof(f39,plain,(\n",
      "  ~d(b,a) | one = b),\n",
      "  inference(resolution,[],[f10,f24])).\n",
      "fof(f22,plain,(\n",
      "  spl0_2),\n",
      "  inference(avatar_split_clause,[],[f9,f19])).\n",
      "fof(f9,axiom,(\n",
      "  m(a,a) = m(two,m(b,b))),\n",
      "  file('/tmp/sqrt.p',unknown)).\n",
      "fof(f17,plain,(\n",
      "  ~spl0_1),\n",
      "  inference(avatar_split_clause,[],[f11,f14])).\n",
      "fof(f11,axiom,(\n",
      "  one != two),\n",
      "  file('/tmp/sqrt.p',unknown)).\n",
      "% SZS output end Proof for sqrt\n",
      "% ------------------------------\n",
      "% Version: Vampire 4.8 (commit 23d7e9c3b on 2024-03-22 08:41:08 +0000)\n",
      "% Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d\n",
      "% Termination reason: Refutation\n",
      "\n",
      "% Memory used [KB]: 6366\n",
      "% Time elapsed: 0.542 s\n",
      "% ------------------------------\n",
      "% ------------------------------\n",
      "% Success in time 0.568 s\n"
     ]
    }
   ],
   "source": [
    "! vampire --mode casc /tmp/sqrt.p"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 2,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "# Preprocessing class: FSSSSMSSSSSNFFN.\n",
      "# Scheduled 1 strats onto 1 cores with 300 seconds (300 total)\n",
      "# Starting G-E--_302_C18_F1_URBAN_RG_S04BN with 300s (1) cores\n",
      "# G-E--_302_C18_F1_URBAN_RG_S04BN with pid 415821 completed with status 0\n",
      "# Result found by G-E--_302_C18_F1_URBAN_RG_S04BN\n",
      "# No SInE strategy applied\n",
      "# Search class: FGUSM-FFSF22-SFFFFFNN\n",
      "# Scheduled 6 strats onto 1 cores with 300 seconds (300 total)\n",
      "# Starting SAT001_MinMin_p005000_rr_RG with 109s (1) cores\n",
      "# SAT001_MinMin_p005000_rr_RG with pid 415822 completed with status 0\n",
      "# Result found by SAT001_MinMin_p005000_rr_RG\n",
      "# Initializing proof state\n",
      "# Scanning for AC axioms\n",
      "# m is AC\n",
      "# AC handling enabled\n",
      "#\n",
      "#cnf(i_0_22, plain, (one!=two)).\n",
      "#\n",
      "#cnf(i_0_13, plain, (m(X1,one)=X1)).\n",
      "#\n",
      "#cnf(i_0_12, plain, (m(one,X1)=X1)).\n",
      "#\n",
      "#cnf(i_0_16, plain, (X2=X3|m(X1,X2)!=m(X1,X3))).\n",
      "#\n",
      "#cnf(i_0_21, plain, (X1=one|~d(X1,a)|~d(X1,b))).\n",
      "#\n",
      "#cnf(i_0_15, plain, (m(X1,X2)=m(X2,X1))).\n",
      "#\n",
      "#cnf(i_0_20, plain, (m(two,m(b,b))=m(a,a))).\n",
      "#\n",
      "#cnf(i_0_18, plain, (d(X1,X3)|m(X1,X2)!=X3)).\n",
      "#\n",
      "#cnf(i_0_18, plain, (d(X1,m(X1,X2)))).\n",
      "#\n",
      "#cnf(i_0_17, plain, (m(X1,f(X1,X2))=X2|~d(X1,X2))).\n",
      "#\n",
      "#cnf(i_0_19, plain, (d(two,X1)|d(two,X2)|~d(two,m(X1,X2)))).\n",
      "#\n",
      "#cnf(i_0_14, plain, (m(m(X1,X2),X3)=m(X1,m(X2,X3)))).\n",
      "# Presaturation interreduction done\n",
      "#\n",
      "#cnf(i_0_22, plain, (one!=two)).\n",
      "#\n",
      "#cnf(i_0_13, plain, (m(X1,one)=X1)).\n",
      "#\n",
      "#cnf(i_0_12, plain, (m(one,X1)=X1)).\n",
      "#\n",
      "#cnf(i_0_16, plain, (X1=X2|m(X3,X1)!=m(X3,X2))).\n",
      "#\n",
      "#cnf(i_0_20, plain, (m(two,m(b,b))=m(a,a))).\n",
      "#\n",
      "#cnf(i_0_26, plain, (X1=one|m(X2,X1)!=X2)).\n",
      "#\n",
      "#cnf(i_0_21, plain, (X1=one|~d(X1,a)|~d(X1,b))).\n",
      "#\n",
      "#cnf(i_0_18, plain, (d(X1,m(X1,X2)))).\n",
      "#\n",
      "#cnf(i_0_15, plain, (m(X1,X2)=m(X2,X1))).\n",
      "#\n",
      "#cnf(i_0_17, plain, (m(X1,f(X1,X2))=X2|~d(X1,X2))).\n",
      "#\n",
      "#cnf(i_0_14, plain, (m(m(X1,X2),X3)=m(X1,m(X2,X3)))).\n",
      "#\n",
      "#cnf(i_0_34, plain, (d(one,X1))).\n",
      "#\n",
      "#cnf(i_0_35, plain, (d(X1,X1))).\n",
      "#\n",
      "#cnf(i_0_73, plain, (b=one|~d(b,a))).\n",
      "##\n",
      "#cnf(i_0_19, plain, (d(two,X2)|d(two,X1)|~d(two,m(X1,X2)))).\n",
      "#\n",
      "#cnf(i_0_40, plain, (X1=one|m(X1,X2)!=X2)).\n",
      "#\n",
      "#cnf(i_0_82, plain, (m(b,b)!=m(a,a))).\n",
      "######\n",
      "#cnf(i_0_31, plain, (m(b,b)=one|m(a,a)!=two)).\n",
      "#\n",
      "#cnf(i_0_100, plain, (m(a,a)!=one|m(a,a)!=two)).\n",
      "##\n",
      "#cnf(i_0_29, plain, (m(b,b)=X1|m(a,a)!=m(two,X1))).\n",
      "#\n",
      "#cnf(i_0_33, plain, (d(two,m(a,a)))).\n",
      "##\n",
      "#cnf(i_0_109, plain, (d(two,a))).\n",
      "#\n",
      "#cnf(i_0_41, plain, (X1=X2|m(X1,X3)!=m(X3,X2))).\n",
      "###\n",
      "#cnf(i_0_50, plain, (f(one,X1)=X1)).\n",
      "#\n",
      "#cnf(i_0_52, plain, (f(X1,X1)=one)).\n",
      "#\n",
      "#cnf(i_0_90, plain, (d(b,one)|m(a,a)!=two)).\n",
      "###\n",
      "#cnf(i_0_39, plain, (d(X1,m(X2,X1)))).\n",
      "####\n",
      "#cnf(i_0_84, plain, (X1=one|f(X1,X2)!=X2|~d(X1,X2))).\n",
      "####\n",
      "#cnf(i_0_119, plain, (X1=X2|m(X1,X3)!=m(X2,X3))).\n",
      "##\n",
      "#cnf(i_0_53, plain, (f(X1,m(X1,X2))=X2)).\n",
      "#######\n",
      "#cnf(i_0_160, plain, (f(b,one)=b|m(a,a)!=two)).\n",
      "####\n",
      "#cnf(i_0_58, plain, (m(X1,m(X2,f(m(X1,X2),X3)))=X3|~d(m(X1,X2),X3))).\n",
      "#\n",
      "#cnf(i_0_165, plain, (f(X1,m(X2,X1))=X2)).\n",
      "##\n",
      "#cnf(i_0_173, plain, (d(X1,X3)|~d(m(X1,X2),X3))).\n",
      "#\n",
      "#cnf(i_0_220, plain, (d(X1,X2)|~d(X3,X2)|~d(X1,X3))).\n",
      "#\n",
      "#cnf(i_0_231, plain, (d(X1,X2)|~d(X1,one))).\n",
      "###\n",
      "#cnf(i_0_216, plain, (d(b,X1)|m(a,a)!=two)).\n",
      "#\n",
      "#cnf(i_0_241, plain, (b=one|m(a,a)!=two)).\n",
      "##\n",
      "#cnf(i_0_60, plain, (d(m(X1,X2),m(X1,m(X2,X3))))).\n",
      "##\n",
      "#cnf(i_0_223, plain, (d(X1,X2)|~d(m(X3,X1),X2))).\n",
      "##\n",
      "#cnf(i_0_229, plain, (d(X1,a)|~d(X1,two))).\n",
      "#\n",
      "#cnf(i_0_283, plain, (b=one|~d(b,two))).\n",
      "######\n",
      "#cnf(i_0_62, plain, (m(X1,m(X2,X3))=m(X3,m(X1,X2)))).\n",
      "####\n",
      "#cnf(i_0_286, plain, (d(X1,a)|~d(X2,two)|~d(X1,X2))).\n",
      "####\n",
      "#cnf(i_0_97, plain, (b=X1|m(a,a)!=two|m(b,X1)!=one)).\n",
      "##\n",
      "#cnf(i_0_113, plain, (X1=b|m(a,a)!=two|m(X1,b)!=one)).\n",
      "###\n",
      "#cnf(i_0_284, plain, (d(X1,a)|~d(m(X1,X2),two))).\n",
      "###\n",
      "#cnf(i_0_285, plain, (d(X1,a)|~d(m(X2,X1),two))).\n",
      "######\n",
      "#cnf(i_0_64, plain, (X1=one|m(X2,m(X3,X1))!=m(X2,X3))).\n",
      "###\n",
      "#cnf(i_0_409, plain, (b=one|m(a,a)!=m(two,b))).\n",
      "########\n",
      "#cnf(i_0_65, plain, (X1=X2|m(X3,m(X4,X1))!=m(X3,m(X4,X2)))).\n",
      "#\n",
      "#cnf(i_0_161, plain, (f(two,m(a,a))=m(b,b))).\n",
      "#########\n",
      "#cnf(i_0_143, plain, (X1=two|m(X1,m(b,b))!=m(a,a))).\n",
      "##\n",
      "#cnf(i_0_133, plain, (d(m(b,b),m(a,a)))).\n",
      "#####\n",
      "#cnf(i_0_476, plain, (d(b,m(a,a)))).\n",
      "####\n",
      "#cnf(i_0_327, plain, (X1=one|m(X3,m(X1,X2))!=m(X2,X3))).\n",
      "#\n",
      "#cnf(i_0_67, plain, (m(two,m(b,m(b,X1)))=m(a,m(a,X1)))).\n",
      "##\n",
      "#cnf(i_0_510, plain, (m(b,m(b,X1))!=m(a,m(a,X1)))).\n",
      "##\n",
      "#cnf(i_0_552, plain, (m(b,m(X1,b))!=m(a,m(a,X1)))).\n",
      "##\n",
      "#cnf(i_0_554, plain, (m(a,m(a,b))!=b|m(a,a)!=two)).\n",
      "#####\n",
      "#cnf(i_0_69, plain, (m(X1,m(f(X1,X2),X3))=m(X2,X3)|~d(X1,X2))).\n",
      "#\n",
      "#cnf(i_0_134, plain, (d(X1,m(X2,m(X3,X1))))).\n",
      "#######\n",
      "#cnf(i_0_407, plain, (X1=one|m(X2,m(X1,X3))!=m(X2,X3))).\n",
      "###\n",
      "#cnf(i_0_71, plain, (m(X2,m(X1,X3))=m(X1,m(X2,X3)))).\n",
      "##\n",
      "#cnf(i_0_649, plain, (m(X1,m(b,b))!=m(X1,m(a,a)))).\n",
      "##\n",
      "#cnf(i_0_809, plain, (m(two,m(a,a))!=m(a,a))).\n",
      "#############\n",
      "#cnf(i_0_808, plain, (m(X1,m(a,a))!=X1|m(a,a)!=two)).\n",
      "####\n",
      "#cnf(i_0_824, plain, (m(a,m(a,X1))!=X1|m(a,a)!=two)).\n",
      "#\n",
      "#cnf(i_0_75, plain, (d(two,m(X2,X3))|d(two,X1)|~d(two,m(X2,m(X3,X1))))).\n",
      "#\n",
      "#cnf(i_0_208, plain, (f(m(b,b),m(a,a))=two)).\n",
      "###\n",
      "#cnf(i_0_826, plain, (m(a,m(X1,a))!=X1|m(a,a)!=two)).\n",
      "#######\n",
      "#cnf(i_0_77, plain, (d(two,f(X1,X2))|d(two,X1)|~d(two,X2)|~d(X1,X2))).\n",
      "###############\n",
      "#cnf(i_0_830, plain, (f(a,X1)!=m(a,X1)|m(a,a)!=two|~d(a,X1))).\n",
      "########\n",
      "#cnf(i_0_226, plain, (d(X1,m(X2,m(X1,X3))))).\n",
      "#\n",
      "#cnf(i_0_887, plain, (f(a,X1)!=m(X1,a)|m(a,a)!=two|~d(a,X1))).\n",
      "####\n",
      "#cnf(i_0_487, plain, (m(X1,m(a,a))!=m(b,m(b,X1)))).\n",
      "#####\n",
      "#cnf(i_0_83, plain, (m(X1,X2)=one|m(X1,m(X2,X3))!=X3)).\n",
      "####\n",
      "#cnf(i_0_951, plain, (m(b,m(b,X1))!=m(a,m(X1,a)))).\n",
      "####\n",
      "#cnf(i_0_970, plain, (m(two,b)=one|m(a,a)!=b)).\n",
      "###\n",
      "#cnf(i_0_91, plain, (m(b,m(b,X1))=X1|m(a,a)!=two)).\n",
      "#####\n",
      "#cnf(i_0_991, plain, (m(b,m(X1,b))!=m(a,m(X1,a)))).\n",
      "#################\n",
      "#cnf(i_0_104, plain, (m(b,b)=X1|m(a,a)!=m(X1,two))).\n",
      "############\n",
      "#cnf(i_0_1009, plain, (f(b,X1)=m(b,X1)|m(a,a)!=two)).\n",
      "#\n",
      "#cnf(i_0_556, plain, (m(X1,m(b,b))!=m(a,m(a,X1)))).\n",
      "##########\n",
      "#cnf(i_0_255, plain, (d(m(two,b),m(a,a)))).\n",
      "#########\n",
      "#cnf(i_0_432, plain, (X1=b|m(two,m(b,X1))!=m(a,a))).\n",
      "######\n",
      "#cnf(i_0_814, plain, (m(b,m(X1,b))!=m(X1,m(a,a)))).\n",
      "##########################\n",
      "#cnf(i_0_470, plain, (X1=two|m(b,m(b,X1))!=m(a,a))).\n",
      "##\n",
      "#cnf(i_0_115, plain, (X1=X2|m(X1,m(X3,X4))!=m(X3,m(X4,X2)))).\n",
      "###########\n",
      "#cnf(i_0_122, plain, (m(X1,X2)=X3|m(X1,m(X2,X4))!=m(X4,X3))).\n",
      "##\n",
      "#cnf(i_0_995, plain, (m(X1,m(b,b))!=m(a,m(X1,a)))).\n",
      "#########\n",
      "#cnf(i_0_124, plain, (X1=X2|m(f(X1,X3),X2)!=X3|~d(X1,X3))).\n",
      "###################\n",
      "#cnf(i_0_472, plain, (X1=two|m(b,m(X1,b))!=m(a,a))).\n",
      "###\n",
      "#cnf(i_0_136, plain, (d(f(X1,X2),X2)|~d(X1,X2))).\n",
      "########\n",
      "#cnf(i_0_494, plain, (X1=one|m(X2,m(X3,X1))!=m(X3,X2))).\n",
      "###\n",
      "#cnf(i_0_144, plain, (X1=m(X2,X3)|m(X1,X4)!=m(X2,m(X3,X4)))).\n",
      "###########\n",
      "#cnf(i_0_146, plain, (X1=X2|m(X1,f(X2,X3))!=X3|~d(X2,X3))).\n",
      "#############\n",
      "#cnf(i_0_549, plain, (m(b,m(X2,m(b,X1)))!=m(a,m(a,m(X1,X2))))).\n",
      "###\n",
      "#cnf(i_0_550, plain, (m(b,m(X1,m(X2,b)))!=m(a,m(a,m(X1,X2))))).\n",
      "#########\n",
      "#cnf(i_0_650, plain, (m(X1,m(b,m(b,X2)))!=m(X1,m(a,m(a,X2))))).\n",
      "###\n",
      "#cnf(i_0_1445, plain, (m(two,m(a,m(a,b)))!=m(a,m(a,b)))).\n",
      "####\n",
      "#cnf(i_0_664, plain, (X1=one|m(X1,m(X3,X2))!=m(X2,X3))).\n",
      "#\n",
      "#cnf(i_0_162, plain, (f(m(X1,X2),m(X1,m(X2,X3)))=X3)).\n",
      "####\n",
      "#cnf(i_0_1501, plain, (m(two,m(a,m(a,X1)))!=m(a,m(a,X1)))).\n",
      "##\n",
      "#cnf(i_0_1602, plain, (m(two,X1)!=X1|~d(m(a,a),X1))).\n",
      "######\n",
      "#cnf(i_0_507, plain, (d(two,m(a,m(a,X1))))).\n",
      "###\n",
      "#cnf(i_0_1609, plain, (m(two,a)!=a|~d(m(a,a),two))).\n",
      "#######\n",
      "#cnf(i_0_170, plain, (m(X1,m(X2,m(X3,f(m(X1,m(X2,X3)),X4))))=X4|~d(m(X1,m(X2,X3)),X4))).\n",
      "##\n",
      "#cnf(i_0_1592, plain, (m(two,m(a,X1))!=m(a,X1)|~d(a,X1))).\n",
      "#\n",
      "#cnf(i_0_1800, plain, (m(two,a)!=a|~d(a,one))).\n",
      "#\n",
      "#cnf(i_0_1801, plain, (m(two,m(X1,a))!=m(X1,a)|~d(a,X1))).\n",
      "#########\n",
      "#cnf(i_0_1805, plain, (m(X1,m(two,a))!=m(a,X1)|~d(a,X1))).\n",
      "####\n",
      "#cnf(i_0_1811, plain, (m(X1,m(two,a))!=m(X1,a)|~d(a,X1))).\n",
      "#####\n",
      "#cnf(i_0_172, plain, (X1=X2|m(X1,m(X3,f(m(X2,X3),X4)))!=X4|~d(m(X2,X3),X4))).\n",
      "#########\n",
      "#cnf(i_0_1134, plain, (X1=b|m(two,m(X1,b))!=m(a,a))).\n",
      "##\n",
      "#cnf(i_0_174, plain, (m(X1,m(X2,m(f(m(X1,X2),X3),X4)))=m(X3,X4)|~d(m(X1,X2),X3))).\n",
      "##\n",
      "#cnf(i_0_1819, plain, (m(two,m(X1,a))!=m(a,X1)|~d(a,X1))).\n",
      "####\n",
      "#cnf(i_0_1824, plain, (m(two,m(a,X1))!=m(X1,a)|~d(a,X1))).\n",
      "###\n",
      "#cnf(i_0_1138, plain, (X1=b|m(X1,m(two,b))!=m(a,a))).\n",
      "##\n",
      "#cnf(i_0_175, plain, (m(X2,f(m(X1,X2),X3))=f(X1,X3)|~d(m(X1,X2),X3))).\n",
      "#\n",
      "#cnf(i_0_2139, plain, (f(b,m(a,a))=m(two,b))).\n",
      "##########\n",
      "#cnf(i_0_176, plain, (X1=one|m(X2,f(m(X1,X2),X3))!=X3|~d(m(X1,X2),X3))).\n",
      "###########\n",
      "#cnf(i_0_177, plain, (m(X1,m(X2,m(f(m(X1,X2),X3),f(X3,X4))))=X4|~d(m(X1,X2),X3)|~d(X3,X4))).\n",
      "#####\n",
      "#cnf(i_0_2201, plain, (m(two,X2)!=X2|~d(m(a,a),X1)|~d(X1,X2))).\n",
      "######\n",
      "#cnf(i_0_178, plain, (d(two,m(X1,f(m(X2,X1),X3)))|d(two,X2)|~d(m(X2,X1),X3)|~d(two,X3))).\n",
      "####\n",
      "#cnf(i_0_1594, plain, (m(two,m(a,X2))!=m(a,X2)|~d(m(a,X1),X2))).\n",
      "##\n",
      "#cnf(i_0_1657, plain, (m(two,X2)!=X2|~d(m(a,m(a,X1)),X2))).\n",
      "####\n",
      "#cnf(i_0_1271, plain, (m(two,b)=X1|m(b,X1)!=m(a,a))).\n",
      "#\n",
      "#cnf(i_0_180, plain, (d(m(X1,f(m(X2,X1),X3)),X3)|~d(m(X2,X1),X3))).\n",
      "##\n",
      "#cnf(i_0_2401, plain, (m(two,X1)!=X1|~d(m(a,X2),X1)|~d(a,X2))).\n",
      "#\n",
      "#cnf(i_0_2498, plain, (m(two,X1)!=X1|~d(a,one))).\n",
      "##\n",
      "#cnf(i_0_2499, plain, (m(two,X1)!=X1|~d(m(X2,a),X1)|~d(a,X2))).\n",
      "######\n",
      "#cnf(i_0_181, plain, (X1=X2|m(X3,m(f(m(X1,X3),X4),X2))!=X4|~d(m(X1,X3),X4))).\n",
      "##############\n",
      "#cnf(i_0_1894, plain, (m(two,m(X1,X2))!=m(X1,X2)|~d(m(a,a),X1))).\n",
      "#\n",
      "#cnf(i_0_2346, plain, (m(two,X1)!=X1|~d(m(X2,m(a,a)),X1))).\n",
      "#########\n",
      "#cnf(i_0_2391, plain, (m(two,m(a,X1))!=m(a,X1)|~d(m(X2,a),X1))).\n",
      "###\n",
      "#cnf(i_0_2412, plain, (m(two,X1)!=X1|~d(m(a,m(X2,a)),X1))).\n",
      "##################\n",
      "#cnf(i_0_909, plain, (d(two,m(X1,m(a,a))))).\n",
      "#####\n",
      "#cnf(i_0_2422, plain, (m(two,a)!=a|~d(m(a,m(a,X1)),two))).\n",
      "######\n",
      "#cnf(i_0_924, plain, (d(b,m(a,m(a,X1))))).\n",
      "##########\n",
      "#cnf(i_0_190, plain, (m(X1,m(f(X1,X2),f(X2,X3)))=X3|~d(X2,X3)|~d(X1,X2))).\n",
      "#\n",
      "#cnf(i_0_1137, plain, (f(m(two,b),m(a,a))=b)).\n",
      "#####\n",
      "#cnf(i_0_2596, plain, (m(two,a)!=a|~d(m(X1,m(a,a)),two))).\n",
      "################\n",
      "#cnf(i_0_192, plain, (m(X1,m(X2,f(m(X2,X1),X3)))=X3|~d(m(X2,X1),X3))).\n",
      "############\n",
      "#cnf(i_0_1627, plain, (d(two,m(a,m(X1,a))))).\n",
      "##\n",
      "#cnf(i_0_1792, plain, (m(two,X1)!=X1|~d(a,f(a,X1))|~d(a,X1))).\n",
      "##########\n",
      "#cnf(i_0_2342, plain, (m(two,X1)!=X1|~d(m(a,a),two)|~d(a,X1))).\n",
      "#\n",
      "#cnf(i_0_2501, plain, (m(two,a)!=a|~d(m(a,X1),two)|~d(a,X1))).\n",
      "##\n",
      "#cnf(i_0_2513, plain, (m(two,a)!=a|~d(m(X1,a),two)|~d(a,X1))).\n",
      "#########\n",
      "#cnf(i_0_2657, plain, (m(two,a)!=a|~d(m(a,a),X1)|~d(X1,two))).\n",
      "###\n",
      "#cnf(i_0_2623, plain, (m(two,a)!=a|~d(m(a,m(X1,a)),two))).\n",
      "###\n",
      "#cnf(i_0_217, plain, (d(two,X1)|~d(m(a,a),X1))).\n",
      "##\n",
      "#cnf(i_0_169, plain, (m(X1,f(m(two,X1),m(a,a)))=m(b,b)|~d(m(two,X1),m(a,a)))).\n",
      "########\n",
      "#cnf(i_0_333, plain, (m(X1,X2)=one|m(X2,m(X3,X1))!=X3)).\n",
      "###\n",
      "#cnf(i_0_179, plain, (m(X1,f(m(X2,X1),X2))=one|~d(m(X2,X1),X2))).\n",
      "#\n",
      "#cnf(i_0_3209, plain, (m(X1,X2)=one|m(X2,X1)!=one)).\n",
      "#\n",
      "#cnf(i_0_3277, plain, (m(two,a)!=a|~d(m(X1,a),X1))).\n",
      "##\n",
      "#cnf(i_0_3422, plain, (m(two,a)!=a|~d(m(a,X1),X1))).\n",
      "#######\n",
      "#cnf(i_0_183, plain, (m(X1,f(m(X2,X1),m(X2,X3)))=X3|~d(m(X2,X1),m(X2,X3)))).\n",
      "##########\n",
      "#cnf(i_0_696, plain, (m(X1,X2)=one|m(X1,m(X3,X2))!=X3)).\n",
      "###\n",
      "#cnf(i_0_3282, plain, (m(two,X1)!=X1|~d(m(X2,a),X2)|~d(a,X1))).\n",
      "##\n",
      "#cnf(i_0_3666, plain, (m(two,X1)!=X1|~d(m(a,X2),X2)|~d(a,X1))).\n",
      "#######\n",
      "#cnf(i_0_185, plain, (m(X2,f(m(X3,X2),m(X1,X3)))=X1|~d(m(X3,X2),m(X1,X3)))).\n",
      "#############\n",
      "#cnf(i_0_3281, plain, (m(two,m(a,X1))!=m(a,X1)|~d(m(X2,a),X2))).\n",
      "##\n",
      "#cnf(i_0_3857, plain, (m(two,m(a,X1))!=m(a,X1)|~d(m(a,X2),X2))).\n",
      "######\n",
      "#cnf(i_0_979, plain, (m(X1,X2)=one|m(X2,m(X1,X3))!=X3)).\n",
      "#\n",
      "#cnf(i_0_209, plain, (f(X1,m(X2,m(X3,X1)))=m(X2,X3))).\n",
      "###########\n",
      "#cnf(i_0_211, plain, (f(f(X1,X2),X2)=X1|~d(X1,X2))).\n",
      "##\n",
      "#cnf(i_0_3428, plain, (m(two,a)!=a|~d(X1,f(a,X1))|~d(a,X1))).\n",
      "#######\n",
      "#cnf(i_0_4028, plain, (m(two,a)!=a|m(a,a)!=two|~d(a,b))).\n",
      "##\n",
      "#cnf(i_0_212, plain, (f(m(X1,f(m(X2,X1),X3)),X3)=X2|~d(m(X2,X1),X3))).\n",
      "###########\n",
      "#cnf(i_0_218, plain, (d(m(X1,X2),X3)|~d(m(X1,m(X2,X4)),X3))).\n",
      "#\n",
      "#cnf(i_0_3955, plain, (f(b,X1)=m(X1,b)|m(a,a)!=two)).\n",
      "##########\n",
      "#cnf(i_0_221, plain, (d(X1,X2)|~d(m(X1,X3),X4)|~d(X4,X2))).\n",
      "#####\n",
      "#cnf(i_0_2579, plain, (m(two,m(a,X1))!=m(a,X1)|~d(m(a,a),two))).\n",
      "##\n",
      "#cnf(i_0_2680, plain, (d(b,X1)|~d(m(a,a),X1))).\n",
      "####\n",
      "#cnf(i_0_230, plain, (d(X1,m(a,a))|~d(X1,two))).\n",
      "###########\n",
      "#cnf(i_0_232, plain, (d(X1,m(X2,X3))|~d(X1,X2))).\n",
      "###\n",
      "#cnf(i_0_4255, plain, (m(two,a)!=a|~d(m(a,m(X1,X2)),X1))).\n",
      "##\n",
      "#cnf(i_0_4300, plain, (m(two,a)!=a|~d(m(a,X1),one))).\n",
      "###\n",
      "#cnf(i_0_4353, plain, (m(two,a)!=a|~d(m(X1,a),one))).\n",
      "###\n",
      "#cnf(i_0_233, plain, (d(X1,m(X2,X3))|~d(X1,X3))).\n",
      "##\n",
      "#cnf(i_0_4340, plain, (m(two,a)!=a|~d(X1,one)|~d(a,X1))).\n",
      "#########\n",
      "#cnf(i_0_240, plain, (d(X1,X2)|m(a,a)!=two|~d(X1,b))).\n",
      "##\n",
      "#cnf(i_0_4305, plain, (m(two,a)!=a|~d(m(X2,X1),X2)|~d(a,X1))).\n",
      "#########\n",
      "#cnf(i_0_244, plain, (d(X1,m(X2,m(X3,X4)))|~d(X1,m(X2,X3)))).\n",
      "##\n",
      "#cnf(i_0_4447, plain, (m(two,a)!=a|~d(m(X2,X1),X1)|~d(a,X2))).\n",
      "#####\n",
      "#cnf(i_0_3242, plain, (m(X1,X2)=one|m(X3,m(X2,X1))!=X3)).\n",
      "####\n",
      "#cnf(i_0_246, plain, (d(m(X1,b),X1)|m(a,a)!=two)).\n",
      "###########\n",
      "#cnf(i_0_247, plain, (d(m(X1,two),m(X1,m(a,a))))).\n",
      "##########\n",
      "#cnf(i_0_3276, plain, (d(b,a)|~d(m(X1,a),X1))).\n",
      "#\n",
      "#cnf(i_0_248, plain, (d(m(X1,m(X2,X3)),m(X1,m(X2,m(X3,X4)))))).\n",
      "#\n",
      "#cnf(i_0_4650, plain, (d(b,a)|~d(a,one))).\n",
      "#\n",
      "#cnf(i_0_4306, plain, (m(two,a)!=a|~d(m(a,X2),X1)|~d(X1,X2))).\n",
      "#####\n",
      "#cnf(i_0_4771, plain, (b=one|~d(a,one))).\n",
      "#\n",
      "#cnf(i_0_3302, plain, (d(X1,one)|~d(m(X2,X1),X2))).\n",
      "###\n",
      "#cnf(i_0_250, plain, (d(m(X1,X2),m(X1,X3))|~d(X2,X3))).\n",
      "#\n",
      "#cnf(i_0_2678, plain, (d(b,m(a,m(X1,a))))).\n",
      "######\n",
      "#cnf(i_0_3313, plain, (d(X1,a)|~d(m(X2,X1),X2))).\n",
      "####\n",
      "#cnf(i_0_251, plain, (d(m(X1,X2),m(X1,X4))|~d(m(X2,X3),X4))).\n",
      "####\n",
      "#cnf(i_0_4324, plain, (m(two,a)!=a|~d(m(X1,a),X2)|~d(X2,X1))).\n",
      "###\n",
      "#cnf(i_0_3327, plain, (d(X1,X2)|~d(m(X3,X1),X3))).\n",
      "####\n",
      "#cnf(i_0_253, plain, (d(m(X1,X2),m(X1,m(X3,X2))))).\n",
      "#\n",
      "#cnf(i_0_2684, plain, (d(b,m(X1,m(a,a))))).\n",
      "##############\n",
      "#cnf(i_0_4260, plain, (m(two,a)!=a|~d(m(X1,m(X2,a)),X1))).\n",
      "##########\n",
      "#cnf(i_0_4319, plain, (m(two,a)!=a|~d(m(a,m(X2,X1)),X1))).\n",
      "########\n",
      "#cnf(i_0_258, plain, (d(m(X1,X2),m(X2,m(X3,X1))))).\n",
      "####\n",
      "#cnf(i_0_4328, plain, (m(two,a)!=a|~d(m(X2,m(a,X1)),X1))).\n",
      "#########\n",
      "#cnf(i_0_4329, plain, (m(two,a)!=a|~d(m(X1,m(a,X2)),X1))).\n",
      "#####################\n",
      "#cnf(i_0_5242, plain, (d(m(X1,X2),m(X2,X1)))).\n",
      "#######\n",
      "#cnf(i_0_4209, plain, (d(b,a)|~d(m(a,a),two))).\n",
      "###\n",
      "#cnf(i_0_265, plain, (d(X2,m(X1,m(f(X1,X2),X3)))|~d(X1,X2))).\n",
      "#####\n",
      "#cnf(i_0_4341, plain, (m(two,a)!=a|~d(X2,one)|~d(a,X1)|~d(X1,X2))).\n",
      "######\n",
      "#cnf(i_0_266, plain, (d(X3,m(X1,m(X2,m(f(m(X1,X2),X3),X4))))|~d(m(X1,X2),X3))).\n",
      "#\n",
      "#cnf(i_0_354, plain, (m(b,m(X1,b))=X1|m(a,a)!=two)).\n",
      "##########\n",
      "#cnf(i_0_268, plain, (d(m(X2,X1),m(X1,m(X2,X3))))).\n",
      "####\n",
      "#cnf(i_0_4363, plain, (m(two,m(X1,X2))!=m(X1,X2)|~d(m(a,a),X2))).\n",
      "#########\n",
      "#cnf(i_0_4378, plain, (m(two,a)!=a|~d(m(X1,m(X2,a)),X2))).\n",
      "#########\n",
      "#cnf(i_0_271, plain, (d(m(b,b),X1)|~d(m(a,a),X1))).\n",
      "###########\n",
      "#cnf(i_0_272, plain, (d(X1,X2)|~d(m(X3,m(X4,X1)),X2))).\n",
      "###########\n",
      "#cnf(i_0_274, plain, (d(f(X1,X2),X3)|~d(X2,X3)|~d(X1,X2))).\n",
      "#######\n",
      "#cnf(i_0_4651, plain, (d(b,a)|~d(m(a,X1),X1))).\n",
      "####\n",
      "#cnf(i_0_275, plain, (d(m(X1,f(m(X2,X1),X3)),X4)|~d(m(X2,X1),X3)|~d(X3,X4))).\n",
      "#\n",
      "#cnf(i_0_483, plain, (d(X1,m(a,a))|~d(X1,b))).\n",
      "######\n",
      "#cnf(i_0_5043, plain, (d(X1,X2)|~d(m(X1,X3),X3))).\n",
      "####\n",
      "#cnf(i_0_288, plain, (m(b,b)=m(X1,X2)|m(X1,m(X2,two))!=m(a,a))).\n",
      "#########\n",
      "#cnf(i_0_5045, plain, (d(X1,X2)|~d(m(a,X1),two))).\n",
      "#############\n",
      "#cnf(i_0_291, plain, (d(m(X1,X2),m(X1,m(X3,m(X4,X2)))))).\n",
      "#########\n",
      "#cnf(i_0_5950, plain, (d(b,X1)|~d(m(a,a),b))).\n",
      "####################################\n",
      "#cnf(i_0_1017, plain, (d(m(b,X1),X1)|m(a,a)!=two)).\n",
      "##########\n",
      "#cnf(i_0_299, plain, (m(X1,m(X2,m(X3,f(m(X2,m(X3,X1)),X4))))=X4|~d(m(X2,m(X3,X1)),X4))).\n",
      "###########\n",
      "#cnf(i_0_300, plain, (d(m(X2,m(X3,X1)),m(X1,m(X2,m(X3,X4)))))).\n",
      "#\n",
      "#cnf(i_0_1018, plain, (f(m(b,X1),X1)=b|m(a,a)!=two)).\n",
      "#######\n",
      "#cnf(i_0_6207, plain, (d(X1,X2)|~d(m(X1,a),two))).\n",
      "###############\n",
      "#cnf(i_0_1035, plain, (d(m(b,b),X1)|m(a,a)!=two)).\n",
      "#####\n",
      "#cnf(i_0_4348, plain, (m(two,a)!=a|~d(m(X2,m(a,X1)),one))).\n",
      "#################\n",
      "#cnf(i_0_1036, plain, (m(X1,m(b,b))=X1|m(a,a)!=two)).\n",
      "###############################################\n",
      "#cnf(i_0_4342, plain, (m(two,a)!=a|~d(m(X1,X2),one)|~d(a,X1))).\n",
      "########\n",
      "#cnf(i_0_6944, plain, (m(b,b)=b|m(a,a)!=two)).\n",
      "#\n",
      "#cnf(i_0_7315, plain, (m(a,a)!=b|m(a,a)!=two)).\n",
      "#######################\n",
      "#cnf(i_0_4343, plain, (m(two,a)!=a|~d(m(a,X1),X2)|~d(X2,one))).\n",
      "######\n",
      "#cnf(i_0_7269, plain, (m(X1,b)=X1|m(a,a)!=two)).\n",
      "######\n",
      "#cnf(i_0_4344, plain, (m(two,a)!=a|~d(m(X1,a),X2)|~d(X2,one))).\n",
      "###########\n",
      "#cnf(i_0_6912, plain, (m(two,a)!=a|~d(m(X1,X2),one)|~d(a,X2))).\n",
      "#######\n",
      "#cnf(i_0_317, plain, (m(b,b)=m(X1,X2)|m(X2,m(two,X1))!=m(a,a))).\n",
      "######\n",
      "#cnf(i_0_4350, plain, (m(two,a)!=a|~d(m(X1,m(X2,a)),one))).\n",
      "################\n",
      "#cnf(i_0_320, plain, (d(m(X1,X2),m(X1,m(X4,m(X2,X3)))))).\n",
      "##########\n",
      "#cnf(i_0_4770, plain, (d(X1,a)|~d(a,one)|~d(X1,b))).\n",
      "############\n",
      "#cnf(i_0_322, plain, (X1=X2|m(X1,m(X3,X4))!=m(X4,m(X2,X3)))).\n",
      "################################\n",
      "#cnf(i_0_534, plain, (b=one|m(a,m(a,X1))!=m(two,m(b,X1)))).\n",
      "#\n",
      "#cnf(i_0_326, plain, (f(X1,m(X3,m(X1,X2)))=m(X2,X3))).\n",
      "###########\n",
      "#cnf(i_0_328, plain, (m(X1,m(X2,m(X3,f(m(X3,m(X1,X2)),X4))))=X4|~d(m(X3,m(X1,X2)),X4))).\n",
      "###########\n",
      "#cnf(i_0_329, plain, (d(m(X3,m(X1,X2)),m(X1,m(X2,m(X3,X4)))))).\n",
      "#########################\n",
      "#cnf(i_0_4414, plain, (m(two,a)!=a|~d(a,f(X1,one))|~d(X1,one))).\n",
      "#####\n",
      "#cnf(i_0_1396, plain, (X1=m(two,b)|m(X1,b)!=m(a,a))).\n",
      "###\n",
      "#cnf(i_0_332, plain, (d(two,m(X1,X2))|d(two,X3)|~d(two,m(X2,m(X3,X1))))).\n",
      "######\n",
      "#cnf(i_0_6152, plain, (m(two,a)!=a|~d(m(a,m(a,a)),b))).\n",
      "################\n",
      "#cnf(i_0_335, plain, (f(m(X1,X2),m(X2,m(X3,X1)))=X3)).\n",
      "###\n",
      "#cnf(i_0_5897, plain, (m(two,m(X1,a))!=m(X1,a)|~d(m(a,a),two))).\n",
      "#######\n",
      "#cnf(i_0_7736, plain, (b=one|m(a,m(X1,a))!=m(two,m(b,X1)))).\n",
      "#\n",
      "#cnf(i_0_336, plain, (X1=X2|m(X4,m(X1,X3))!=m(X3,m(X4,X2)))).\n",
      "####\n",
      "#cnf(i_0_1600, plain, (m(two,m(a,m(X1,a)))!=m(a,m(X1,a)))).\n",
      "########\n",
      "#cnf(i_0_8417, plain, (f(m(X1,X2),m(X2,X1))=one)).\n",
      "##\n",
      "#cnf(i_0_1603, plain, (m(two,m(X1,m(a,a)))!=m(X1,m(a,a)))).\n",
      "########\n",
      "#cnf(i_0_338, plain, (d(X1,X2)|~d(m(X4,m(X1,X3)),X2))).\n",
      "###########\n",
      "#cnf(i_0_339, plain, (m(X1,X2)=X3|m(X2,m(X4,X1))!=m(X4,X3))).\n",
      "###########\n",
      "#cnf(i_0_340, plain, (d(m(X1,X2),X3)|~d(m(X2,m(X4,X1)),X3))).\n",
      "######################\n",
      "#cnf(i_0_342, plain, (X1=m(X2,X3)|m(X1,X4)!=m(X3,m(X4,X2)))).\n",
      "############################################\n",
      "#cnf(i_0_349, plain, (m(X1,m(f(X1,X2),X3))=m(X3,X2)|~d(X1,X2))).\n",
      "###########\n",
      "#cnf(i_0_350, plain, (m(X1,m(X2,m(f(m(X1,X2),X3),X4)))=m(X4,X3)|~d(m(X1,X2),X3))).\n",
      "#########\n",
      "#cnf(i_0_421, plain, (X1=b|m(X2,m(b,X1))!=X2|m(a,a)!=two)).\n",
      "##\n",
      "#cnf(i_0_352, plain, (m(X1,m(X2,X3))=m(X3,m(X2,X1)))).\n",
      "#######################################################\n",
      "#cnf(i_0_359, plain, (m(X3,m(X2,m(f(m(X2,X3),X4),X1)))=m(X1,X4)|~d(m(X2,X3),X4))).\n",
      "#################################\n",
      "#cnf(i_0_368, plain, (m(X1,X2)=b|m(X2,m(b,X1))!=one|m(a,a)!=two)).\n",
      "###########\n",
      "#cnf(i_0_369, plain, (m(X1,X2)=b|m(X1,m(X2,b))!=one|m(a,a)!=two)).\n",
      "######################\n",
      "#cnf(i_0_367, plain, (m(X1,f(m(b,X1),one))=b|m(a,a)!=two|~d(m(b,X1),one))).\n",
      "#\n",
      "#cnf(i_0_1339, plain, (d(f(X1,one),X2)|~d(X1,one))).\n",
      "##########\n",
      "#cnf(i_0_379, plain, (d(m(X1,X2),a)|~d(m(X1,m(X2,X3)),two))).\n",
      "###########\n",
      "#cnf(i_0_382, plain, (d(X1,a)|~d(m(X1,X2),X3)|~d(X3,two))).\n",
      "#\n",
      "#cnf(i_0_1548, plain, (f(m(X1,b),X1)=b|m(a,a)!=two)).\n",
      "####\n",
      "#cnf(i_0_8729, plain, (m(X1,m(two,m(a,a)))!=m(X1,m(a,a)))).\n",
      "######\n",
      "#cnf(i_0_383, plain, (d(X1,a)|~d(m(X3,m(X1,X2)),two))).\n",
      "#\n",
      "#cnf(i_0_1564, plain, (f(m(b,b),X1)=X1|m(a,a)!=two)).\n",
      "##########\n",
      "#cnf(i_0_384, plain, (d(X1,a)|~d(m(X2,m(X3,X1)),two))).\n",
      "###########\n",
      "#cnf(i_0_389, plain, (d(m(b,b),a)|~d(m(a,a),two))).\n",
      "#\n",
      "#cnf(i_0_1619, plain, (d(two,m(a,X1))|~d(a,X1))).\n",
      "###\n",
      "#cnf(i_0_11356, plain, (m(two,m(a,m(a,X1)))!=m(X1,m(a,a)))).\n",
      "######\n",
      "#cnf(i_0_9771, plain, (m(b,X1)=b|m(a,a)!=two|m(X2,X1)!=X2)).\n",
      "########\n",
      "#cnf(i_0_9772, plain, (m(X1,b)=b|m(a,a)!=two|m(X2,X1)!=X2)).\n",
      "####\n",
      "#cnf(i_0_392, plain, (d(f(X1,X2),a)|~d(X2,two)|~d(X1,X2))).\n",
      "####\n",
      "#cnf(i_0_11606, plain, (m(two,m(a,m(X1,a)))!=m(X1,m(a,a)))).\n",
      "#######\n",
      "#cnf(i_0_393, plain, (d(m(X1,f(m(X2,X1),X3)),a)|~d(m(X2,X1),X3)|~d(X3,two))).\n",
      "###########\n",
      "#cnf(i_0_394, plain, (d(m(X1,X2),a)|~d(m(X2,m(X3,X1)),two))).\n",
      "######################\n",
      "#cnf(i_0_400, plain, (m(b,b)=one|m(X1,m(a,a))!=m(X1,two))).\n",
      "##\n",
      "#cnf(i_0_2403, plain, (m(two,X1)!=X1|~d(m(a,X3),X1)|~d(m(a,X2),X3))).\n",
      "##\n",
      "#cnf(i_0_11922, plain, (m(two,m(X1,a))!=m(X1,a)|~d(m(a,X2),X1))).\n",
      "#######\n",
      "#cnf(i_0_401, plain, (X1=one|m(X2,m(X3,m(X4,X1)))!=m(X2,m(X3,X4)))).\n",
      "##\n",
      "#cnf(i_0_11931, plain, (m(two,m(X1,a))!=m(X1,a)|~d(m(X2,a),X2))).\n",
      "##\n",
      "#cnf(i_0_11953, plain, (m(two,m(X1,a))!=m(X1,a)|~d(m(X2,a),X1))).\n",
      "#######\n",
      "#cnf(i_0_403, plain, (f(X1,X2)=one|m(X3,X2)!=m(X3,X1)|~d(X1,X2))).\n",
      "####\n",
      "#cnf(i_0_12073, plain, (m(two,m(X1,a))!=m(X1,a)|~d(m(a,X2),X2))).\n",
      "######\n",
      "#cnf(i_0_11660, plain, (m(b,X1)=b|m(a,a)!=two|m(X1,X2)!=X2)).\n",
      "#\n",
      "#cnf(i_0_404, plain, (m(X1,f(m(X2,X1),X3))=one|m(X4,X3)!=m(X4,X2)|~d(m(X2,X1),X3))).\n",
      "##########\n",
      "#cnf(i_0_12177, plain, (f(X1,one)=one|m(X2,X1)!=X2|~d(X1,one))).\n",
      "#\n",
      "#cnf(i_0_405, plain, (m(X1,X2)=one|m(X3,m(X2,m(X4,X1)))!=m(X3,X4))).\n",
      "#####\n",
      "#cnf(i_0_11894, plain, (m(two,X1)!=X1|~d(m(a,X2),one)|~d(a,X1))).\n",
      "######\n",
      "#cnf(i_0_406, plain, (m(X1,X2)=one|m(X3,m(X1,m(X2,X4)))!=m(X3,X4))).\n",
      "###\n",
      "#cnf(i_0_12532, plain, (m(two,X1)!=X1|~d(X2,one)|~d(a,X1)|~d(a,X2))).\n",
      "##########\n",
      "#cnf(i_0_12553, plain, (m(two,X1)!=X1|~d(m(X2,a),one)|~d(a,X1))).\n",
      "#########\n",
      "#cnf(i_0_422, plain, (X1=m(b,b)|m(X2,m(two,X1))!=m(X2,m(a,a)))).\n",
      "#\n",
      "#cnf(i_0_2667, plain, (d(b,m(a,X1))|~d(a,X1))).\n",
      "#######\n",
      "#cnf(i_0_12414, plain, (f(X1,one)=one|m(X1,X2)!=X2|~d(X1,one))).\n",
      "###\n",
      "#cnf(i_0_423, plain, (X1=X2|m(X3,m(X4,m(X5,X1)))!=m(X3,m(X4,m(X5,X2))))).\n",
      "###\n",
      "#cnf(i_0_2880, plain, (m(two,X1)!=X1|~d(m(a,X3),X1)|~d(m(X2,a),X3))).\n",
      "###\n",
      "#cnf(i_0_4296, plain, (m(two,a)!=a|m(a,a)!=two|~d(m(a,X1),b))).\n",
      "#####\n",
      "#cnf(i_0_425, plain, (X1=f(X2,X3)|m(X4,m(X2,X1))!=m(X4,X3)|~d(X2,X3))).\n",
      "###\n",
      "#cnf(i_0_4857, plain, (m(two,m(X1,X2))!=m(X1,X2)|~d(a,X1)|~d(a,X2))).\n",
      "#####\n",
      "#cnf(i_0_13213, plain, (m(b,b)!=two|m(a,a)!=two|~d(a,b))).\n",
      "##\n",
      "#cnf(i_0_12021, plain, (b=one|m(X1,m(a,a))!=m(X1,m(two,b)))).\n",
      "#\n",
      "#cnf(i_0_426, plain, (X1=m(X2,f(m(X3,X2),X4))|m(X5,m(X3,X1))!=m(X5,X4)|~d(m(X3,X2),X4))).\n",
      "##\n",
      "#cnf(i_0_11895, plain, (m(two,X1)!=X1|~d(m(X2,a),X1)|~d(m(a,X3),X2))).\n",
      "#########\n",
      "#cnf(i_0_427, plain, (X1=m(X2,X3)|m(X4,m(X5,X1))!=m(X4,m(X3,m(X5,X2))))).\n",
      "###\n",
      "#cnf(i_0_12994, plain, (m(two,X1)!=X1|~d(m(X2,a),X1)|~d(m(X3,a),X2))).\n",
      "###\n",
      "#cnf(i_0_5948, plain, (m(two,a)!=a|~d(m(a,a),b))).\n",
      "####\n",
      "#cnf(i_0_1006, plain, (X1=b|m(X1,m(b,X2))!=X2|m(a,a)!=two)).\n",
      "#\n",
      "#cnf(i_0_428, plain, (X1=m(X2,X3)|m(X4,m(X5,X1))!=m(X4,m(X2,m(X3,X5))))).\n",
      "###########\n",
      "#cnf(i_0_430, plain, (X1=X2|m(X3,m(X4,X1))!=m(X3,m(X2,X4)))).\n",
      "##\n",
      "#cnf(i_0_13226, plain, (m(X2,m(two,X1))!=m(X1,X2)|~d(a,X1)|~d(a,X2))).\n",
      "#\n",
      "#cnf(i_0_13939, plain, (m(X1,two)!=X1|~d(a,one))).\n",
      "###\n",
      "#cnf(i_0_6889, plain, (m(two,a)!=a|m(a,a)!=two|~d(a,m(b,b)))).\n",
      "#########\n",
      "#cnf(i_0_13227, plain, (m(X1,m(two,X2))!=m(X1,X2)|~d(a,X1)|~d(a,X2))).\n",
      "#########\n",
      "#cnf(i_0_13228, plain, (m(X2,m(X1,two))!=m(X1,X2)|~d(a,X1)|~d(a,X2))).\n",
      "#######\n",
      "#cnf(i_0_1180, plain, (X1=b|m(X1,m(X2,b))!=X2|m(a,a)!=two)).\n",
      "####\n",
      "#cnf(i_0_13229, plain, (m(X1,m(X2,two))!=m(X1,X2)|~d(a,X1)|~d(a,X2))).\n",
      "######\n",
      "#cnf(i_0_1225, plain, (b=X1|m(b,m(X2,X1))!=X2|m(a,a)!=two)).\n",
      "############\n",
      "#cnf(i_0_1270, plain, (m(b,b)=X1|m(a,a)!=two|m(X2,X1)!=X2)).\n",
      "##################\n",
      "#cnf(i_0_13959, plain, (m(two,m(X1,X2))!=m(X2,X1)|~d(a,X2)|~d(a,X1))).\n",
      "###############\n",
      "#cnf(i_0_1395, plain, (X1=m(b,b)|m(a,a)!=two|m(X1,X2)!=X2)).\n",
      "###############################\n",
      "#cnf(i_0_2178, plain, (X1=one|f(X1,X3)!=X3|~d(m(X1,X2),X3))).\n",
      "######\n",
      "#cnf(i_0_14309, plain, (f(two,X1)!=X1|~d(m(a,a),X1))).\n",
      "###\n",
      "#cnf(i_0_14310, plain, (f(two,X1)!=X1|~d(m(a,m(a,X2)),X1))).\n",
      "########\n",
      "#cnf(i_0_14387, plain, (f(two,X1)!=X1|~d(m(X2,a),X2))).\n",
      "#\n",
      "#cnf(i_0_14451, plain, (f(two,X1)!=X1|~d(a,one))).\n",
      "#####\n",
      "#cnf(i_0_14454, plain, (f(two,X1)!=X1|~d(m(a,a),two))).\n",
      "##\n",
      "#cnf(i_0_14452, plain, (f(two,X1)!=X1|~d(m(a,X2),X2))).\n",
      "########\n",
      "#cnf(i_0_14468, plain, (m(two,X1)!=X1|~d(m(a,a),two))).\n",
      "#\n",
      "#cnf(i_0_14469, plain, (m(X1,two)!=X1|~d(m(a,a),two))).\n",
      "#########\n",
      "#cnf(i_0_14426, plain, (f(two,a)!=a|~d(m(a,m(a,X1)),two))).\n",
      "#\n",
      "#cnf(i_0_14455, plain, (f(two,X1)!=X1|~d(m(a,m(a,a)),two))).\n",
      "#\n",
      "#cnf(i_0_14456, plain, (f(two,X1)!=X1|~d(m(a,m(a,a)),b))).\n",
      "#####\n",
      "#cnf(i_0_14460, plain, (m(X1,m(X2,two))!=m(X1,X2)|~d(a,one))).\n",
      "###\n",
      "#cnf(i_0_14461, plain, (m(X1,m(two,X2))!=m(X2,X1)|~d(a,one))).\n",
      "###\n",
      "#cnf(i_0_14529, plain, (m(b,b)!=two|m(a,a)!=two|~d(a,one))).\n",
      "####\n",
      "#cnf(i_0_3032, plain, (d(two,m(X1,a))|~d(a,X1))).\n",
      "#\n",
      "#cnf(i_0_14527, plain, (m(X1,m(two,X2))!=m(X1,X2)|~d(a,one))).\n",
      "##############\n",
      "#cnf(i_0_14537, plain, (m(X2,m(X1,two))!=m(X1,X2)|~d(a,one))).\n",
      "########\n",
      "#cnf(i_0_14538, plain, (m(two,m(X2,X1))!=m(X1,X2)|~d(a,one))).\n",
      "################\n",
      "#cnf(i_0_14354, plain, (b=one|f(b,a)!=a|~d(m(a,a),two))).\n",
      "#\n",
      "#cnf(i_0_14370, plain, (b=one|f(b,X1)!=X1|~d(m(a,a),X1))).\n",
      "#########################\n",
      "#cnf(i_0_468, plain, (m(X1,X2)=two|m(X1,m(X2,m(b,b)))!=m(a,a))).\n",
      "###########\n",
      "#cnf(i_0_475, plain, (d(X1,m(a,a))|~d(X1,m(b,b)))).\n",
      "###########\n",
      "#cnf(i_0_488, plain, (m(X1,X2)=one|m(X3,m(X1,m(X2,X4)))!=m(X4,X3))).\n",
      "####\n",
      "#cnf(i_0_14458, plain, (f(two,X1)!=X1|~d(m(X2,m(X3,a)),X2))).\n",
      "##\n",
      "#cnf(i_0_14877, plain, (f(two,X1)!=X1|~d(m(X2,a),one))).\n",
      "##\n",
      "#cnf(i_0_14896, plain, (f(two,X1)!=X1|~d(m(a,X2),one))).\n",
      "###\n",
      "#cnf(i_0_489, plain, (X1=one|m(f(X1,X3),X2)!=m(X2,X3)|~d(X1,X3))).\n",
      "######\n",
      "#cnf(i_0_14875, plain, (m(a,a)!=two|f(two,X1)!=X1|~d(a,b))).\n",
      "###\n",
      "#cnf(i_0_14904, plain, (f(two,X1)!=X1|~d(X2,one)|~d(a,X2))).\n",
      "##\n",
      "#cnf(i_0_490, plain, (X1=one|m(X3,m(f(m(X1,X3),X4),X2))!=m(X2,X4)|~d(m(X1,X3),X4))).\n",
      "##\n",
      "#cnf(i_0_14986, plain, (m(two,X1)!=X1|~d(X2,one)|~d(a,X2))).\n",
      "##\n",
      "#cnf(i_0_14987, plain, (m(X1,two)!=X1|~d(X2,one)|~d(a,X2))).\n",
      "###\n",
      "#cnf(i_0_14980, plain, (m(a,a)!=two|m(two,X1)!=X1|~d(a,b))).\n",
      "#\n",
      "#cnf(i_0_14981, plain, (m(a,a)!=two|m(X1,two)!=X1|~d(a,b))).\n",
      "###\n",
      "#cnf(i_0_491, plain, (X1=one|m(X2,m(X4,m(X1,X3)))!=m(X3,m(X4,X2)))).\n",
      "###########\n",
      "#cnf(i_0_492, plain, (X1=one|m(X2,m(X3,m(X4,X1)))!=m(X3,m(X4,X2)))).\n",
      "##\n",
      "#cnf(i_0_14880, plain, (f(two,X1)!=X1|~d(m(X3,a),X2)|~d(X2,X3))).\n",
      "#########\n",
      "#cnf(i_0_497, plain, (X1=one|m(X2,m(X3,m(X1,X4)))!=m(X4,m(X2,X3)))).\n",
      "##\n",
      "#cnf(i_0_15311, plain, (f(two,X1)!=X1|~d(m(X2,X3),X2)|~d(a,X3))).\n",
      "###\n",
      "#cnf(i_0_15431, plain, (f(two,X1)!=X1|~d(m(X3,X2),X2)|~d(a,X3))).\n",
      "######\n",
      "#cnf(i_0_501, plain, (X1=one|m(X2,f(m(X2,X1),X3))!=X3|~d(m(X2,X1),X3))).\n",
      "###########\n",
      "#cnf(i_0_504, plain, (d(m(X1,two),m(X1,m(a,m(a,X2)))))).\n",
      "###########\n",
      "#cnf(i_0_505, plain, (m(b,m(b,m(X2,m(X1,two))))=m(X1,m(a,m(a,X2))))).\n",
      "##\n",
      "#cnf(i_0_14881, plain, (f(two,X1)!=X1|~d(m(a,X3),X2)|~d(X2,X3))).\n",
      "#########\n",
      "#cnf(i_0_506, plain, (X1=two|m(X1,m(b,m(b,X2)))!=m(a,m(a,X2)))).\n",
      "###########\n",
      "#cnf(i_0_509, plain, (f(two,m(a,m(a,X1)))=m(b,m(b,X1)))).\n",
      "######\n",
      "#cnf(i_0_14479, plain, (m(a,a)!=two|f(two,X1)!=X1|~d(a,m(b,b)))).\n",
      "#######\n",
      "#cnf(i_0_14459, plain, (f(two,X1)!=X1|~d(m(X2,m(X3,a)),X3))).\n",
      "###########\n",
      "#cnf(i_0_14505, plain, (f(two,X1)!=X1|~d(m(a,m(X2,X3)),X2))).\n",
      "########\n",
      "#cnf(i_0_15440, plain, (f(two,X1)!=X1|~d(m(a,a),b))).\n",
      "####\n",
      "#cnf(i_0_16226, plain, (m(two,X1)!=X1|~d(m(a,a),b))).\n",
      "#\n",
      "#cnf(i_0_16227, plain, (m(X1,two)!=X1|~d(m(a,a),b))).\n",
      "#######\n",
      "#cnf(i_0_516, plain, (m(b,m(b,X1))=one|m(a,m(a,X1))!=two)).\n",
      "##\n",
      "#cnf(i_0_16406, plain, (m(a,m(a,a))!=two|f(two,X1)!=X1)).\n",
      "####\n",
      "#cnf(i_0_16403, plain, (m(a,m(a,a))!=two|m(two,a)!=a)).\n",
      "#####\n",
      "#cnf(i_0_517, plain, (d(m(b,m(b,X1)),m(a,m(a,X1))))).\n",
      "####\n",
      "#cnf(i_0_14506, plain, (f(two,X1)!=X1|~d(m(a,m(X2,X3)),X3))).\n",
      "##\n",
      "#cnf(i_0_16621, plain, (m(a,m(a,a))!=two|m(two,X1)!=X1)).\n",
      "#\n",
      "#cnf(i_0_16622, plain, (m(a,m(a,a))!=two|m(X1,two)!=X1)).\n",
      "####\n",
      "#cnf(i_0_518, plain, (f(m(b,m(b,X1)),m(a,m(a,X1)))=two)).\n",
      "###########\n",
      "#cnf(i_0_519, plain, (d(m(b,m(b,X1)),a)|~d(m(a,m(a,X1)),two))).\n",
      "##\n",
      "#cnf(i_0_14873, plain, (f(two,X1)!=X1|~d(m(X2,m(a,X3)),X2))).\n",
      "####################\n",
      "#cnf(i_0_521, plain, (d(two,X1)|~d(m(a,m(a,X2)),X1))).\n",
      "###\n",
      "#cnf(i_0_14878, plain, (f(two,X1)!=X1|~d(m(X3,m(a,X2)),X2))).\n",
      "####\n",
      "#cnf(i_0_16258, plain, (d(b,one)|m(a,m(a,X1))!=two)).\n",
      "###\n",
      "#cnf(i_0_16273, plain, (d(b,a)|m(a,m(a,X1))!=two)).\n",
      "#\n",
      "#cnf(i_0_522, plain, (m(X1,m(a,m(a,X2)))!=m(b,m(b,m(X2,X1))))).\n",
      "######\n",
      "#cnf(i_0_16522, plain, (m(a,m(a,b))!=one|m(a,m(a,b))!=two)).\n",
      "##\n",
      "#cnf(i_0_16324, plain, (d(b,X1)|m(a,m(a,X2))!=two)).\n",
      "###\n",
      "#cnf(i_0_523, plain, (m(b,m(b,X1))=X2|m(a,m(a,X1))!=m(two,X2))).\n",
      "#########\n",
      "#cnf(i_0_16460, plain, (d(X1,one)|m(a,m(a,X1))!=two)).\n",
      "#\n",
      "#cnf(i_0_16467, plain, (d(X1,a)|m(a,m(a,X1))!=two)).\n",
      "#\n",
      "#cnf(i_0_524, plain, (m(b,m(b,X1))=one|m(X2,m(a,m(a,X1)))!=m(X2,two))).\n",
      "#######\n",
      "#cnf(i_0_16471, plain, (d(X1,X2)|m(a,m(a,X1))!=two)).\n",
      "####\n",
      "#cnf(i_0_525, plain, (d(m(b,m(b,X1)),X2)|~d(m(a,m(a,X1)),X2))).\n",
      "########\n",
      "#cnf(i_0_17412, plain, (d(b,X1)|m(a,m(X2,a))!=two)).\n",
      "##########\n",
      "#cnf(i_0_17423, plain, (d(b,X1)|m(X2,m(a,a))!=two)).\n",
      "####\n",
      "#cnf(i_0_527, plain, (X1=m(b,m(b,X2))|m(X1,two)!=m(a,m(a,X2)))).\n",
      "##########\n",
      "#cnf(i_0_17414, plain, (d(b,X1)|~d(m(a,a),two))).\n",
      "#\n",
      "#cnf(i_0_528, plain, (m(b,m(b,X1))=X2|m(X3,m(a,m(a,X1)))!=m(X3,m(two,X2)))).\n",
      "######\n",
      "#cnf(i_0_16562, plain, (m(X1,m(a,a))!=one|m(a,m(a,X1))!=two)).\n",
      "###\n",
      "#cnf(i_0_17710, plain, (d(X1,X2)|m(a,m(X1,a))!=two)).\n",
      "########\n",
      "#cnf(i_0_16564, plain, (m(a,m(a,X1))!=one|m(a,m(a,X1))!=two)).\n",
      "#####\n",
      "#cnf(i_0_530, plain, (d(m(two,b),m(a,m(a,X1))))).\n",
      "##\n",
      "#cnf(i_0_18045, plain, (a!=one|a!=two|~d(m(X1,a),X1))).\n",
      "#\n",
      "#cnf(i_0_18142, plain, (a!=one|a!=two|~d(a,one))).\n",
      "#\n",
      "#cnf(i_0_18143, plain, (a!=one|a!=two|~d(m(a,X1),X1))).\n",
      "###\n",
      "#cnf(i_0_18145, plain, (a!=one|a!=two|~d(m(a,a),two))).\n",
      "####\n",
      "#cnf(i_0_531, plain, (m(b,X1)=one|m(a,m(a,X1))!=m(two,b))).\n",
      "####\n",
      "#cnf(i_0_18046, plain, (m(a,X1)!=one|m(a,X1)!=two|~d(a,X1))).\n",
      "#######\n",
      "#cnf(i_0_532, plain, (m(b,X1)=X2|m(a,m(a,X1))!=m(two,m(b,X2)))).\n",
      "######\n",
      "#cnf(i_0_16565, plain, (m(a,m(X1,a))!=one|m(a,m(a,X1))!=two)).\n",
      "###\n",
      "#cnf(i_0_17721, plain, (d(X1,X2)|m(X1,m(a,a))!=two)).\n",
      "###\n",
      "#cnf(i_0_4887, plain, (d(m(X1,X2),X1)|~d(X2,one))).\n",
      "##########\n",
      "#cnf(i_0_535, plain, (m(b,m(b,X1))=m(b,b)|m(a,m(a,X1))!=m(a,a))).\n",
      "#######\n",
      "#cnf(i_0_18386, plain, (d(two,X1)|~d(a,one))).\n",
      "####\n",
      "#cnf(i_0_536, plain, (m(a,m(a,f(b,X1)))=m(two,m(b,X1))|~d(b,X1))).\n",
      "###########\n",
      "#cnf(i_0_537, plain, (m(a,m(a,m(X1,f(m(b,X1),X2))))=m(two,m(b,X2))|~d(m(b,X1),X2))).\n",
      "#\n",
      "#cnf(i_0_4921, plain, (d(b,m(X1,a))|~d(a,X1))).\n",
      "##########\n",
      "#cnf(i_0_538, plain, (m(two,m(b,m(X2,m(b,X1))))=m(a,m(a,m(X1,X2))))).\n",
      "########\n",
      "#cnf(i_0_18450, plain, (m(two,X1)!=X1|~d(m(a,a),one))).\n",
      "###\n",
      "#cnf(i_0_539, plain, (m(two,m(b,m(X1,m(X2,b))))=m(a,m(a,m(X1,X2))))).\n",
      "###########\n",
      "#cnf(i_0_541, plain, (m(two,m(b,m(X1,b)))=m(a,m(a,X1)))).\n",
      "######################\n",
      "#cnf(i_0_543, plain, (m(a,m(a,b))=m(two,b)|m(a,a)!=two)).\n",
      "##\n",
      "#cnf(i_0_20907, plain, (m(two,b)!=b|m(a,a)!=two)).\n",
      "####\n",
      "#cnf(i_0_20924, plain, (m(two,b)!=one|m(a,a)!=two)).\n",
      "###\n",
      "#cnf(i_0_20842, plain, (b=a|m(two,b)!=a|m(a,a)!=two)).\n",
      "##\n",
      "#cnf(i_0_544, plain, (m(a,m(a,f(m(b,b),X1)))=m(two,X1)|~d(m(b,b),X1))).\n",
      "########\n",
      "#cnf(i_0_20631, plain, (a=one|m(a,b)!=m(two,b)|m(a,a)!=two)).\n",
      "###\n",
      "#cnf(i_0_545, plain, (m(two,m(X1,m(b,b)))=m(a,m(a,X1)))).\n",
      "######################\n",
      "#cnf(i_0_547, plain, (m(a,m(a,f(b,X1)))!=m(b,X1)|~d(b,X1))).\n",
      "#####\n",
      "#cnf(i_0_21610, plain, (m(two,m(b,X1))!=m(b,X1)|~d(b,X1))).\n",
      "#####\n",
      "#cnf(i_0_21637, plain, (m(two,b)!=b|~d(b,one))).\n",
      "#\n",
      "#cnf(i_0_548, plain, (m(a,m(a,m(X1,f(m(b,X1),X2))))!=m(b,X2)|~d(m(b,X1),X2))).\n",
      "##\n",
      "#cnf(i_0_21642, plain, (m(a,a)!=two|m(two,X1)!=X1)).\n",
      "####\n",
      "#cnf(i_0_21658, plain, (m(b,b)!=two|m(a,a)!=two)).\n",
      "#####\n",
      "#cnf(i_0_555, plain, (m(a,m(a,f(m(b,b),X1)))!=X1|~d(m(b,b),X1))).\n",
      "#######\n",
      "#cnf(i_0_21755, plain, (m(two,X1)!=X1|~d(m(b,b),X1))).\n",
      "####\n",
      "#cnf(i_0_565, plain, (m(a,m(a,m(f(b,X1),X2)))=m(two,m(b,m(X1,X2)))|~d(b,X1))).\n",
      "#########\n",
      "#cnf(i_0_21759, plain, (m(two,a)!=a|~d(m(b,b),two))).\n",
      "##\n",
      "#cnf(i_0_566, plain, (m(f(b,X1),X2)=b|m(a,a)!=two|m(X1,X2)!=one)).\n",
      "###########\n",
      "#cnf(i_0_567, plain, (m(a,m(a,m(f(b,X1),X2)))!=m(b,m(X1,X2))|~d(b,X1))).\n",
      "###########\n",
      "#cnf(i_0_568, plain, (m(a,m(a,f(b,X1)))!=m(X1,b)|~d(b,X1))).\n",
      "###########\n",
      "#cnf(i_0_569, plain, (m(f(two,X1),X2)=m(b,b)|m(a,a)!=m(X1,X2)|~d(two,X1))).\n",
      "##\n",
      "#cnf(i_0_21638, plain, (m(two,m(X1,b))!=m(X1,b)|~d(b,X1))).\n",
      "#########\n",
      "#cnf(i_0_570, plain, (m(X2,m(f(f(X1,X2),X3),X4))=m(X1,m(X3,X4))|~d(f(X1,X2),X3)|~d(X1,X2))).\n",
      "###\n",
      "#cnf(i_0_21655, plain, (m(X1,m(two,b))!=m(b,X1)|~d(b,X1))).\n",
      "#############\n",
      "#cnf(i_0_22739, plain, (m(two,m(b,X1))!=m(X1,b)|~d(b,X1))).\n",
      "######\n",
      "#cnf(i_0_573, plain, (d(m(X1,X2),m(X1,m(X3,X4)))|~d(X2,X3))).\n",
      "###########\n",
      "#cnf(i_0_574, plain, (m(f(X2,X3),m(X4,m(X1,X2)))=m(X1,m(X3,X4))|~d(X2,X3))).\n",
      "##\n",
      "#cnf(i_0_23228, plain, (m(X1,m(two,b))!=m(X1,b)|~d(b,X1))).\n",
      "#########\n",
      "#cnf(i_0_575, plain, (X1=X2|m(X1,m(f(X2,X3),X4))!=m(X3,X4)|~d(X2,X3))).\n",
      "#############\n",
      "#cnf(i_0_23749, plain, (m(two,m(X1,b))!=m(b,X1)|~d(b,X1))).\n",
      "#########\n",
      "#cnf(i_0_578, plain, (f(X1,m(X2,X3))=m(f(X1,X2),X3)|~d(X1,X2))).\n",
      "#\n",
      "#cnf(i_0_24590, plain, (m(a,f(two,a))=m(b,b))).\n",
      "##\n",
      "#cnf(i_0_24913, plain, (f(two,X1)!=X1|~d(m(b,b),one))).\n",
      "##\n",
      "#cnf(i_0_24939, plain, (m(two,X1)!=X1|~d(m(b,b),one))).\n",
      "##\n",
      "#cnf(i_0_24882, plain, (a=two|m(b,b)!=a)).\n",
      "####\n",
      "#cnf(i_0_579, plain, (X1=one|m(f(X1,X2),X3)!=m(X2,X3)|~d(X1,X2))).\n",
      "##\n",
      "#cnf(i_0_24940, plain, (m(X1,two)!=X1|~d(m(b,b),one))).\n",
      "#\n",
      "#cnf(i_0_25011, plain, (m(a,a)!=two|m(X1,two)!=X1)).\n",
      "####\n",
      "#cnf(i_0_24952, plain, (a=two|m(a,a)!=two|a!=one)).\n",
      "###\n",
      "#cnf(i_0_24738, plain, (a=one|f(two,a)!=m(b,b))).\n",
      "#\n",
      "#cnf(i_0_580, plain, (m(X1,m(f(X1,X2),m(X3,f(m(X2,X3),X4))))=X4|~d(m(X2,X3),X4)|~d(X1,X2))).\n",
      "#\n",
      "#cnf(i_0_24735, plain, (d(a,m(b,b)))).\n",
      "#\n",
      "#cnf(i_0_14479, plain, (m(a,a)!=two|f(two,X1)!=X1)).\n",
      "#####\n",
      "#cnf(i_0_24775, plain, (f(two,a)=one|m(b,b)!=a)).\n",
      "#\n",
      "#cnf(i_0_25599, plain, (m(b,b)!=a|a!=one)).\n",
      "###\n",
      "#cnf(i_0_581, plain, (d(m(X2,X3),m(X1,m(f(X1,X2),m(X3,X4))))|~d(X1,X2))).\n",
      "##\n",
      "#cnf(i_0_25636, plain, (m(a,a)!=two|a!=one)).\n",
      "###\n",
      "#cnf(i_0_25607, plain, (m(b,b)!=a|m(X1,a)!=X1)).\n",
      "#\n",
      "#cnf(i_0_25609, plain, (m(b,b)!=a|m(a,X1)!=X1)).\n",
      "#####\n",
      "#cnf(i_0_582, plain, (m(X1,m(f(X1,X2),m(X3,X4)))=m(X4,m(X2,X3))|~d(X1,X2))).\n",
      "#\n",
      "#cnf(i_0_25588, plain, (d(a,one)|m(a,a)!=two)).\n",
      "##########\n",
      "#cnf(i_0_583, plain, (d(X1,a)|~d(m(X2,X3),two)|~d(X1,X2))).\n",
      "#\n",
      "#cnf(i_0_26397, plain, (d(a,X1)|m(a,a)!=two)).\n",
      "######\n",
      "#cnf(i_0_26469, plain, (m(a,a)!=two)).\n",
      "####\n",
      "#cnf(i_0_584, plain, (d(two,m(f(X1,X2),X3))|d(two,X1)|~d(two,m(X2,X3))|~d(X1,X2))).\n",
      "#\n",
      "#cnf(i_0_24737, plain, (f(a,m(b,b))=f(two,a))).\n",
      "##########\n",
      "#cnf(i_0_586, plain, (d(m(f(X1,X2),X3),m(X2,X3))|~d(X1,X2))).\n",
      "#\n",
      "#cnf(i_0_25598, plain, (m(a,X1)=m(two,X1)|m(b,b)!=a)).\n",
      "#\n",
      "#cnf(i_0_25012, plain, (m(X1,two)!=X1|~d(m(a,a),one))).\n",
      "#######\n",
      "#cnf(i_0_24731, plain, (f(two,X1)!=X1|~d(m(a,m(b,b)),a))).\n",
      "#\n",
      "#cnf(i_0_24924, plain, (f(two,X1)!=X1|~d(m(a,m(b,b)),X1))).\n",
      "#\n",
      "#cnf(i_0_587, plain, (f(m(f(X1,X2),X3),m(X2,X3))=X1|~d(X1,X2))).\n",
      "########\n",
      "#cnf(i_0_24715, plain, (X1=a|m(X1,f(two,a))!=m(b,b))).\n",
      "###\n",
      "#cnf(i_0_588, plain, (d(m(f(X1,X2),X3),a)|~d(m(X2,X3),two)|~d(X1,X2))).\n",
      "###########\n",
      "#cnf(i_0_589, plain, (X1=X2|m(f(X1,X3),m(X4,X2))!=m(X3,X4)|~d(X1,X3))).\n",
      "############\n",
      "#cnf(i_0_24776, plain, (d(f(two,a),m(b,b)))).\n",
      "######\n",
      "#cnf(i_0_26562, plain, (a=X1|m(f(two,a),X1)!=m(b,b))).\n",
      "####\n",
      "#cnf(i_0_591, plain, (d(X1,X2)|~d(m(X3,X4),X2)|~d(X1,X3))).\n",
      "#\n",
      "#cnf(i_0_24777, plain, (f(f(two,a),m(b,b))=a)).\n",
      "##########\n",
      "#cnf(i_0_592, plain, (X1=one|m(f(X1,X3),m(X4,X2))!=m(X2,m(X3,X4))|~d(X1,X3))).\n",
      "#\n",
      "#cnf(i_0_24806, plain, (d(a,m(X1,m(b,b))))).\n",
      "########\n",
      "#cnf(i_0_24814, plain, (a=one|f(a,X1)!=X1|~d(m(b,b),X1))).\n",
      "##\n",
      "#cnf(i_0_593, plain, (m(f(X1,X2),X3)=X4|m(X2,X3)!=m(X1,X4)|~d(X1,X2))).\n",
      "#\n",
      "#cnf(i_0_24885, plain, (d(two,m(a,m(b,b))))).\n",
      "##########\n",
      "#cnf(i_0_594, plain, (m(f(X1,X2),X3)=one|m(X4,m(X2,X3))!=m(X4,X1)|~d(X1,X2))).\n",
      "#\n",
      "#cnf(i_0_25605, plain, (m(two,X1)=m(X1,a)|m(b,b)!=a)).\n",
      "##########\n",
      "#cnf(i_0_595, plain, (d(m(f(X1,X2),X3),X4)|~d(m(X2,X3),X4)|~d(X1,X2))).\n",
      "#################\n",
      "#cnf(i_0_18067, plain, (m(a,m(X1,a))!=one|m(a,m(X1,a))!=two)).\n",
      "##\n",
      "#cnf(i_0_27759, plain, (m(X1,a)!=one|m(X1,a)!=two|~d(a,X1))).\n",
      "###\n",
      "#cnf(i_0_597, plain, (X1=m(f(X2,X3),X4)|m(X1,X2)!=m(X3,X4)|~d(X2,X3))).\n",
      "#\n",
      "#cnf(i_0_27057, plain, (d(f(two,a),m(a,a)))).\n",
      "##########\n",
      "#cnf(i_0_598, plain, (m(f(X1,X2),X3)=X4|m(X5,m(X2,X3))!=m(X5,m(X1,X4))|~d(X1,X2))).\n",
      "#######\n",
      "#cnf(i_0_18388, plain, (d(X1,a)|~d(m(X2,X1),one))).\n",
      "###\n",
      "#cnf(i_0_18390, plain, (d(X1,a)|~d(m(X1,X2),one))).\n",
      "###\n",
      "#cnf(i_0_14388, plain, (f(two,X1)!=X1|~d(m(a,X2),X1)|~d(a,X2))).\n",
      "###\n",
      "#cnf(i_0_28257, plain, (f(two,m(a,X1))!=m(a,X1)|~d(a,X1))).\n",
      "###\n",
      "#cnf(i_0_28314, plain, (m(two,a)!=a|~d(a,two))).\n",
      "###\n",
      "#cnf(i_0_600, plain, (m(X2,f(m(X1,f(X1,X2)),X3))=X3|~d(m(X1,f(X1,X2)),X3)|~d(X1,X2))).\n",
      "#\n",
      "#cnf(i_0_27333, plain, (d(a,m(b,m(b,X1))))).\n",
      "#\n",
      "#cnf(i_0_28273, plain, (f(two,m(X1,a))!=m(X1,a)|~d(a,X1))).\n",
      "####\n",
      "#cnf(i_0_18078, plain, (m(X1,m(a,a))!=one|m(X1,m(a,a))!=two)).\n",
      "##\n",
      "#cnf(i_0_18447, plain, (d(X1,X2)|~d(m(X3,X1),one))).\n",
      "###\n",
      "#cnf(i_0_601, plain, (d(m(X1,f(X1,X2)),m(X2,X3))|~d(X1,X2))).\n",
      "#####\n",
      "#cnf(i_0_28315, plain, (m(f(two,a),X1)!=m(a,X1)|~d(a,X1))).\n",
      "###\n",
      "#cnf(i_0_18449, plain, (d(X1,X2)|~d(m(X1,X3),one))).\n",
      "###\n",
      "#cnf(i_0_602, plain, (m(X3,m(X1,f(X1,X2)))=m(X2,X3)|~d(X1,X2))).\n",
      "###########\n",
      "#cnf(i_0_603, plain, (X1=one|m(X3,X1)!=m(X2,f(X2,X3))|~d(X2,X3))).\n",
      "###\n",
      "#cnf(i_0_28944, plain, (m(X1,f(two,a))!=m(a,X1)|~d(a,X1))).\n",
      "######\n",
      "#cnf(i_0_29506, plain, (X1=one|m(b,b)!=a|m(a,X1)!=two)).\n",
      "##\n",
      "#cnf(i_0_604, plain, (X1=X2|m(X4,X1)!=m(X3,m(f(X3,X4),X2))|~d(X3,X4))).\n",
      "############\n",
      "#cnf(i_0_27335, plain, (d(a,m(b,m(X1,b))))).\n",
      "##\n",
      "#cnf(i_0_14478, plain, (f(two,X1)!=X1|~d(X2,f(a,X2))|~d(a,X2))).\n",
      "########\n",
      "#cnf(i_0_606, plain, (f(X1,X2)=one|m(X2,X3)!=m(X3,X1)|~d(X1,X2))).\n",
      "##\n",
      "#cnf(i_0_14907, plain, (f(two,X1)!=X1|~d(m(X2,X3),one)|~d(a,X2))).\n",
      "############\n",
      "#cnf(i_0_14908, plain, (f(two,X1)!=X1|~d(m(X3,X2),one)|~d(a,X2))).\n",
      "#######\n",
      "#cnf(i_0_24803, plain, (d(a,X1)|~d(m(b,b),X1))).\n",
      "#\n",
      "#cnf(i_0_613, plain, (m(X2,f(f(X1,X2),X3))=m(X1,X3)|~d(f(X1,X2),X3)|~d(X1,X2))).\n",
      "#\n",
      "#cnf(i_0_30024, plain, (d(a,b)|~d(b,one))).\n",
      "#\n",
      "#cnf(i_0_14909, plain, (f(two,X1)!=X1|~d(m(a,X2),X3)|~d(X3,one))).\n",
      "#######\n",
      "#cnf(i_0_24934, plain, (d(b,X1)|m(a,m(b,b))!=two)).\n",
      "##\n",
      "#cnf(i_0_614, plain, (m(X2,m(X3,f(m(f(X1,X2),X3),X4)))=m(X1,X4)|~d(m(f(X1,X2),X3),X4)|~d(X1,X2))).\n",
      "##\n",
      "#cnf(i_0_14910, plain, (f(two,X1)!=X1|~d(m(X2,a),X3)|~d(X3,one))).\n",
      "#######\n",
      "#cnf(i_0_27327, plain, (d(a,one)|m(a,m(a,b))!=two)).\n",
      "##\n",
      "#cnf(i_0_615, plain, (m(X1,m(X4,m(f(X1,X2),X3)))=m(X2,m(X3,X4))|~d(X1,X2))).\n",
      "####\n",
      "#cnf(i_0_28238, plain, (f(two,X1)!=X1|~d(m(X2,a),X1)|~d(a,X2))).\n",
      "###\n",
      "#cnf(i_0_28694, plain, (d(a,b)|~d(m(X1,b),X1))).\n",
      "###\n",
      "#cnf(i_0_28722, plain, (d(a,b)|m(a,m(a,X1))!=two)).\n",
      "#\n",
      "#cnf(i_0_616, plain, (m(X1,m(X3,m(X4,f(X1,X2))))=m(X2,m(X3,X4))|~d(X1,X2))).\n",
      "#########\n",
      "#cnf(i_0_28740, plain, (d(a,one)|m(a,m(a,X1))!=two)).\n",
      "##\n",
      "#cnf(i_0_618, plain, (m(X1,m(X3,f(X1,X2)))=m(X2,X3)|~d(X1,X2))).\n",
      "##################\n",
      "#cnf(i_0_28812, plain, (d(b,X1)|~d(m(a,a),one))).\n",
      "##\n",
      "#cnf(i_0_28951, plain, (d(a,X1)|~d(m(b,b),one))).\n",
      "#\n",
      "#cnf(i_0_28956, plain, (d(two,X1)|~d(m(a,a),one))).\n",
      "############\n",
      "#cnf(i_0_622, plain, (d(X1,m(X2,m(X3,m(X1,X4)))))).\n",
      "###########\n",
      "#cnf(i_0_623, plain, (d(X1,m(X2,m(X3,m(X4,X1)))))).\n",
      "##########\n",
      "#cnf(i_0_31688, plain, (d(a,b)|~d(m(b,X1),X1))).\n",
      "#\n",
      "#cnf(i_0_624, plain, (d(X1,m(X2,m(X3,X4)))|~d(X1,X4))).\n",
      "#########\n",
      "#cnf(i_0_31690, plain, (d(a,b)|~d(m(a,b),two))).\n",
      "##\n",
      "#cnf(i_0_626, plain, (d(m(b,b),m(X1,m(a,a))))).\n",
      "#######\n",
      "#cnf(i_0_31732, plain, (d(a,b)|m(a,m(X1,a))!=two)).\n",
      "###\n",
      "#cnf(i_0_31734, plain, (d(a,b)|m(a,m(b,b))!=two)).\n",
      "#\n",
      "#cnf(i_0_627, plain, (d(m(b,m(b,X1)),m(X2,m(a,m(a,X1)))))).\n",
      "#######\n",
      "#cnf(i_0_31751, plain, (d(a,b)|m(X1,m(a,a))!=two)).\n",
      "#######\n",
      "#cnf(i_0_14411, plain, (f(two,X1)!=X1|~d(m(a,a),X2)|~d(X2,X1))).\n",
      "########\n",
      "#cnf(i_0_630, plain, (d(f(X1,X2),m(X3,X2))|~d(X1,X2))).\n",
      "#\n",
      "#cnf(i_0_29476, plain, (m(X1,two)=m(a,X1)|m(b,b)!=a)).\n",
      "#######\n",
      "#cnf(i_0_32385, plain, (d(a,one)|m(a,m(X1,a))!=two)).\n",
      "###\n",
      "#cnf(i_0_631, plain, (d(m(f(X1,X2),X3),m(X4,m(X2,X3)))|~d(X1,X2))).\n",
      "#######\n",
      "#cnf(i_0_32387, plain, (d(a,one)|m(a,m(b,b))!=two)).\n",
      "#\n",
      "#cnf(i_0_32405, plain, (d(a,one)|m(X1,m(a,a))!=two)).\n",
      "###\n",
      "#cnf(i_0_632, plain, (d(m(X1,f(m(X2,X1),X3)),m(X4,X3))|~d(m(X2,X1),X3))).\n",
      "####\n",
      "#cnf(i_0_14906, plain, (f(two,X1)!=X1|~d(X3,one)|~d(a,X2)|~d(X2,X3))).\n",
      "#######\n",
      "#cnf(i_0_633, plain, (d(m(X1,X2),m(X3,m(X2,m(X4,X1)))))).\n",
      "#####\n",
      "#cnf(i_0_34259, plain, (m(two,X1)!=X1|~d(X2,one)|~d(a,X3)|~d(X3,X2))).\n",
      "######\n",
      "#cnf(i_0_634, plain, (d(m(X1,X2),m(X3,m(X1,m(X2,X4)))))).\n",
      "###\n",
      "#cnf(i_0_34260, plain, (m(X1,two)!=X1|~d(X2,one)|~d(a,X3)|~d(X3,X2))).\n",
      "########\n",
      "#cnf(i_0_639, plain, (d(m(b,X1),m(a,m(a,X1))))).\n",
      "##############\n",
      "#cnf(i_0_14409, plain, (f(two,X1)!=X1|~d(m(a,m(X2,a)),X1))).\n",
      "###\n",
      "#cnf(i_0_34935, plain, (f(two,a)!=a|~d(m(a,m(X1,a)),two))).\n",
      "##########\n",
      "#cnf(i_0_14420, plain, (f(two,X1)!=X1|~d(m(X2,m(a,a)),X1))).\n",
      "####\n",
      "#cnf(i_0_34981, plain, (f(two,a)!=a|~d(m(X1,m(a,a)),two))).\n",
      "#############\n",
      "#cnf(i_0_652, plain, (X1=one|m(X2,f(X1,X3))!=m(X2,X3)|~d(X1,X3))).\n",
      "###########\n",
      "#cnf(i_0_653, plain, (X1=one|m(X2,m(f(X1,X3),X4))!=m(X2,m(X3,X4))|~d(X1,X3))).\n",
      "##\n",
      "#cnf(i_0_14445, plain, (f(two,m(a,X1))!=m(a,X1)|~d(m(a,X2),X1))).\n",
      "#########\n",
      "#cnf(i_0_654, plain, (X1=one|m(X2,m(X3,f(m(X1,X3),X4)))!=m(X2,X4)|~d(m(X1,X3),X4))).\n",
      "####\n",
      "#cnf(i_0_14894, plain, (f(two,X1)!=X1|~d(m(X2,m(X3,a)),one))).\n",
      "#######\n",
      "#cnf(i_0_655, plain, (X1=one|m(X2,m(X4,m(X1,X3)))!=m(X2,m(X3,X4)))).\n",
      "####\n",
      "#cnf(i_0_14918, plain, (f(two,X1)!=X1|~d(m(X3,m(a,X2)),one))).\n",
      "##################\n",
      "#cnf(i_0_662, plain, (X1=one|m(X2,m(X3,m(X1,X4)))!=m(X2,m(X3,X4)))).\n",
      "###########\n",
      "#cnf(i_0_666, plain, (f(X1,X2)=one|m(X2,X3)!=m(X1,X3)|~d(X1,X2))).\n",
      "#####\n",
      "#cnf(i_0_18149, plain, (a!=one|a!=two|~d(m(X1,m(X2,a)),X1))).\n",
      "##\n",
      "#cnf(i_0_35811, plain, (m(a,m(a,a))!=two|a!=one|a!=two)).\n",
      "#\n",
      "#cnf(i_0_35815, plain, (a!=one|a!=two|~d(m(X1,a),one))).\n",
      "######\n",
      "#cnf(i_0_35848, plain, (a!=one|a!=two|~d(m(a,X1),one))).\n",
      "##########\n",
      "#cnf(i_0_35859, plain, (a!=one|a!=two|~d(X1,one)|~d(a,X1))).\n",
      "#########\n",
      "#cnf(i_0_672, plain, (m(b,m(X1,m(b,X2)))!=m(a,m(a,m(X1,X2))))).\n",
      "######\n",
      "#cnf(i_0_35889, plain, (a!=one|a!=two|~d(m(b,b),one))).\n",
      "#####\n",
      "#cnf(i_0_677, plain, (m(b,b)=m(X1,X2)|m(X1,m(two,X2))!=m(a,a))).\n",
      "###\n",
      "#cnf(i_0_35818, plain, (a!=one|a!=two|~d(m(X2,a),X1)|~d(X1,X2))).\n",
      "########\n",
      "#cnf(i_0_678, plain, (m(X1,m(X3,m(f(X1,X2),X4)))=m(X2,m(X3,X4))|~d(X1,X2))).\n",
      "##\n",
      "#cnf(i_0_36112, plain, (a!=one|a!=two|~d(m(X1,X2),X1)|~d(a,X2))).\n",
      "###########\n",
      "#cnf(i_0_36842, plain, (a!=one|a!=two|~d(m(X2,X1),X1)|~d(a,X2))).\n",
      "#####################\n",
      "#cnf(i_0_18507, plain, (d(m(X2,X1),X1)|~d(X2,one))).\n",
      "#####################\n",
      "#cnf(i_0_684, plain, (X1=X2|m(X1,m(X3,X4))!=m(X3,m(X2,X4)))).\n",
      "##############################\n",
      "#cnf(i_0_37019, plain, (X1=one|f(X1,X2)!=X2|~d(X1,one))).\n",
      "##############\n",
      "#cnf(i_0_690, plain, (m(X1,m(X2,m(X3,f(m(X2,m(X1,X3)),X4))))=X4|~d(m(X2,m(X1,X3)),X4))).\n",
      "###########\n",
      "#cnf(i_0_691, plain, (d(m(X2,m(X1,X3)),m(X1,m(X2,m(X3,X4)))))).\n",
      "######\n",
      "#cnf(i_0_23787, plain, (m(b,m(X1,b))!=two|m(a,m(a,X1))!=two)).\n",
      "###########\n",
      "#cnf(i_0_24911, plain, (m(a,m(b,b))!=one|m(a,m(b,b))!=two)).\n",
      "######################\n",
      "#cnf(i_0_27279, plain, (m(X2,m(a,X1))!=m(X1,X2)|m(b,b)!=a)).\n",
      "#####\n",
      "#cnf(i_0_695, plain, (d(two,m(X1,X2))|d(two,X3)|~d(two,m(X1,m(X3,X2))))).\n",
      "###\n",
      "#cnf(i_0_35819, plain, (a!=one|a!=two|~d(m(a,X2),X1)|~d(X1,X2))).\n",
      "#########\n",
      "#cnf(i_0_24720, plain, (d(X1,m(b,b))|~d(X1,a))).\n",
      "######\n",
      "#cnf(i_0_38677, plain, (d(two,b))).\n",
      "\n",
      "# Proof found!\n",
      "# SZS status Unsatisfiable\n",
      "# SZS output start CNFRefutation\n",
      "cnf(sos, axiom, (d(X1,X3)|m(X1,X2)!=X3), file('/tmp/sqrt.p', sos)).\n",
      "cnf(sos, axiom, (m(X1,f(X1,X2))=X2|~d(X1,X2)), file('/tmp/sqrt.p', sos)).\n",
      "cnf(sos, axiom, (m(X1,m(X2,X3))=m(m(X1,X2),X3)), file('/tmp/sqrt.p', sos)).\n",
      "cnf(sos, axiom, (X2=X3|m(X1,X2)!=m(X1,X3)), file('/tmp/sqrt.p', sos)).\n",
      "cnf(sos, axiom, (d(two,X1)|d(two,X2)|~d(two,m(X1,X2))), file('/tmp/sqrt.p', sos)).\n",
      "cnf(sos, axiom, (m(a,a)=m(two,m(b,b))), file('/tmp/sqrt.p', sos)).\n",
      "cnf(sos, axiom, (m(X1,X2)=m(X2,X1)), file('/tmp/sqrt.p', sos)).\n",
      "cnf(sos, axiom, (X1=one|~d(X1,a)|~d(X1,b)), file('/tmp/sqrt.p', sos)).\n",
      "cnf(sos, axiom, (two!=one), file('/tmp/sqrt.p', sos)).\n",
      "cnf(c_0_9, plain, (d(X1,X3)|m(X1,X2)!=X3), inference(fof_simplification,[status(thm)],[sos])).\n",
      "cnf(c_0_10, plain, (m(X1,f(X1,X2))=X2|~d(X1,X2)), inference(fof_simplification,[status(thm)],[sos])).\n",
      "cnf(c_0_11, plain, (d(X1,X3)|m(X1,X2)!=X3), c_0_9).\n",
      "cnf(c_0_12, plain, (m(X1,f(X1,X2))=X2|~d(X1,X2)), c_0_10).\n",
      "cnf(c_0_13, axiom, (m(X1,m(X2,X3))=m(m(X1,X2),X3)), sos).\n",
      "cnf(c_0_14, plain, (X2=X3|m(X1,X2)!=m(X1,X3)), inference(fof_simplification,[status(thm)],[sos])).\n",
      "cnf(c_0_15, plain, (d(X1,m(X1,X2))), inference(er,[status(thm)],[c_0_11])).\n",
      "cnf(c_0_16, plain, (m(X1,m(X2,f(m(X1,X2),X3)))=X3|~d(m(X1,X2),X3)), inference(spm,[status(thm)],[c_0_12, c_0_13])).\n",
      "cnf(c_0_17, plain, (X2=X3|m(X1,X2)!=m(X1,X3)), c_0_14).\n",
      "cnf(c_0_18, plain, (d(two,X1)|d(two,X2)|~d(two,m(X1,X2))), inference(fof_simplification,[status(thm)],[sos])).\n",
      "cnf(c_0_19, axiom, (m(a,a)=m(two,m(b,b))), sos).\n",
      "cnf(c_0_20, plain, (d(X1,X2)|~d(m(X1,X3),X2)), inference(spm,[status(thm)],[c_0_15, c_0_16])).\n",
      "cnf(c_0_21, plain, (f(X1,m(X1,X2))=X2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_12])]), c_0_15])])).\n",
      "cnf(c_0_22, plain, (m(X1,m(f(X1,X2),X3))=m(X2,X3)|~d(X1,X2)), inference(spm,[status(thm)],[c_0_13, c_0_12])).\n",
      "cnf(c_0_23, plain, (d(two,X1)|d(two,X2)|~d(two,m(X1,X2))), c_0_18).\n",
      "cnf(c_0_24, plain, (d(two,m(a,a))), inference(spm,[status(thm)],[c_0_15, c_0_19])).\n",
      "cnf(c_0_25, plain, (d(X1,X2)|~d(X3,X2)|~d(X1,X3)), inference(spm,[status(thm)],[c_0_20, c_0_12])).\n",
      "cnf(c_0_26, plain, (f(two,m(a,a))=m(b,b)), inference(spm,[status(thm)],[c_0_21, c_0_19])).\n",
      "cnf(c_0_27, plain, (f(X1,m(X2,X3))=m(f(X1,X2),X3)|~d(X1,X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).\n",
      "cnf(c_0_28, plain, (d(two,a)), inference(spm,[status(thm)],[c_0_23, c_0_24])).\n",
      "cnf(c_0_29, axiom, (m(X1,X2)=m(X2,X1)), sos).\n",
      "cnf(c_0_30, plain, (d(X1,m(X2,X3))|~d(X1,X2)), inference(spm,[status(thm)],[c_0_25, c_0_15])).\n",
      "cnf(c_0_31, plain, (m(a,f(two,a))=m(b,b)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])]), c_0_29])).\n",
      "cnf(c_0_32, plain, (X1=one|~d(X1,a)|~d(X1,b)), inference(fof_simplification,[status(thm)],[sos])).\n",
      "cnf(c_0_33, plain, (d(X1,m(b,b))|~d(X1,a)), inference(spm,[status(thm)],[c_0_30, c_0_31])).\n",
      "cnf(c_0_34, plain, (two!=one), inference(fof_simplification,[status(thm)],[sos])).\n",
      "cnf(c_0_35, plain, (X1=one|~d(X1,a)|~d(X1,b)), c_0_32).\n",
      "cnf(c_0_36, plain, (d(two,b)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_33]), c_0_28])])).\n",
      "cnf(c_0_37, plain, (two!=one), c_0_34).\n",
      "cnf(c_0_38, plain, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_28])]), c_0_37]), ['proof']).\n",
      "# SZS output end CNFRefutation\n"
     ]
    }
   ],
   "source": [
    "! eprover-ho --auto-schedule --proof-object /tmp/sqrt.p"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "Anyhow, with that guidance, we can attempt to write a proof like thing in python.\n",
    "\n",
    "I've been tinkering with trying to enforce a proof discipline / kernel in python, but I've been leaning towards reducing it all to idioms until I'm more sure what I want. \n",
    "\n",
    "- https://github.com/philzook58/knuckledragger github\n",
    "- https://www.philipzucker.com/knuckledrag2/ second blog post\n",
    "- https://www.philipzucker.com/python-itp/ first blog post\n",
    "\n",
    "The kernel is only enforcing rules that I think make sense anyway.\n",
    "\n",
    "A small wrapper enables us the make stepwise proofs. We only give python names to things that we want to accept as axioms or things that come out of `lemma`. The justifications that go into `by` should only be previously named theorems.\n",
    "\n",
    "This avoids all the rigamorole of keeping a separate `Theorem` datatype or theorem database, but is more susceptible to accidental error than the styles I've used in those other posts."
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 12,
   "metadata": {},
   "outputs": [],
   "source": [
    "def lemma(thm, by=[]):\n",
    "    prove(Implies(And(by), thm))\n",
    "    return thm"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "The following let's me see last assignment, prewtty printing the statement of the lemma in a nice form."
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 14,
   "metadata": {},
   "outputs": [],
   "source": [
    "from IPython.core.interactiveshell import InteractiveShell\n",
    "InteractiveShell.ast_node_interactivity = \"last_expr_or_assign\""
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "The strategy is to make opaque functions that we can hide the definitions that may be confusing the solver. The approach semi fails."
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 23,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "proved\n"
     ]
    },
    {
     "data": {
      "text/html": [
       "&forall;x : mul(x, 1) = x"
      ],
      "text/plain": [
       "ForAll(x, mul(x, 1) == x)"
      ]
     },
     "execution_count": 23,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "Z = IntSort()\n",
    "B = BoolSort()\n",
    "x,y,z = Ints(\"x y z\")\n",
    "mul = Function(\"mul\", Z, Z, Z)\n",
    "mul_def = ForAll([x,y], mul(x, y) == x*y)"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "Z3 easily proves the basic algebraic properties of multiplication"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 33,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "proved\n"
     ]
    },
    {
     "data": {
      "text/html": [
       "&forall;x : mul(x, 1) = x"
      ],
      "text/plain": [
       "ForAll(x, mul(x, 1) == x)"
      ]
     },
     "execution_count": 33,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "mul_x_1 = lemma(ForAll([x], mul(x, 1) == x), by=[mul_def])"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 17,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "proved\n"
     ]
    },
    {
     "data": {
      "text/html": [
       "&forall;x : mul(1, x) = x"
      ],
      "text/plain": [
       "ForAll(x, mul(1, x) == x)"
      ]
     },
     "execution_count": 17,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "mul_1_x = lemma(ForAll([x], mul(1, x) == x), by=[mul_def])"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 18,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "proved\n"
     ]
    },
    {
     "data": {
      "text/html": [
       "&forall;x, y, z : mul(x, mul(y, z)) = mul(mul(x, y), z)"
      ],
      "text/plain": [
       "ForAll([x, y, z], mul(x, mul(y, z)) == mul(mul(x, y), z))"
      ]
     },
     "execution_count": 18,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "mul_assoc = lemma(ForAll([x,y,z], mul(x, mul(y,z)) == mul(mul(x,y), z)), by=[mul_def])"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 19,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "proved\n"
     ]
    },
    {
     "data": {
      "text/html": [
       "&forall;x, y : mul(x, y) = mul(y, x)"
      ],
      "text/plain": [
       "ForAll([x, y], mul(x, y) == mul(y, x))"
      ]
     },
     "execution_count": 19,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "mul_comm = lemma(ForAll([x,y], mul(x,y) == mul(y,x)), by=[mul_def])"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 20,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "proved\n"
     ]
    },
    {
     "data": {
      "text/html": [
       "&forall;x, y, z : x &ne; 0 &and; mul(x, y) = mul(x, z) &rArr; y = z"
      ],
      "text/plain": [
       "ForAll([x, y, z],\n",
       "       Implies(And(x != 0, mul(x, y) == mul(x, z)), y == z))"
      ]
     },
     "execution_count": 20,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "mul_cancel = lemma(ForAll([x,y,z], Implies(And(x != 0, mul(x,y) == mul(x,z)), y == z)), by=[mul_def])"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "We can define a [divides](https://en.wikipedia.org/wiki/Divisor) relation. I define it in terms of `x/y` which is flooring division. Is there a built in divides relation in SMT?"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 24,
   "metadata": {},
   "outputs": [
    {
     "data": {
      "text/html": [
       "&forall;x, y : divides(x, y) = (x = y&middot;(x/y))"
      ],
      "text/plain": [
       "ForAll([x, y], divides(x, y) == (x == y*(x/y)))"
      ]
     },
     "execution_count": 24,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "divides = Function(\"divides\", Z, Z, B)\n",
    "divides_def = ForAll([x,y], divides(x,y) == (x == y*(x/y)))"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 26,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "proved\n"
     ]
    },
    {
     "data": {
      "text/html": [
       "&forall;x, y : divides(x, y) &rArr; mul(x/y, y) = x"
      ],
      "text/plain": [
       "ForAll([x, y], Implies(divides(x, y), mul(x/y, y) == x))"
      ]
     },
     "execution_count": 26,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "div_mul_cancel = lemma(ForAll([x,y], Implies(divides(x,y), mul(x/y, y) == x)), by=[divides_def, mul_def])"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "Here we have two properties that I can't get to pass. They are basic facts. We don't get a counterexample, so perhaps that lends a mild amount of credence."
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 29,
   "metadata": {},
   "outputs": [
    {
     "data": {
      "text/html": [
       "&forall;x, y, z : mul(x, y) = z &rArr; divides(z, x)"
      ],
      "text/plain": [
       "ForAll([x, y, z], Implies(mul(x, y) == z, divides(z, x)))"
      ]
     },
     "execution_count": 29,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "#mul_divides = lemma(ForAll([x,y,z], Implies(mul(x,y) == z, divides(z,x))), by=[mul_def, divides_def])\n",
    "# can't get this one to prove\n",
    "mul_divides = ForAll([x,y,z], Implies(mul(x,y) == z, divides(z,x)))"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 27,
   "metadata": {},
   "outputs": [
    {
     "data": {
      "text/html": [
       "&forall;x, y :\n",
       " divides(mul(x, y), 2) &rArr; divides(x, 2) &or; divides(y, 2)"
      ],
      "text/plain": [
       "ForAll([x, y],\n",
       "       Implies(divides(mul(x, y), 2),\n",
       "               Or(divides(x, 2), divides(y, 2))))"
      ]
     },
     "execution_count": 27,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "#two_prime = lemma(ForAll([x,y], Implies(divides(mul(x,y),2), Or(divides(x,2), divides(y,2)))), by=[divides_def, mul_def])\n",
    "# can't get this one to pass\n",
    "two_prime = ForAll([x,y], Implies(divides(mul(x,y),2), Or(divides(x,2), divides(y,2))))"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "Part of the cleverness of the Otter proof is abusing this notion that every fraction can be expressed in lowest common terms. This is a nontrivial seeming fact.\n",
    "\n",
    "Here I use an expanded form of `solve` because I was fiddling with the smt formula."
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 28,
   "metadata": {},
   "outputs": [
    {
     "data": {
      "text/html": [
       "<b>unsat</b>"
      ],
      "text/plain": [
       "unsat"
      ]
     },
     "execution_count": 28,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "no_common_factor = ForAll([z], Implies(And(divides(x,z), divides(y,z)), z == 1))\n",
    "prime_thm = Not(Exists([x,y], And(no_common_factor, x != 0, mul(x,x) == mul(2,mul(y,y)))))\n",
    "lemmas = [mul_1_x, mul_x_1, mul_comm, mul_assoc, mul_divides, mul_cancel, div_mul_cancel, two_prime]\n",
    "s = Solver()\n",
    "s.add(lemmas)\n",
    "s.add(Not(prime_thm))\n",
    "with open(\"/tmp/prime.smt2\", \"w\") as f:\n",
    "    f.write(s.sexpr()) \n",
    "#lemma(Not(Exists([x,y], And(no_factor, mul(x,x) == mul(2,mul(y,y))))), by=[mul_1_x, mul_x_1, mul_assoc, two_prime, div_mul_cancel])\n",
    "s.check()"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "## Bits and Bobbles\n",
    "\n",
    "I am not satisfied. There are big holes and there was some nontrivial workarounds. I'm not particularly sure my thrashing even helped that much.\n",
    "\n"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 31,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "% Running in auto input_syntax mode. Trying SMTLIB2\n",
      "tff(func_def_0, type, mul: ($int * $int) > $int).\n",
      "tff(func_def_12, type, sK3: $int).\n",
      "tff(func_def_13, type, sK4: $int).\n",
      "tff(pred_def_1, type, divides: ($int * $int) > $o).\n",
      "tff(pred_def_5, type, lG2: ($int * $int) > $o).\n",
      "tff(u63,axiom,\n",
      "    (lG2(sK4,sK3) | ~sLP1)).\n",
      "\n",
      "tff(u62,axiom,\n",
      "    ((0 != sK3) | ~sLP1)).\n",
      "\n",
      "tff(u61,axiom,\n",
      "    ((mul(2,mul(sK4,sK4)) = mul(sK3,sK3)) | ~sLP1)).\n",
      "\n",
      "tff(u64,axiom,\n",
      "    (![X2 : $int, X0 : $int, X1 : $int] : (((1 = X2) | ~divides(X1,X2) | ~divides(X0,X2) | ~lG2(X0,X1))))).\n",
      "\n",
      "tff(u35,axiom,\n",
      "    (![X0 : $o] : (((($true = X0)) | (($false = X0)))))).\n",
      "\n",
      "tff(u34,axiom,\n",
      "    (($true != $false))).\n",
      "\n",
      "tff(u33,axiom,\n",
      "    (![X0 : $int, X1 : $int] : ((~$less(X0,X1) | ~$less(X1,$sum(X0,1)))))).\n",
      "\n",
      "tff(u32,axiom,\n",
      "    (![X2 : $int, X1 : $int] : (((0 = X2) | ~$less($sum($abs(X2),$uminus(1)),$remainder_e(X1,X2)))))).\n",
      "\n",
      "tff(u31,axiom,\n",
      "    (![X2 : $int, X1 : $int] : (((0 = X2) | ~$less($remainder_e(X1,X2),0))))).\n",
      "\n",
      "tff(u30,axiom,\n",
      "    (![X2 : $int, X1 : $int] : (((0 = X2) | ($sum($remainder_e(X1,X2),$product(X2,$quotient_e(X1,X2))) = X1))))).\n",
      "\n",
      "tff(u29,axiom,\n",
      "    (![X1 : $int] : ((~$less(X1,0) | ($uminus(X1) = $abs(X1)))))).\n",
      "\n",
      "tff(u28,axiom,\n",
      "    (![X1 : $int] : ((~$less(0,X1) | ($abs(X1) = X1))))).\n",
      "\n",
      "tff(u74,axiom,\n",
      "    (![X2 : $int, X0 : $int, X3 : $int] : (((0 = X0) | ($product(X0,X2) != $product(X0,X3)) | (X2 = X3))))).\n",
      "\n",
      "tff(u26,axiom,\n",
      "    (![X2 : $int, X0 : $int, X1 : $int] : (($product(X0,$sum(X1,X2)) = $sum($product(X0,X1),$product(X0,X2)))))).\n",
      "\n",
      "tff(u25,axiom,\n",
      "    (![X0 : $int] : ((0 = $product(X0,0))))).\n",
      "\n",
      "tff(u24,axiom,\n",
      "    (![X0 : $int] : (($product(X0,1) = X0)))).\n",
      "\n",
      "tff(u23,axiom,\n",
      "    (![X2 : $int, X0 : $int, X1 : $int] : (($product(X0,$product(X1,X2)) = $product($product(X0,X1),X2))))).\n",
      "\n",
      "tff(u22,axiom,\n",
      "    (![X0 : $int, X1 : $int] : (($product(X0,X1) = $product(X1,X0))))).\n",
      "\n",
      "tff(u21,axiom,\n",
      "    (![X0 : $int] : (($uminus($uminus(X0)) = X0)))).\n",
      "\n",
      "tff(u20,axiom,\n",
      "    (![X0 : $int, X1 : $int] : (($less(X0,X1) | $less(X1,$sum(X0,1)))))).\n",
      "\n",
      "tff(u19,axiom,\n",
      "    (![X2 : $int, X0 : $int, X1 : $int] : ((~$less(X0,X1) | $less($sum(X0,X2),$sum(X1,X2)))))).\n",
      "\n",
      "tff(u18,axiom,\n",
      "    (![X0 : $int, X1 : $int] : (($less(X0,X1) | $less(X1,X0) | (X0 = X1))))).\n",
      "\n",
      "tff(u17,axiom,\n",
      "    (![X2 : $int, X0 : $int, X1 : $int] : ((~$less(X0,X1) | ~$less(X1,X2) | $less(X0,X2))))).\n",
      "\n",
      "tff(u16,axiom,\n",
      "    (![X0 : $int] : (~$less(X0,X0)))).\n",
      "\n",
      "tff(u15,axiom,\n",
      "    (![X0 : $int] : ((0 = $sum(X0,$uminus(X0)))))).\n",
      "\n",
      "tff(u14,axiom,\n",
      "    (![X0 : $int, X1 : $int] : (($uminus($sum(X0,X1)) = $sum($uminus(X1),$uminus(X0)))))).\n",
      "\n",
      "tff(u13,axiom,\n",
      "    (![X0 : $int] : (($sum(X0,0) = X0)))).\n",
      "\n",
      "tff(u12,axiom,\n",
      "    (![X2 : $int, X0 : $int, X1 : $int] : (($sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2))))).\n",
      "\n",
      "tff(u11,axiom,\n",
      "    (![X0 : $int, X1 : $int] : (($sum(X0,X1) = $sum(X1,X0))))).\n",
      "\n",
      "tff(u65,hypothesis,\n",
      "    (![X0 : $int] : ((mul(1,X0) = X0)))).\n",
      "\n",
      "tff(u66,hypothesis,\n",
      "    (![X0 : $int] : ((mul(X0,1) = X0)))).\n",
      "\n",
      "tff(u67,hypothesis,\n",
      "    (![X0 : $int, X1 : $int] : ((mul(X1,X0) = mul(X0,X1))))).\n",
      "\n",
      "tff(u68,hypothesis,\n",
      "    (![X2 : $int, X0 : $int, X1 : $int] : ((mul(mul(X1,X0),X2) = mul(X1,mul(X0,X2)))))).\n",
      "\n",
      "tff(u75,hypothesis,\n",
      "    (![X0 : $int, X1 : $int] : (divides(mul(X1,X0),X1)))).\n",
      "\n",
      "tff(u70,hypothesis,\n",
      "    (![X2 : $int, X0 : $int, X1 : $int] : (((X0 = X2) | (0 = X1) | (mul(X1,X0) != mul(X1,X2)))))).\n",
      "\n",
      "tff(u71,hypothesis,\n",
      "    (![X0 : $int, X1 : $int] : (((mul($quotient_e(X1,X0),X0) = X1) | ~divides(X1,X0))))).\n",
      "\n",
      "tff(u72,hypothesis,\n",
      "    (![X0 : $int, X1 : $int] : ((divides(X1,2) | divides(X0,2) | ~divides(mul(X1,X0),2))))).\n",
      "\n",
      "tff(u73,hypothesis,\n",
      "    sLP1).\n",
      "\n"
     ]
    }
   ],
   "source": [
    "!vampire --mode tclausify /tmp/prime.smt2  # vampire is useful for converting to clauses or smt to tptp"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "I think I have less interest in rigor than others who might be interested in such things.\n",
    "\n",
    "\n",
    "I was once stumped on how to derive the irrationality of the square root of 2 in Z3. Here I'm going to recpliate it.\n",
    "I suspect ATP systems like vampire or eprover will not get stuck\n",
    "\n",
    "It can be formulated as proving `not exists m n, n^2 = 2 m^2`  $\\neg \\exists m n, n^2 = 2 m^2$\n",
    "\n",
    "Cody also suggested infinitude of primes.\n",
    "\n",
    "This is all math shit. These are not my theorems. I think maybe they are the wrong ones to attack.\n",
    "\n",
    "Facts about evens and odds."
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 1,
   "metadata": {},
   "outputs": [],
   "source": [
    "#from cvc5.pythonic import *\n",
    "from z3 import *\n",
    "x,y,z = Ints(\"x y z\")\n",
    "p = Real(\"p\")\n",
    "Z = IntSort()\n",
    "R = RealSort()\n",
    "B = BoolSort()"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "The paper mentions that many systems don't have the reals. Z3 and CVC5 do. So you can state something closer to what you meant.\n",
    "\n",
    "In a certain sense, the SMT \"reals\" are not the reals at all. They can prove some under approximation of reasonable facts about the reals. They are based around the rationals, linear equations, linear inequalities, polynomials and algebraic numbers.\n"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 2,
   "metadata": {},
   "outputs": [
    {
     "data": {
      "text/html": [
       "&forall;p : (&exist;x, y : p = ToReal(x)/ToReal(y)) = rational(p)"
      ],
      "text/plain": [
       "ForAll(p,\n",
       "       (Exists([x, y], p == ToReal(x)/ToReal(y))) ==\n",
       "       rational(p))"
      ]
     },
     "execution_count": 2,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "rational = Function(\"rational\", R, B)\n",
    "rat_def = ForAll([p], rational(p) == Exists([x,y], p == ToReal(x)/ToReal(y)))\n",
    "rat_def"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 3,
   "metadata": {},
   "outputs": [],
   "source": [
    "step1 = Implies(And(rational(p), p*p == 2), Exists([x,y], x*x == 2*y*y))\n",
    "prove(Implies(rat_def, step1))\n"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": null,
   "metadata": {},
   "outputs": [
    {
     "ename": "NameError",
     "evalue": "name 'rational' is not defined",
     "output_type": "error",
     "traceback": [
      "\u001b[0;31m---------------------------------------------------------------------------\u001b[0m",
      "\u001b[0;31mNameError\u001b[0m                                 Traceback (most recent call last)",
      "Cell \u001b[0;32mIn[13], line 1\u001b[0m\n\u001b[0;32m----> 1\u001b[0m prove(ForAll([p], Implies(\u001b[43mrational\u001b[49m(p), Not(p\u001b[38;5;241m*\u001b[39mp \u001b[38;5;241m==\u001b[39m \u001b[38;5;241m2\u001b[39m))))\n",
      "\u001b[0;31mNameError\u001b[0m: name 'rational' is not defined"
     ]
    }
   ],
   "source": [
    "#prove(ForAll([p], Implies(rational(p), Not(p*p == 2))))"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "Lemmas about even and oddness"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "I have a suspicion that it is wise to define things in \"open\" and explicitly skolemized form.\n",
    "\n",
    "`evenp(x,y)` is the \"evidence\" form of evenness, where you give explicitly the number that is half of x. This perhaps gives you the option of helping the solver find the needed existential by just giving it.\n",
    " `even(x) = exists p, evenp(x,p)`\n",
    "\n",
    "Explicit skolem function help you refer to something that is guaranteed to exist.\n",
    "\n",
    "Choice functions would make this easier probably.\n",
    "\n"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": null,
   "metadata": {},
   "outputs": [],
   "source": [
    "evenp = Function(\"evenp\", I, I, B)\n",
    "evenp_def = ForAll([x,y], evenp(x,y) == (x == 2*y))\n",
    "\n",
    "\n",
    "xsq_even = Implies(x*x == 2*y*y, evenp(x*x, y*y))\n",
    "prove(Implies(evenp_def, xsq_even))\n",
    "\n",
    "\n",
    "ForAll([x,y], Implies( evenp(x,y), even(x)))\n",
    "#ForAll([x,y], Implies( evenp(x,y), even(y)))\n",
    "\n",
    "\n",
    "\n",
    "\n",
    "\n"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 3,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "proved\n"
     ]
    }
   ],
   "source": [
    "rearrange = Implies(x*x == 2*y*y, y*y == (x*x)/2)\n",
    "prove(rearrange)"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": null,
   "metadata": {},
   "outputs": [],
   "source": [
    "rearrange = Implies(x*x == 2*y*y, y*y == (x*x)/2)\n",
    "prove(rearrange)"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": null,
   "metadata": {},
   "outputs": [],
   "source": [
    "xsq_even = Implies(x*x == 2*y, even(x*x))\n",
    "prove(Implies(even_def, xsq_even))\n",
    "xsq_even"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": null,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "proved\n",
      "proved\n",
      "proved\n"
     ]
    }
   ],
   "source": [
    "#from cvc5.pythonic import *\n",
    "#from z3 \n",
    "x,y,z = Ints(\"x y z\")\n",
    "even = Function(\"even\", IntSort(), BoolSort())\n",
    "even_def = ForAll(x, even(x) == Exists([y], x == 2*y))\n",
    "odd = Function(\"odd\", IntSort(), BoolSort())\n",
    "odd_def = ForAll(x, odd(x) == Exists([y], x == 2*y+1))\n",
    "\n",
    "even_not_odd = ForAll([x], Implies(even(x), Not(odd(x))))\n",
    "prove(Implies(And(even_def, odd_def), even_not_odd))\n",
    "\n",
    "even_or_odd = ForAll(x, even(x) != odd(x))\n",
    "prove(Implies(And(even_def, odd_def), even_or_odd))\n",
    "\n"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "Facts about even odd"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": null,
   "metadata": {},
   "outputs": [],
   "source": [
    "even_not_odd = ForAll(x, Implies(even(x), Not(odd(x))))\n",
    "prove(Implies(And(even_def, odd_def), even_not_odd))\n",
    "\n",
    "odd_not_even = ForAll(x, Implies(odd(x), Not(even(x))))\n",
    "prove(Implies(And(even_def, odd_def), odd_not_even))\n",
    "\n",
    "\n",
    "not_even_odd = ForAll(x, Not(And(even(x),odd(x))))\n",
    "prove(Implies(And(even_def, odd_def), not_even_odd))\n",
    "\n",
    "\n",
    "\n",
    "\n",
    "def induct(P):\n",
    "    return Implies(\n",
    "            And(P(0), \n",
    "                ForAll([x], Implies(P(x),  And(P(x + 1), P(x - 1))))),\n",
    "            ForAll([x], P(x)))\n",
    "\n",
    "even_or_odd = ForAll(x, Or(even(x), odd(x)))\n",
    "#prove(Implies(And(even_def, odd_def, induct(lambda n: Or(even(n), odd(n)))), even_or_odd))\n",
    "\n",
    "evenodd = lambda n: Or(even(n), odd(n))\n",
    "prove(Implies(And(even_def, odd_def), Or(even(0), odd(0))))\n",
    "prove(Implies(And(even_def, odd_def), ForAll([x], Implies(evenodd(x), evenodd(x+1)))))\n",
    "#prove(Implies(And(even_def, odd_def), ForAll([x], Implies(evenodd(x), evenodd(x-1)))))\n",
    "\n",
    "prove(Implies(And(even_def, odd_def), ForAll([x], Implies(odd(x), Exists([y], x == 2*y-1)))))\n",
    "\n",
    "#prove(Implies(And(even_def, odd_def), ForAll([x], odd(x) == Exists([y], x == 2*y-1))))\n",
    "\n",
    "even_or_odd = ForAll(x, even(x) != odd(x))\n",
    "# Z3 can't prove this for some reason. I suppose this might be a complex diophantine equation or something.\n",
    "prove(Implies(And(even_def, odd_def), even_or_odd))\n",
    "\n",
    "\n",
    "sqrt_thm = ForAll([x,y], Implies(x != 0, 2*x*x != y*y))\n",
    "#prove(sqrt_thm)\n",
    "\n",
    "#solve(2*x*x == y*y)\n",
    "\n",
    "even_sq = ForAll([x,y], Implies(even(x), even(x*x)))\n",
    "#prove(Implies(even_def, even_sq))\n",
    "\n",
    "\n"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 3,
   "metadata": {},
   "outputs": [
    {
     "ename": "NameError",
     "evalue": "name 'Function' is not defined",
     "output_type": "error",
     "traceback": [
      "\u001b[0;31m---------------------------------------------------------------------------\u001b[0m",
      "\u001b[0;31mNameError\u001b[0m                                 Traceback (most recent call last)",
      "Cell \u001b[0;32mIn[3], line 4\u001b[0m\n\u001b[1;32m      1\u001b[0m \u001b[38;5;66;03m#def divides(x,y):\u001b[39;00m\n\u001b[1;32m      2\u001b[0m \u001b[38;5;66;03m#    return x == y*(x/y)\u001b[39;00m\n\u001b[0;32m----> 4\u001b[0m divides \u001b[38;5;241m=\u001b[39m \u001b[43mFunction\u001b[49m(\u001b[38;5;124m\"\u001b[39m\u001b[38;5;124mdivides\u001b[39m\u001b[38;5;124m\"\u001b[39m, I, I, B)\n\u001b[1;32m      5\u001b[0m divides_def \u001b[38;5;241m=\u001b[39m ForAll([x,y], divides(x,y) \u001b[38;5;241m==\u001b[39m (x \u001b[38;5;241m==\u001b[39m y\u001b[38;5;241m*\u001b[39m(x\u001b[38;5;241m/\u001b[39my)))\n\u001b[1;32m      6\u001b[0m prove(ForAll([x], divides(x, \u001b[38;5;241m0\u001b[39m)))\n",
      "\u001b[0;31mNameError\u001b[0m: name 'Function' is not defined"
     ]
    }
   ],
   "source": [
    "#def divides(x,y):\n",
    "#    return x == y*(x/y)\n",
    "\n",
    "divides = Function(\"divides\", I, I, B)\n",
    "divides_def = ForAll([x,y], divides(x,y) == (x == y*(x/y)))\n",
    "prove(ForAll([x], divides(x, 0)))\n",
    "\n"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 89,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "Writing /tmp/prime2.smt2\n"
     ]
    }
   ],
   "source": [
    "%%file /tmp/prime2.smt2\n",
    "(declare-sort I 0)\n",
    "(declare-fun mul (I I) I)\n",
    "(declare-fun divsk (I I) I)\n",
    "(declare-fun divides (I I) Bool)\n",
    "(declare-const one I)\n",
    "(declare-const two I)\n",
    "(declare-const a I)\n",
    "(declare-const b I)\n",
    "(assert (distinct one two))\n",
    "(assert (forall ((x I)) (= (mul one x) x)))\n",
    "(assert (forall ((x I)) (= (mul x one) x)))\n",
    "(assert (forall ((x I) (y I) (z I)) (= (mul x (mul y z)) (mul (mul x y) z))))\n",
    "(assert (forall ((x I) (y I)) (= (mul x y) (mul y x))))\n",
    "(assert (forall ((x I) (y I) (z I))\n",
    "  (=> (= (mul x y) (mul x z)) (= y z))))\n",
    "(assert (forall ((x I) (y I)) (=> (divides x y) (= (mul y (divsk x y)) x))))\n",
    "(assert (forall ((x I) (y I))\n",
    "  (=> (divides (mul x y) two) (or (divides x two) (divides y two)))))\n",
    "(assert (= (mul a a) (mul two (mul b b)))) ; \n",
    "(assert (forall ((x I)) (=> (and (divides a x) (divides b x)) (= x one)))) ; lowest common terms\n",
    ";(assert (let ((a!1 (exists ((x I) (y I))\n",
    ";             (let ((a!1 (forall ((z I))\n",
    ";                         (=> (and (divides x z) (divides y z)) (= z one)))))\n",
    ";               (and a!1 (= (mul x x) (mul two (mul y y))))))))\n",
    ";  (not (not a!1))))"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 10,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "Overwriting /tmp/sqrt2.in\n"
     ]
    }
   ],
   "source": []
  },
  {
   "cell_type": "code",
   "execution_count": 66,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "============================== Prover9 ===============================\n",
      "Prover9 (64) version 2017-11A (CIIRC), November 2017.\n",
      "Process 259687 was started by philip on philip-ThinkPad-Z13-Gen-2,\n",
      "Mon Apr 29 17:27:10 2024\n",
      "The command was \"prover9 -f /tmp/sqrt2.in\".\n",
      "============================== end of head ===========================\n",
      "\n",
      "============================== INPUT =================================\n",
      "\n",
      "% Reading from file /tmp/sqrt2.in\n",
      "\n",
      "set(auto).\n",
      "    % set(auto) -> set(auto_inference).\n",
      "    % set(auto) -> set(auto_setup).\n",
      "    % set(auto_setup) -> set(predicate_elim).\n",
      "    % set(auto_setup) -> assign(eq_defs, unfold).\n",
      "    % set(auto) -> set(auto_limits).\n",
      "    % set(auto_limits) -> assign(max_weight, \"100.000\").\n",
      "    % set(auto_limits) -> assign(sos_limit, 20000).\n",
      "    % set(auto) -> set(auto_denials).\n",
      "    % set(auto) -> set(auto_process).\n",
      "\n",
      "formulas(sos).\n",
      "x = x.\n",
      "m(1,x) = x.\n",
      "m(x,1) = x.\n",
      "m(x,m(y,z)) = m(m(x,y),z).\n",
      "m(x,y) = m(y,x).\n",
      "m(x,y) != m(x,z) | y = z.\n",
      "-d(x,y) | m(x,f(x,y)) = y.\n",
      "m(x,z) != y | d(x,y).\n",
      "-d(2,m(x,y)) | d(2,x) | d(2,y).\n",
      "m(a,a) = m(2,m(b,b)).\n",
      "-d(x,a) | -d(x,b) | x = 1.\n",
      "2 != 1.\n",
      "end_of_list.\n",
      "\n",
      "============================== end of input ==========================\n",
      "\n",
      "============================== PROCESS NON-CLAUSAL FORMULAS ==========\n",
      "\n",
      "% Formulas that are not ordinary clauses:\n",
      "\n",
      "============================== end of process non-clausal formulas ===\n",
      "\n",
      "============================== PROCESS INITIAL CLAUSES ===============\n",
      "\n",
      "% Clauses before input processing:\n",
      "\n",
      "formulas(usable).\n",
      "end_of_list.\n",
      "\n",
      "formulas(sos).\n",
      "x = x.  [assumption].\n",
      "m(1,x) = x.  [assumption].\n",
      "m(x,1) = x.  [assumption].\n",
      "m(x,m(y,z)) = m(m(x,y),z).  [assumption].\n",
      "m(x,y) = m(y,x).  [assumption].\n",
      "m(x,y) != m(x,z) | y = z.  [assumption].\n",
      "-d(x,y) | m(x,f(x,y)) = y.  [assumption].\n",
      "m(x,y) != z | d(x,z).  [assumption].\n",
      "-d(2,m(x,y)) | d(2,x) | d(2,y).  [assumption].\n",
      "m(a,a) = m(2,m(b,b)).  [assumption].\n",
      "-d(x,a) | -d(x,b) | x = 1.  [assumption].\n",
      "2 != 1.  [assumption].\n",
      "end_of_list.\n",
      "\n",
      "formulas(demodulators).\n",
      "end_of_list.\n",
      "\n",
      "============================== PREDICATE ELIMINATION =================\n",
      "\n",
      "No predicates eliminated.\n",
      "\n",
      "============================== end predicate elimination =============\n",
      "\n",
      "Auto_denials:  (non-Horn, no changes).\n",
      "\n",
      "Term ordering decisions:\n",
      "Predicate symbol precedence:  predicate_order([ =, d ]).\n",
      "Function symbol precedence:  function_order([ 2, 1, a, b, m, f ]).\n",
      "After inverse_order:  (no changes).\n",
      "Unfolding symbols: (none).\n",
      "\n",
      "Auto_inference settings:\n",
      "  % set(paramodulation).  % (positive equality literals)\n",
      "  % set(binary_resolution).  % (non-Horn)\n",
      "  % set(neg_ur_resolution).  % (non-Horn, less than 100 clauses)\n",
      "\n",
      "Auto_process settings:\n",
      "  % set(factor).  % (non-Horn)\n",
      "  % set(unit_deletion).  % (non-Horn)\n",
      "\n",
      "kept:      1 m(1,x) = x.  [assumption].\n",
      "kept:      2 m(x,1) = x.  [assumption].\n",
      "           3 m(x,m(y,z)) = m(m(x,y),z).  [assumption].\n",
      "kept:      4 m(m(x,y),z) = m(x,m(y,z)).  [copy(3),flip(a)].\n",
      "\n",
      "% Operation m is commutative; C redundancy checks enabled.\n",
      "kept:      5 m(x,y) = m(y,x).  [assumption].\n",
      "kept:      6 m(x,y) != m(x,z) | y = z.  [assumption].\n",
      "kept:      7 -d(x,y) | m(x,f(x,y)) = y.  [assumption].\n",
      "kept:      8 m(x,y) != z | d(x,z).  [assumption].\n",
      "kept:      9 -d(2,m(x,y)) | d(2,x) | d(2,y).  [assumption].\n",
      "           10 m(a,a) = m(2,m(b,b)).  [assumption].\n",
      "kept:      11 m(2,m(b,b)) = m(a,a).  [copy(10),flip(a)].\n",
      "           12 -d(x,a) | -d(x,b) | x = 1.  [assumption].\n",
      "kept:      13 -d(x,a) | -d(x,b) | 1 = x.  [copy(12),flip(c)].\n",
      "           14 2 != 1.  [assumption].\n",
      "kept:      15 1 != 2.  [copy(14),flip(a)].\n",
      "kept:      16 -d(2,m(x,x)) | d(2,x).  [factor(9,b,c)].\n",
      "\n",
      "============================== end of process initial clauses ========\n",
      "\n",
      "============================== CLAUSES FOR SEARCH ====================\n",
      "\n",
      "% Clauses after input processing:\n",
      "\n",
      "formulas(usable).\n",
      "end_of_list.\n",
      "\n",
      "formulas(sos).\n",
      "1 m(1,x) = x.  [assumption].\n",
      "2 m(x,1) = x.  [assumption].\n",
      "4 m(m(x,y),z) = m(x,m(y,z)).  [copy(3),flip(a)].\n",
      "5 m(x,y) = m(y,x).  [assumption].\n",
      "6 m(x,y) != m(x,z) | y = z.  [assumption].\n",
      "7 -d(x,y) | m(x,f(x,y)) = y.  [assumption].\n",
      "8 m(x,y) != z | d(x,z).  [assumption].\n",
      "9 -d(2,m(x,y)) | d(2,x) | d(2,y).  [assumption].\n",
      "11 m(2,m(b,b)) = m(a,a).  [copy(10),flip(a)].\n",
      "13 -d(x,a) | -d(x,b) | 1 = x.  [copy(12),flip(c)].\n",
      "15 1 != 2.  [copy(14),flip(a)].\n",
      "16 -d(2,m(x,x)) | d(2,x).  [factor(9,b,c)].\n",
      "end_of_list.\n",
      "\n",
      "formulas(demodulators).\n",
      "1 m(1,x) = x.  [assumption].\n",
      "2 m(x,1) = x.  [assumption].\n",
      "4 m(m(x,y),z) = m(x,m(y,z)).  [copy(3),flip(a)].\n",
      "5 m(x,y) = m(y,x).  [assumption].\n",
      "        % (lex-dep)\n",
      "11 m(2,m(b,b)) = m(a,a).  [copy(10),flip(a)].\n",
      "end_of_list.\n",
      "\n",
      "============================== end of clauses for search =============\n",
      "\n",
      "============================== SEARCH ================================\n",
      "\n",
      "% Starting search at 0.00 seconds.\n",
      "\n",
      "given #1 (I,wt=5): 1 m(1,x) = x.  [assumption].\n",
      "\n",
      "given #2 (I,wt=5): 2 m(x,1) = x.  [assumption].\n",
      "\n",
      "given #3 (I,wt=11): 4 m(m(x,y),z) = m(x,m(y,z)).  [copy(3),flip(a)].\n",
      "\n",
      "given #4 (I,wt=7): 5 m(x,y) = m(y,x).  [assumption].\n",
      "\n",
      "% Operation m is associative-commutative; CAC redundancy checks enabled.\n",
      "\n",
      "given #5 (I,wt=10): 6 m(x,y) != m(x,z) | y = z.  [assumption].\n",
      "\n",
      "given #6 (I,wt=10): 7 -d(x,y) | m(x,f(x,y)) = y.  [assumption].\n",
      "\n",
      "given #7 (I,wt=8): 8 m(x,y) != z | d(x,z).  [assumption].\n",
      "\n",
      "given #8 (I,wt=11): 9 -d(2,m(x,y)) | d(2,x) | d(2,y).  [assumption].\n",
      "\n",
      "given #9 (I,wt=9): 11 m(2,m(b,b)) = m(a,a).  [copy(10),flip(a)].\n",
      "\n",
      "given #10 (I,wt=9): 13 -d(x,a) | -d(x,b) | 1 = x.  [copy(12),flip(c)].\n",
      "\n",
      "given #11 (I,wt=3): 15 1 != 2.  [copy(14),flip(a)].\n",
      "\n",
      "given #12 (I,wt=8): 16 -d(2,m(x,x)) | d(2,x).  [factor(9,b,c)].\n",
      "\n",
      "given #13 (A,wt=11): 17 m(x,m(y,z)) = m(y,m(x,z)).  [para(5(a,1),4(a,1,1)),rewrite([4(2)])].\n",
      "\n",
      "given #14 (F,wt=5): 38 m(x,2) != x.  [ur(6,b,15,a),rewrite([2(2)]),flip(a)].\n",
      "\n",
      "given #15 (F,wt=5): 45 m(2,x) != x.  [para(5(a,1),38(a,1))].\n",
      "\n",
      "given #16 (F,wt=7): 47 m(b,b) != m(a,a).  [para(11(a,1),45(a,1)),flip(a)].\n",
      "\n",
      "given #17 (F,wt=9): 44 m(x,m(y,2)) != m(x,y).  [ur(6,b,38,a)].\n",
      "\n",
      "given #18 (T,wt=3): 24 d(x,x).  [resolve(8,a,2,a)].\n",
      "\n",
      "given #19 (T,wt=3): 25 d(1,x).  [resolve(8,a,1,a)].\n",
      "\n",
      "given #20 (T,wt=5): 22 d(x,m(y,x)).  [resolve(8,a,5,a)].\n",
      "\n",
      "given #21 (T,wt=5): 27 d(x,m(x,y)).  [resolve(8,a,2,a(flip)),rewrite([5(3),1(3)])].\n",
      "\n",
      "given #22 (A,wt=8): 18 m(x,y) != x | 1 = y.  [para(2(a,1),6(a,1)),flip(a)].\n",
      "\n",
      "given #23 (F,wt=9): 46 m(x,m(2,y)) != m(x,y).  [ur(6,b,45,a)].\n",
      "\n",
      "given #24 (F,wt=9): 50 m(x,m(2,y)) != m(y,x).  [para(5(a,1),44(a,1)),rewrite([4(3)])].\n",
      "\n",
      "given #25 (F,wt=9): 51 m(x,m(y,2)) != m(y,x).  [para(5(a,1),44(a,2))].\n",
      "\n",
      "given #26 (F,wt=9): 66 m(2,m(x,y)) != m(y,x).  [para(5(a,1),46(a,1)),rewrite([4(3)])].\n",
      "\n",
      "given #27 (T,wt=5): 33 d(2,m(a,a)).  [resolve(11,a,8,a)].\n",
      "\n",
      "given #28 (T,wt=3): 79 d(2,a).  [resolve(33,a,16,a)].\n",
      "\n",
      "given #29 (T,wt=5): 55 f(1,x) = x.  [resolve(25,a,7,a),rewrite([1(4)])].\n",
      "\n",
      "given #30 (T,wt=6): 28 x != y | d(x,y).  [para(2(a,1),8(a,1))].\n",
      "\n",
      "given #31 (A,wt=14): 19 m(x,m(y,z)) != m(x,m(y,u)) | z = u.  [para(4(a,1),6(a,1)),rewrite([4(4)])].\n",
      "\n",
      "given #32 (F,wt=3): 82 -d(2,b).  [ur(13,a,79,a,c,15,a)].\n",
      "\n",
      "given #33 (F,wt=3): 100 b != 2.  [ur(28,b,82,a),flip(a)].\n",
      "\n",
      "given #34 (F,wt=5): 101 -d(2,m(b,b)).  [ur(16,b,82,a)].\n",
      "\n",
      "given #35 (F,wt=5): 102 m(2,x) != b.  [ur(8,b,82,a)].\n",
      "\n",
      "given #36 (T,wt=6): 53 -d(b,a) | b = 1.  [resolve(24,a,13,b),flip(b)].\n",
      "\n",
      "given #37 (T,wt=7): 39 d(x,m(y,m(x,z))).  [resolve(17,a,8,a)].\n",
      "\n",
      "given #38 (T,wt=5): 118 d(b,m(a,a)).  [para(11(a,1),39(a,2))].\n",
      "\n",
      "given #39 (T,wt=7): 54 m(x,f(x,x)) = x.  [resolve(24,a,7,a)].\n",
      "\n",
      "given #40 (A,wt=10): 20 m(x,y) != m(z,x) | z = y.  [para(5(a,1),6(a,1)),flip(a)].\n",
      "\n",
      "given #41 (F,wt=5): 105 m(b,b) != 2.  [ur(28,b,101,a),flip(a)].\n",
      "\n",
      "given #42 (F,wt=5): 111 m(x,2) != b.  [para(5(a,1),102(a,1))].\n",
      "\n",
      "given #43 (F,wt=5): 112 m(a,a) != b.  [para(11(a,1),102(a,1))].\n",
      "\n",
      "given #44 (F,wt=7): 104 m(x,b) != m(x,2).  [ur(6,b,100,a)].\n",
      "\n",
      "given #45 (T,wt=5): 120 f(x,x) = 1.  [resolve(54,a,18,a),flip(a)].\n",
      "\n",
      "given #46 (T,wt=6): 131 x != y | y = x.  [para(1(a,1),20(a,1)),rewrite([2(2)])].\n",
      "\n",
      "given #47 (T,wt=7): 57 d(x,m(y,m(z,x))).  [para(4(a,1),22(a,2))].\n",
      "\n",
      "given #48 (T,wt=7): 58 d(m(b,b),m(a,a)).  [para(11(a,1),22(a,2))].\n",
      "\n",
      "given #49 (A,wt=10): 21 m(x,y) != m(z,x) | y = z.  [para(5(a,1),6(a,2))].\n",
      "\n",
      "given #50 (F,wt=7): 107 -d(2,m(b,m(b,b))).  [ur(9,b,82,a,c,101,a)].\n",
      "\n",
      "given #51 (F,wt=7): 108 m(b,b) != m(2,x).  [ur(8,b,101,a),flip(a)].\n",
      "\n",
      "given #52 (F,wt=7): 113 m(x,m(2,y)) != b.  [para(17(a,1),102(a,1))].\n",
      "\n",
      "given #53 (F,wt=7): 122 m(b,x) != m(x,2).  [ur(20,b,100,a),flip(a)].\n",
      "\n",
      "given #54 (T,wt=7): 81 m(2,f(2,a)) = a.  [resolve(79,a,7,a)].\n",
      "\n",
      "given #55 (T,wt=5): 191 d(f(2,a),a).  [para(81(a,1),22(a,2))].\n",
      "\n",
      "given #56 (T,wt=5): 197 d(2,m(x,a)).  [para(81(a,1),39(a,2,2))].\n",
      "\n",
      "given #57 (T,wt=5): 206 d(2,m(a,x)).  [para(5(a,1),197(a,2))].\n",
      "\n",
      "given #58 (A,wt=9): 23 d(m(x,y),m(x,m(y,z))).  [resolve(8,a,4,a)].\n",
      "\n",
      "given #59 (F,wt=3): 196 b != a.  [para(81(a,1),102(a,1)),flip(a)].\n",
      "\n",
      "given #60 (F,wt=5): 190 f(2,a) != a.  [para(81(a,1),45(a,1)),flip(a)].\n",
      "\n",
      "given #61 (F,wt=5): 201 m(b,b) != a.  [para(81(a,1),108(a,2))].\n",
      "\n",
      "given #62 (F,wt=5): 202 m(x,a) != b.  [para(81(a,1),113(a,1,2))].\n",
      "\n",
      "given #63 (T,wt=6): 188 a != x | d(2,x).  [para(81(a,1),8(a,1))].\n",
      "\n",
      "given #64 (T,wt=7): 84 d(m(x,y),m(y,x)).  [resolve(28,a,5,a)].\n",
      "\n",
      "given #65 (T,wt=7): 200 d(f(2,a),m(x,a)).  [para(81(a,1),57(a,2,2))].\n",
      "\n",
      "given #66 (T,wt=7): 205 d(2,m(x,m(y,a))).  [para(4(a,1),197(a,2))].\n",
      "\n",
      "given #67 (A,wt=12): 29 m(x,m(y,z)) != u | d(m(x,y),u).  [para(4(a,1),8(a,1))].\n",
      "\n",
      "given #68 (F,wt=5): 233 m(a,x) != b.  [para(5(a,1),202(a,1))].\n",
      "\n",
      "given #69 (F,wt=3): 255 -d(a,b).  [ur(7,b,233,a)].\n",
      "\n",
      "given #70 (F,wt=7): 128 m(2,x) != m(x,b).  [ur(20,b,100,a(flip)),flip(a)].\n",
      "\n",
      "given #71 (F,wt=7): 147 m(x,m(y,2)) != b.  [para(4(a,1),111(a,1))].\n",
      "\n",
      "given #72 (T,wt=7): 208 d(2,m(x,m(a,y))).  [para(17(a,1),206(a,2))].\n",
      "\n",
      "given #73 (T,wt=7): 214 d(m(2,b),m(a,a)).  [para(11(a,1),23(a,2))].\n",
      "\n",
      "given #74 (T,wt=7): 217 d(m(x,2),m(x,a)).  [para(81(a,1),23(a,2,2))].\n",
      "\n",
      "given #75 (T,wt=7): 241 d(f(2,a),m(a,x)).  [para(5(a,1),200(a,2))].\n",
      "\n",
      "given #76 (A,wt=8): 30 m(x,y) != z | d(y,z).  [para(5(a,1),8(a,1))].\n",
      "\n",
      "given #77 (F,wt=7): 164 m(b,m(b,b)) != 2.  [ur(28,b,107,a),flip(a)].\n",
      "\n",
      "given #78 (F,wt=7): 172 m(b,b) != m(x,2).  [para(5(a,1),108(a,2))].\n",
      "\n",
      "given #79 (F,wt=7): 184 m(b,x) != m(2,x).  [para(5(a,1),122(a,2))].\n",
      "\n",
      "given #80 (F,wt=7): 218 m(a,x) != m(x,b).  [ur(21,b,196,a),flip(a)].\n",
      "\n",
      "given #81 (T,wt=7): 275 d(m(2,x),m(x,a)).  [para(5(a,1),217(a,1))].\n",
      "\n",
      "given #82 (T,wt=7): 276 d(m(x,2),m(a,x)).  [para(5(a,1),217(a,2))].\n",
      "\n",
      "given #83 (T,wt=7): 304 d(m(2,x),m(a,x)).  [para(5(a,1),275(a,2))].\n",
      "\n",
      "given #84 (T,wt=8): 37 m(a,a) != x | d(2,x).  [para(11(a,1),8(a,1))].\n",
      "\n",
      "given #85 (A,wt=15): 31 -d(2,m(x,m(y,z))) | d(2,m(x,y)) | d(2,z).  [para(4(a,1),9(a,2))].\n",
      "\n",
      "given #86 (F,wt=7): 219 m(b,x) != m(x,a).  [ur(20,b,196,a),flip(a)].\n",
      "\n",
      "given #87 (F,wt=7): 221 m(x,b) != m(x,a).  [ur(6,b,196,a)].\n",
      "\n",
      "given #88 (F,wt=7): 232 m(x,m(y,a)) != b.  [para(4(a,1),202(a,1))].\n",
      "\n",
      "given #89 (F,wt=7): 234 m(b,m(b,b)) != a.  [ur(188,b,107,a),flip(a)].\n",
      "\n",
      "given #90 (T,wt=8): 62 m(x,y) != y | 1 = x.  [para(5(a,1),18(a,1))].\n",
      "\n",
      "given #91 (T,wt=8): 192 a != 2 | f(2,a) = 1.  [para(81(a,1),18(a,1)),flip(b)].\n",
      "\n",
      "given #92 (T,wt=8): 282 a != x | d(f(2,a),x).  [para(81(a,1),30(a,1))].\n",
      "\n",
      "given #93 (T,wt=9): 59 d(m(x,y),m(x,m(z,y))).  [para(17(a,1),22(a,2))].\n",
      "\n",
      "given #94 (A,wt=14): 32 -d(2,m(x,m(y,m(x,y)))) | d(2,m(x,y)).  [factor(31,b,c)].\n",
      "\n",
      "given #95 (F,wt=7): 257 m(x,m(a,y)) != b.  [para(17(a,1),233(a,1))].\n",
      "\n",
      "given #96 (F,wt=7): 263 m(b,f(2,a)) != a.  [para(81(a,1),128(a,1)),rewrite([5(6)]),flip(a)].\n",
      "\n",
      "given #97 (F,wt=7): 301 m(b,x) != m(a,x).  [para(5(a,1),218(a,2)),flip(a)].\n",
      "\n",
      "given #98 (F,wt=9): 106 -d(2,m(b,m(b,m(b,b)))).  [ur(16,b,101,a),rewrite([17(8),5(7)])].\n",
      "\n",
      "given #99 (T,wt=9): 116 d(x,m(y,m(z,m(x,u)))).  [para(4(a,1),39(a,2))].\n",
      "\n",
      "given #100 (T,wt=7): 368 d(b,m(x,m(a,a))).  [para(11(a,1),116(a,2,2))].\n",
      "\n",
      "given #101 (T,wt=7): 371 d(b,m(a,m(a,x))).  [para(5(a,1),368(a,2)),rewrite([4(5)])].\n",
      "\n",
      "given #102 (T,wt=7): 372 d(b,m(a,m(x,a))).  [para(17(a,1),368(a,2))].\n",
      "\n",
      "given #103 (A,wt=13): 34 m(2,m(b,m(b,x))) = m(a,m(a,x)).  [para(11(a,1),4(a,1,1)),rewrite([4(4),4(9)]),flip(a)].\n",
      "\n",
      "given #104 (F,wt=9): 110 m(x,m(2,y)) != m(x,b).  [ur(6,b,102,a)].\n",
      "\n",
      "given #105 (F,wt=9): 121 m(2,m(x,y)) != m(y,b).  [ur(20,b,102,a),rewrite([4(5)]),flip(a)].\n",
      "\n",
      "given #106 (F,wt=9): 127 m(b,x) != m(x,m(2,y)).  [ur(20,b,102,a(flip)),flip(a)].\n",
      "\n",
      "given #107 (F,wt=9): 139 m(b,m(b,x)) != m(x,2).  [ur(20,b,105,a),rewrite([4(6)]),flip(a)].\n",
      "\n",
      "given #108 (T,wt=9): 158 d(x,m(y,m(z,m(u,x)))).  [para(4(a,1),57(a,2,2))].\n",
      "\n",
      "given #109 (T,wt=9): 159 d(m(b,b),m(x,m(a,a))).  [para(11(a,1),57(a,2,2))].\n",
      "\n",
      "given #110 (T,wt=9): 211 d(m(x,y),m(y,m(x,z))).  [para(5(a,1),23(a,1))].\n",
      "\n",
      "given #111 (T,wt=9): 212 d(m(x,y),m(y,m(z,x))).  [para(5(a,1),23(a,2)),rewrite([4(3)])].\n",
      "\n",
      "given #112 (A,wt=13): 35 m(b,m(b,m(x,2))) = m(x,m(a,a)).  [para(11(a,1),4(a,2,2)),rewrite([17(6),5(5)])].\n",
      "\n",
      "given #113 (F,wt=9): 141 m(x,m(b,b)) != m(x,2).  [ur(6,b,105,a)].\n",
      "\n",
      "given #114 (F,wt=7): 513 m(a,a) != m(2,2).  [para(11(a,1),141(a,1))].\n",
      "\n",
      "given #115 (F,wt=9): 142 m(2,x) != m(x,m(b,b)).  [ur(20,b,105,a(flip)),flip(a)].\n",
      "\n",
      "given #116 (F,wt=9): 143 m(x,m(2,y)) != m(y,b).  [ur(20,b,111,a),rewrite([4(5)]),flip(a)].\n",
      "\n",
      "given #117 (T,wt=9): 213 d(m(x,2),m(x,m(a,a))).  [para(11(a,1),23(a,2,2))].\n",
      "\n",
      "given #118 (T,wt=9): 240 d(f(2,a),m(x,m(y,a))).  [para(4(a,1),200(a,2))].\n",
      "\n",
      "given #119 (T,wt=9): 243 d(2,m(x,m(y,m(z,a)))).  [para(4(a,1),205(a,2,2))].\n",
      "\n",
      "given #120 (T,wt=9): 270 d(2,m(x,m(y,m(a,z)))).  [para(4(a,1),208(a,2))].\n",
      "\n",
      "given #121 (A,wt=12): 36 m(a,a) != m(2,x) | m(b,b) = x.  [para(11(a,1),6(a,1))].\n",
      "\n",
      "given #122 (F,wt=7): 548 m(a,a) != m(2,a).  [ur(36,b,201,a)].\n",
      "\n",
      "given #123 (F,wt=9): 145 m(x,m(y,2)) != m(x,b).  [ur(6,b,111,a)].\n",
      "\n",
      "given #124 (F,wt=9): 146 m(b,x) != m(x,m(y,2)).  [ur(20,b,111,a(flip)),flip(a)].\n",
      "\n",
      "given #125 (F,wt=9): 167 m(b,m(b,b)) != m(2,x).  [ur(8,b,107,a),flip(a)].\n",
      "\n",
      "given #126 (T,wt=9): 278 d(f(2,a),m(x,m(a,y))).  [para(17(a,1),241(a,2))].\n",
      "\n",
      "given #127 (T,wt=9): 305 d(m(a,a),m(a,m(b,b))).  [para(11(a,1),275(a,1)),rewrite([5(8)])].\n",
      "\n",
      "given #128 (T,wt=9): 339 d(m(x,y),m(z,m(y,x))).  [para(5(a,1),59(a,2)),rewrite([4(3)])].\n",
      "\n",
      "given #129 (T,wt=9): 342 d(m(x,f(2,a)),m(x,a)).  [para(81(a,1),59(a,2,2))].\n",
      "\n",
      "given #130 (A,wt=14): 40 m(x,m(y,z)) != m(y,u) | m(x,z) = u.  [para(17(a,1),6(a,1))].\n",
      "\n",
      "given #131 (F,wt=9): 173 m(b,b) != m(x,m(2,y)).  [para(17(a,1),108(a,2))].\n",
      "\n",
      "given #132 (F,wt=7): 639 m(b,b) != m(x,a).  [para(81(a,1),173(a,2,2))].\n",
      "\u0007-------- Proof 1 -------- \n",
      "\n",
      "============================== PROOF =================================\n",
      "\n",
      "% Proof 1 at 0.02 (+ 0.06) seconds.\n",
      "% Length of proof is 27.\n",
      "% Level of proof is 10.\n",
      "% Maximum clause weight is 12.000.\n",
      "% Given clauses 132 (13.200 givens/level).\n",
      "\n",
      "3 m(x,m(y,z)) = m(m(x,y),z).  [assumption].\n",
      "4 m(m(x,y),z) = m(x,m(y,z)).  [copy(3),flip(a)].\n",
      "5 m(x,y) = m(y,x).  [assumption].\n",
      "6 m(x,y) != m(x,z) | y = z.  [assumption].\n",
      "7 -d(x,y) | m(x,f(x,y)) = y.  [assumption].\n",
      "8 m(x,y) != z | d(x,z).  [assumption].\n",
      "9 -d(2,m(x,y)) | d(2,x) | d(2,y).  [assumption].\n",
      "10 m(a,a) = m(2,m(b,b)).  [assumption].\n",
      "11 m(2,m(b,b)) = m(a,a).  [copy(10),flip(a)].\n",
      "12 -d(x,a) | -d(x,b) | x = 1.  [assumption].\n",
      "13 -d(x,a) | -d(x,b) | 1 = x.  [copy(12),flip(c)].\n",
      "14 2 != 1.  [assumption].\n",
      "15 1 != 2.  [copy(14),flip(a)].\n",
      "16 -d(2,m(x,x)) | d(2,x).  [factor(9,b,c)].\n",
      "17 m(x,m(y,z)) = m(y,m(x,z)).  [para(5(a,1),4(a,1,1)),rewrite([4(2)])].\n",
      "33 d(2,m(a,a)).  [resolve(11,a,8,a)].\n",
      "36 m(a,a) != m(2,x) | m(b,b) = x.  [para(11(a,1),6(a,1))].\n",
      "79 d(2,a).  [resolve(33,a,16,a)].\n",
      "81 m(2,f(2,a)) = a.  [resolve(79,a,7,a)].\n",
      "82 -d(2,b).  [ur(13,a,79,a,c,15,a)].\n",
      "101 -d(2,m(b,b)).  [ur(16,b,82,a)].\n",
      "108 m(b,b) != m(2,x).  [ur(8,b,101,a),flip(a)].\n",
      "173 m(b,b) != m(x,m(2,y)).  [para(17(a,1),108(a,2))].\n",
      "186 m(2,m(f(2,a),x)) = m(a,x).  [para(81(a,1),4(a,1,1)),flip(a)].\n",
      "639 m(b,b) != m(x,a).  [para(81(a,1),173(a,2,2))].\n",
      "642 m(a,a) != m(2,m(x,a)).  [ur(36,b,639,a)].\n",
      "643 $F.  [resolve(642,a,186,a(flip))].\n",
      "\n",
      "============================== end of proof ==========================\n",
      "\n",
      "============================== STATISTICS ============================\n",
      "\n",
      "Given=132. Generated=2152. Kept=638. proofs=1.\n",
      "Usable=127. Sos=485. Demods=56. Limbo=1, Disabled=36. Hints=0.\n",
      "Kept_by_rule=0, Deleted_by_rule=0.\n",
      "Forward_subsumed=1514. Back_subsumed=15.\n",
      "Sos_limit_deleted=0. Sos_displaced=0. Sos_removed=0.\n",
      "New_demodulators=58 (2 lex), Back_demodulated=9. Back_unit_deleted=0.\n",
      "Demod_attempts=18123. Demod_rewrites=1480.\n",
      "Res_instance_prunes=0. Para_instance_prunes=0. Basic_paramod_prunes=0.\n",
      "Nonunit_fsub_feature_tests=1506. Nonunit_bsub_feature_tests=528.\n",
      "Megabytes=0.81.\n",
      "User_CPU=0.02, System_CPU=0.06, Wall_clock=0.\n",
      "\n",
      "============================== end of statistics =====================\n",
      "\n",
      "============================== end of search =========================\n",
      "\n",
      "THEOREM PROVED\n",
      "\n",
      "THEOREM PROVED\n",
      "\n",
      "Exiting with 1 proof.\n",
      "\n",
      "------ process 259687 exit (max_proofs) ------\n",
      "\u0007\n",
      "Process 259687 exit (max_proofs) Mon Apr 29 17:27:10 2024\n"
     ]
    }
   ],
   "source": [
    "! prover9 -f /tmp/sqrt2.in"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 95,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "Overwriting /tmp/sqrt.p\n"
     ]
    }
   ],
   "source": [
    "\n"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 94,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "Overwriting /tmp/prime3.smt2\n"
     ]
    }
   ],
   "source": [
    "%%file /tmp/prime3.smt2\n",
    "(declare-sort I 0)\n",
    "(declare-fun m (I I) I)\n",
    "(declare-fun f (I I) I)\n",
    "(declare-fun d (I I) Bool)\n",
    "(declare-const one I)\n",
    "(declare-const two I)\n",
    "(declare-const a I)\n",
    "(declare-const b I)\n",
    "\n",
    "(assert (forall ((x I)) (= (m one x) x)))\n",
    "(assert (forall ((x I)) (= (m x one) x)))\n",
    "(assert (forall ((x I) (y I) (z I)) (= (m x (m y z)) (m (m x y) z))))\n",
    "(assert (forall ((x I) (y I)) (= (m x y) (m y x))))\n",
    "(assert (forall ((x I) (y I) (z I))\n",
    "  (=> (= (m x y) (m x z)) (= y z))))\n",
    "(assert (forall ((x I) (y I)) (=> (d y x) (= (m y (f y x)) x))))\n",
    "(assert (forall ((x I) (y I))\n",
    "  (=> (d two (m x y)) (or (d two x) (d two y)))))\n",
    "(assert (= (m a a) (m two (m b b)))) ; \n",
    "(assert (forall ((x I)) (=> (and (d x a) (d x b)) (= x one)))) ; lowest common terms\n",
    "(assert (distinct one two))\n",
    "(check-sat)"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 88,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "% Running in auto input_syntax mode. Trying TPTP\n",
      "% Refutation found. Thanks to Tanya!\n",
      "% SZS status Unsatisfiable for sqrt\n",
      "% SZS output start Proof for sqrt\n",
      "3. m(X0,m(X1,X2)) = m(m(X0,X1),X2) [input]\n",
      "4. m(X0,X1) = m(X1,X0) [input]\n",
      "5. m(X0,X1) != m(X0,X2) | X1 = X2 [input]\n",
      "6. ~d(X0,X1) | m(X0,f(X0,X1)) = X1 [input]\n",
      "7. d(X0,X2) | m(X0,X1) != X2 [input]\n",
      "8. ~d(two,m(X0,X1)) | d(two,X0) | d(two,X1) [input]\n",
      "9. m(a,a) = m(two,m(b,b)) [input]\n",
      "10. ~d(X0,b) | one = X0 | ~d(X0,a) [input]\n",
      "11. one != two [input]\n",
      "12. d(X0,m(X0,X1)) [equality resolution 7]\n",
      "21. d(X0,m(X1,X0)) [superposition 12,4]\n",
      "25. d(two,m(a,a)) [superposition 12,9]\n",
      "46. m(a,a) != m(two,X0) | m(b,b) = X0 [superposition 5,9]\n",
      "53. m(X0,X1) = m(X0,f(X0,m(X0,X1))) [resolution 6,12]\n",
      "57. m(a,a) = m(two,f(two,m(a,a))) [resolution 6,25]\n",
      "74. d(X2,m(X0,m(X1,X2))) [superposition 21,3]\n",
      "88. d(two,a) | d(two,a) [resolution 8,25]\n",
      "98. d(two,a) [duplicate literal removal 88]\n",
      "99. a = m(two,f(two,a)) [resolution 98,6]\n",
      "198. m(X0,X1) != m(X0,X2) | f(X0,m(X0,X1)) = X2 [superposition 5,53]\n",
      "264. d(X1,m(X2,m(X1,X0))) [superposition 74,4]\n",
      "486. m(a,X0) = m(two,m(f(two,a),X0)) [superposition 3,99]\n",
      "684. d(two,m(X0,a)) [superposition 264,99]\n",
      "1125. m(a,a) != m(a,a) | m(b,b) = f(two,m(a,a)) [superposition 46,57]\n",
      "1160. m(b,b) = f(two,m(a,a)) [trivial inequality removal 1125]\n",
      "7642. f(X0,m(X0,X1)) = X1 [equality resolution 198]\n",
      "58952. m(f(two,a),X0) = f(two,m(a,X0)) [superposition 7642,486]\n",
      "154801. m(b,b) = m(f(two,a),a) [superposition 58952,1160]\n",
      "154861. d(two,m(b,b)) [superposition 684,154801]\n",
      "163619. d(two,b) | d(two,b) [resolution 154861,8]\n",
      "163621. d(two,b) [duplicate literal removal 163619]\n",
      "163624. one = two | ~d(two,a) [resolution 163621,10]\n",
      "163626. ~d(two,a) [subsumption resolution 163624,11]\n",
      "163627. $false [subsumption resolution 163626,98]\n",
      "% SZS output end Proof for sqrt\n",
      "% ------------------------------\n",
      "% Version: Vampire 4.8 (commit 23d7e9c3b on 2024-03-22 08:41:08 +0000)\n",
      "% Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d\n",
      "% Termination reason: Refutation\n",
      "\n",
      "% Memory used [KB]: 34500\n",
      "% Time elapsed: 10.284 s\n",
      "perf_event_open failed (instruction limiting will be disabled): Permission denied\n",
      "(If you are seeing 'Permission denied' ask your admin to run 'sudo sysctl -w kernel.perf_event_paranoid=-1' for you.)\n",
      "% ------------------------------\n",
      "% ------------------------------\n",
      "\n",
      "real\t0m10.291s\n",
      "user\t0m10.242s\n",
      "sys\t0m0.045s\n"
     ]
    }
   ],
   "source": [
    "! time vampire /tmp/sqrt.p"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 107,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "Overwriting /tmp/sqrtcluase.p\n"
     ]
    }
   ],
   "source": [
    "%%file /tmp/sqrtcluase.p\n",
    "\n",
    "cnf(sos,axiom,\n",
    "    m(one,X0) = X0).\n",
    "\n",
    "cnf(sos,axiom,\n",
    "    m(X0,one) = X0).\n",
    "\n",
    "cnf(sos,axiom,\n",
    "    m(m(X1,X0),X2) = m(X1,m(X0,X2))).\n",
    "\n",
    "cnf(sos,axiom,\n",
    "    m(X0,X1) = m(X1,X0)).\n",
    "\n",
    "cnf(sos,axiom,\n",
    "    X0 = X2 | m(X1,X2) != m(X1,X0)).\n",
    "\n",
    "cnf(sos,axiom,\n",
    "    m(X0,f(X0,X1)) = X1 | ~d(X0,X1)).\n",
    "\n",
    "cnf(sos,axiom,\n",
    "    d(two,X1) | d(two,X0) | ~d(two,m(X1,X0))).\n",
    "\n",
    "cnf(sos,axiom,\n",
    "    m(two,m(b,b)) = m(a,a)).\n",
    "\n",
    "cnf(sos,axiom,\n",
    "    one = X0 | ~d(X0,a) | ~d(X0,b)).\n",
    "\n",
    "cnf(sos,axiom,\n",
    "    one != two).\n",
    "\n",
    "%cnf(sos,axiom,~ d(A,B) | m(A,f(A,B)) = B).\n",
    "cnf(sos,axiom, m(A,B) != C | d(A,C)).\n",
    "%cnf(sos,axiom,~ d(two,m(A,B)) | d(two,A) | d(two,B)).\n",
    "%cnf(sos,axiom,m(a,a) = m(two,m(b,b))).\n",
    "%cnf(sos,axiom,~ d(A,a) | ~ d(A,b) | A = one).\n",
    "%cnf(sos,axiom,~ two = one)."
   ]
  },
  {
   "cell_type": "code",
   "execution_count": null,
   "metadata": {},
   "outputs": [],
   "source": [
    "! time eprover-ho /tmp/sqrt.p"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 21,
   "metadata": {},
   "outputs": [
    {
     "data": {
      "text/plain": [
       "'(declare-fun succ (Nat) Nat)'"
      ]
     },
     "execution_count": 21,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "import z3\n",
    "\n",
    "x = z3.Int(\"x\")\n",
    "e = x+2\n",
    "e.sexpr()\n",
    "x.decl().sexpr()\n",
    "def collect_symbols(expr, symbols):\n",
    "    \"\"\" Recursively collect constants and function symbols from an expression. \"\"\"\n",
    "    if expr.num_args() == 0:\n",
    "        # It's a constant if it has no arguments (excluding bound variables in quantifiers)\n",
    "        if expr.decl().kind() == z3.Z3_OP_UNINTERPRETED:\n",
    "            symbols.add(expr)\n",
    "    else:\n",
    "        # If the node is a function application, add the function declaration\n",
    "        if expr.decl().kind() == z3.Z3_OP_UNINTERPRETED or expr.decl().kind() == z3.Z3_OP_APPLY:\n",
    "            symbols.add(expr.decl())\n",
    "        # Go deeper into the expression for nested function applications\n",
    "        for i in range(expr.num_args()):\n",
    "            collect_symbols(expr.arg(i), symbols)\n",
    "# naw this is ridiculous. I'll just collect up my constants manually.\n",
    "#collect_symbols(e, set())\n",
    "z3.ForAll([x],e > 2).sexpr()\n",
    "Nat = z3.Datatype(\"Nat\")\n",
    "Nat.declare(\"zero\")\n",
    "Nat.declare(\"succ\", (\"prev\", Nat))\n",
    "Nat = Nat.create()\n",
    "Nat.sexpr()\n",
    "Nat.succ.sexpr()"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "## PySMT\n",
    "I thought maybe I should use pysmt. But I dunno. I'm so familiar with the z3 bindings.\n"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 3,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "Serialization of the formula:\n",
      "((((1 <= h) & (h < 10)) & ((1 <= e) & (e < 10)) & ((1 <= l) & (l < 10)) & ((1 <= o) & (o < 10)) & ((1 <= w) & (w < 10)) & ((1 <= r) & (r < 10)) & ((1 <= d) & (d < 10))) & (((h + e + l + l + o) = (w + o + r + l + d)) & ((h + e + l + l + o) = 25)))\n",
      "d := 9\n",
      "o := 5\n",
      "r := 1\n",
      "w := 1\n",
      "h := 1\n",
      "l := 9\n",
      "e := 1\n"
     ]
    }
   ],
   "source": [
    "from pysmt.shortcuts import Symbol, And, GE, LT, Plus, Equals, Int, get_model\n",
    "from pysmt.typing import INT\n",
    "\n",
    "hello = [Symbol(s, INT) for s in \"hello\"]\n",
    "world = [Symbol(s, INT) for s in \"world\"]\n",
    "letters = set(hello+world)\n",
    "domains = And([And(GE(l, Int(1)),\n",
    "                   LT(l, Int(10))) for l in letters])\n",
    "\n",
    "sum_hello = Plus(hello) # n-ary operators can take lists\n",
    "sum_world = Plus(world) # as arguments\n",
    "problem = And(Equals(sum_hello, sum_world),\n",
    "              Equals(sum_hello, Int(25)))\n",
    "formula = And(domains, problem)\n",
    "\n",
    "print(\"Serialization of the formula:\")\n",
    "print(formula)\n",
    "\n",
    "model = get_model(formula)\n",
    "if model:\n",
    "  print(model)\n",
    "else:\n",
    "  print(\"No solution found\")"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 1,
   "metadata": {},
   "outputs": [],
   "source": [
    "from pysmt.logics import QF_UFLRA, QF_UFIDL, QF_LRA, QF_IDL, QF_LIA\n",
    "from pysmt.shortcuts import get_env, GT, Solver, Symbol\n",
    "logics = [QF_UFLRA, QF_UFIDL, QF_LIA],  # Some of the supported logics\n",
    "\n",
    "env = get_env()\n",
    "\n",
    "# Add the solver to the environment\n",
    "env.factory.add_generic_solver(\"vampire\",\"/usr/local/bin/vampire\" , logics)"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 2,
   "metadata": {},
   "outputs": [
    {
     "ename": "AttributeError",
     "evalue": "'list' object has no attribute 'theory'",
     "output_type": "error",
     "traceback": [
      "\u001b[0;31m---------------------------------------------------------------------------\u001b[0m",
      "\u001b[0;31mAttributeError\u001b[0m                            Traceback (most recent call last)",
      "Cell \u001b[0;32mIn[2], line 1\u001b[0m\n\u001b[0;32m----> 1\u001b[0m \u001b[38;5;28;01mwith\u001b[39;00m \u001b[43mSolver\u001b[49m\u001b[43m(\u001b[49m\u001b[43mname\u001b[49m\u001b[38;5;241;43m=\u001b[39;49m\u001b[38;5;124;43m\"\u001b[39;49m\u001b[38;5;124;43mvampire\u001b[39;49m\u001b[38;5;124;43m\"\u001b[39;49m\u001b[43m,\u001b[49m\u001b[43m \u001b[49m\u001b[43mlogic\u001b[49m\u001b[38;5;241;43m=\u001b[39;49m\u001b[43mQF_LIA\u001b[49m\u001b[43m)\u001b[49m \u001b[38;5;28;01mas\u001b[39;00m s:\n\u001b[1;32m      2\u001b[0m     s\u001b[38;5;241m.\u001b[39madd_assertion(formula)\n\u001b[1;32m      3\u001b[0m     \u001b[38;5;28;01mif\u001b[39;00m s\u001b[38;5;241m.\u001b[39msolve():\n",
      "File \u001b[0;32m~/.local/lib/python3.10/site-packages/pysmt/shortcuts.py:908\u001b[0m, in \u001b[0;36mSolver\u001b[0;34m(name, logic, **kwargs)\u001b[0m\n\u001b[1;32m    901\u001b[0m \u001b[38;5;28;01mdef\u001b[39;00m \u001b[38;5;21mSolver\u001b[39m(name\u001b[38;5;241m=\u001b[39m\u001b[38;5;28;01mNone\u001b[39;00m, logic\u001b[38;5;241m=\u001b[39m\u001b[38;5;28;01mNone\u001b[39;00m, \u001b[38;5;241m*\u001b[39m\u001b[38;5;241m*\u001b[39mkwargs):\n\u001b[1;32m    902\u001b[0m \u001b[38;5;250m    \u001b[39m\u001b[38;5;124;03m\"\"\"Returns a solver.\u001b[39;00m\n\u001b[1;32m    903\u001b[0m \n\u001b[1;32m    904\u001b[0m \u001b[38;5;124;03m    :param name: Specify the name of the solver\u001b[39;00m\n\u001b[1;32m    905\u001b[0m \u001b[38;5;124;03m    :param logic: Specify the logic that is going to be used.\u001b[39;00m\n\u001b[1;32m    906\u001b[0m \u001b[38;5;124;03m    :rtype: Solver\u001b[39;00m\n\u001b[1;32m    907\u001b[0m \u001b[38;5;124;03m    \"\"\"\u001b[39;00m\n\u001b[0;32m--> 908\u001b[0m     \u001b[38;5;28;01mreturn\u001b[39;00m \u001b[43mget_env\u001b[49m\u001b[43m(\u001b[49m\u001b[43m)\u001b[49m\u001b[38;5;241;43m.\u001b[39;49m\u001b[43mfactory\u001b[49m\u001b[38;5;241;43m.\u001b[39;49m\u001b[43mSolver\u001b[49m\u001b[43m(\u001b[49m\u001b[43mname\u001b[49m\u001b[38;5;241;43m=\u001b[39;49m\u001b[43mname\u001b[49m\u001b[43m,\u001b[49m\n\u001b[1;32m    909\u001b[0m \u001b[43m                                    \u001b[49m\u001b[43mlogic\u001b[49m\u001b[38;5;241;43m=\u001b[39;49m\u001b[43mlogic\u001b[49m\u001b[43m,\u001b[49m\n\u001b[1;32m    910\u001b[0m \u001b[43m                                    \u001b[49m\u001b[38;5;241;43m*\u001b[39;49m\u001b[38;5;241;43m*\u001b[39;49m\u001b[43mkwargs\u001b[49m\u001b[43m)\u001b[49m\n",
      "File \u001b[0;32m~/.local/lib/python3.10/site-packages/pysmt/factory.py:437\u001b[0m, in \u001b[0;36mFactory.Solver\u001b[0;34m(self, name, logic, **options)\u001b[0m\n\u001b[1;32m    436\u001b[0m \u001b[38;5;28;01mdef\u001b[39;00m \u001b[38;5;21mSolver\u001b[39m(\u001b[38;5;28mself\u001b[39m, name\u001b[38;5;241m=\u001b[39m\u001b[38;5;28;01mNone\u001b[39;00m, logic\u001b[38;5;241m=\u001b[39m\u001b[38;5;28;01mNone\u001b[39;00m, \u001b[38;5;241m*\u001b[39m\u001b[38;5;241m*\u001b[39moptions):\n\u001b[0;32m--> 437\u001b[0m     \u001b[38;5;28;01mreturn\u001b[39;00m \u001b[38;5;28;43mself\u001b[39;49m\u001b[38;5;241;43m.\u001b[39;49m\u001b[43mget_solver\u001b[49m\u001b[43m(\u001b[49m\u001b[43mname\u001b[49m\u001b[38;5;241;43m=\u001b[39;49m\u001b[43mname\u001b[49m\u001b[43m,\u001b[49m\n\u001b[1;32m    438\u001b[0m \u001b[43m                           \u001b[49m\u001b[43mlogic\u001b[49m\u001b[38;5;241;43m=\u001b[39;49m\u001b[43mlogic\u001b[49m\u001b[43m,\u001b[49m\n\u001b[1;32m    439\u001b[0m \u001b[43m                           \u001b[49m\u001b[38;5;241;43m*\u001b[39;49m\u001b[38;5;241;43m*\u001b[39;49m\u001b[43moptions\u001b[49m\u001b[43m)\u001b[49m\n",
      "File \u001b[0;32m~/.local/lib/python3.10/site-packages/pysmt/factory.py:91\u001b[0m, in \u001b[0;36mFactory.get_solver\u001b[0;34m(self, name, logic, **options)\u001b[0m\n\u001b[1;32m     89\u001b[0m \u001b[38;5;28;01mdef\u001b[39;00m \u001b[38;5;21mget_solver\u001b[39m(\u001b[38;5;28mself\u001b[39m, name\u001b[38;5;241m=\u001b[39m\u001b[38;5;28;01mNone\u001b[39;00m, logic\u001b[38;5;241m=\u001b[39m\u001b[38;5;28;01mNone\u001b[39;00m, \u001b[38;5;241m*\u001b[39m\u001b[38;5;241m*\u001b[39moptions):\n\u001b[1;32m     90\u001b[0m     SolverClass, closer_logic \u001b[38;5;241m=\u001b[39m \\\n\u001b[0;32m---> 91\u001b[0m        \u001b[38;5;28;43mself\u001b[39;49m\u001b[38;5;241;43m.\u001b[39;49m\u001b[43m_get_solver_class\u001b[49m\u001b[43m(\u001b[49m\u001b[43msolver_list\u001b[49m\u001b[38;5;241;43m=\u001b[39;49m\u001b[38;5;28;43mself\u001b[39;49m\u001b[38;5;241;43m.\u001b[39;49m\u001b[43m_all_solvers\u001b[49m\u001b[43m,\u001b[49m\n\u001b[1;32m     92\u001b[0m \u001b[43m                              \u001b[49m\u001b[43msolver_type\u001b[49m\u001b[38;5;241;43m=\u001b[39;49m\u001b[38;5;124;43m\"\u001b[39;49m\u001b[38;5;124;43mSolver\u001b[39;49m\u001b[38;5;124;43m\"\u001b[39;49m\u001b[43m,\u001b[49m\n\u001b[1;32m     93\u001b[0m \u001b[43m                              \u001b[49m\u001b[43mpreference_list\u001b[49m\u001b[38;5;241;43m=\u001b[39;49m\u001b[38;5;28;43mself\u001b[39;49m\u001b[38;5;241;43m.\u001b[39;49m\u001b[43msolver_preference_list\u001b[49m\u001b[43m,\u001b[49m\n\u001b[1;32m     94\u001b[0m \u001b[43m                              \u001b[49m\u001b[43mdefault_logic\u001b[49m\u001b[38;5;241;43m=\u001b[39;49m\u001b[38;5;28;43mself\u001b[39;49m\u001b[38;5;241;43m.\u001b[39;49m\u001b[43mdefault_logic\u001b[49m\u001b[43m,\u001b[49m\n\u001b[1;32m     95\u001b[0m \u001b[43m                              \u001b[49m\u001b[43mname\u001b[49m\u001b[38;5;241;43m=\u001b[39;49m\u001b[43mname\u001b[49m\u001b[43m,\u001b[49m\n\u001b[1;32m     96\u001b[0m \u001b[43m                              \u001b[49m\u001b[43mlogic\u001b[49m\u001b[38;5;241;43m=\u001b[39;49m\u001b[43mlogic\u001b[49m\u001b[43m)\u001b[49m\n\u001b[1;32m     98\u001b[0m     \u001b[38;5;28;01mreturn\u001b[39;00m SolverClass(environment\u001b[38;5;241m=\u001b[39m\u001b[38;5;28mself\u001b[39m\u001b[38;5;241m.\u001b[39menvironment,\n\u001b[1;32m     99\u001b[0m                        logic\u001b[38;5;241m=\u001b[39mcloser_logic,\n\u001b[1;32m    100\u001b[0m                        \u001b[38;5;241m*\u001b[39m\u001b[38;5;241m*\u001b[39moptions)\n",
      "File \u001b[0;32m~/.local/lib/python3.10/site-packages/pysmt/factory.py:155\u001b[0m, in \u001b[0;36mFactory._get_solver_class\u001b[0;34m(self, solver_list, solver_type, preference_list, default_logic, name, logic)\u001b[0m\n\u001b[1;32m    150\u001b[0m \u001b[38;5;28;01mif\u001b[39;00m name \u001b[38;5;129;01mnot\u001b[39;00m \u001b[38;5;129;01min\u001b[39;00m solver_list:\n\u001b[1;32m    151\u001b[0m     \u001b[38;5;28;01mraise\u001b[39;00m NoSolverAvailableError(\u001b[38;5;124m\"\u001b[39m\u001b[38;5;132;01m%s\u001b[39;00m\u001b[38;5;124m \u001b[39m\u001b[38;5;124m'\u001b[39m\u001b[38;5;132;01m%s\u001b[39;00m\u001b[38;5;124m'\u001b[39m\u001b[38;5;124m is not available\u001b[39m\u001b[38;5;124m\"\u001b[39m \u001b[38;5;241m%\u001b[39m \\\n\u001b[1;32m    152\u001b[0m                                  (solver_type, name))\n\u001b[1;32m    154\u001b[0m \u001b[38;5;28;01mif\u001b[39;00m logic \u001b[38;5;129;01mis\u001b[39;00m \u001b[38;5;129;01mnot\u001b[39;00m \u001b[38;5;28;01mNone\u001b[39;00m \u001b[38;5;129;01mand\u001b[39;00m \\\n\u001b[0;32m--> 155\u001b[0m    name \u001b[38;5;129;01mnot\u001b[39;00m \u001b[38;5;129;01min\u001b[39;00m \u001b[38;5;28;43mself\u001b[39;49m\u001b[38;5;241;43m.\u001b[39;49m\u001b[43m_filter_solvers\u001b[49m\u001b[43m(\u001b[49m\u001b[43msolver_list\u001b[49m\u001b[43m,\u001b[49m\u001b[43m \u001b[49m\u001b[43mlogic\u001b[49m\u001b[38;5;241;43m=\u001b[39;49m\u001b[43mlogic\u001b[49m\u001b[43m)\u001b[49m:\n\u001b[1;32m    156\u001b[0m     \u001b[38;5;28;01mraise\u001b[39;00m NoSolverAvailableError(\u001b[38;5;124m\"\u001b[39m\u001b[38;5;132;01m%s\u001b[39;00m\u001b[38;5;124m \u001b[39m\u001b[38;5;124m'\u001b[39m\u001b[38;5;132;01m%s\u001b[39;00m\u001b[38;5;124m'\u001b[39m\u001b[38;5;124m does not support logic \u001b[39m\u001b[38;5;132;01m%s\u001b[39;00m\u001b[38;5;124m\"\u001b[39m\u001b[38;5;241m%\u001b[39m\n\u001b[1;32m    157\u001b[0m                                  (solver_type, name, logic))\n\u001b[1;32m    159\u001b[0m SolverClass \u001b[38;5;241m=\u001b[39m solver_list[name]\n",
      "File \u001b[0;32m~/.local/lib/python3.10/site-packages/pysmt/factory.py:361\u001b[0m, in \u001b[0;36mFactory._filter_solvers\u001b[0;34m(self, solver_list, logic)\u001b[0m\n\u001b[1;32m    359\u001b[0m \u001b[38;5;28;01mfor\u001b[39;00m s, v \u001b[38;5;129;01min\u001b[39;00m solver_list\u001b[38;5;241m.\u001b[39mitems():\n\u001b[1;32m    360\u001b[0m     \u001b[38;5;28;01mfor\u001b[39;00m l \u001b[38;5;129;01min\u001b[39;00m v\u001b[38;5;241m.\u001b[39mLOGICS:\n\u001b[0;32m--> 361\u001b[0m         \u001b[38;5;28;01mif\u001b[39;00m \u001b[43mlogic\u001b[49m\u001b[43m \u001b[49m\u001b[38;5;241;43m<\u001b[39;49m\u001b[38;5;241;43m=\u001b[39;49m\u001b[43m \u001b[49m\u001b[43ml\u001b[49m:\n\u001b[1;32m    362\u001b[0m             res[s] \u001b[38;5;241m=\u001b[39m v\n\u001b[1;32m    363\u001b[0m             \u001b[38;5;28;01mbreak\u001b[39;00m\n",
      "File \u001b[0;32m~/.local/lib/python3.10/site-packages/pysmt/logics.py:273\u001b[0m, in \u001b[0;36mLogic.__le__\u001b[0;34m(self, other)\u001b[0m\n\u001b[1;32m    272\u001b[0m \u001b[38;5;28;01mdef\u001b[39;00m \u001b[38;5;21m__le__\u001b[39m(\u001b[38;5;28mself\u001b[39m, other):\n\u001b[0;32m--> 273\u001b[0m     \u001b[38;5;28;01mreturn\u001b[39;00m (\u001b[38;5;28mself\u001b[39m\u001b[38;5;241m.\u001b[39mtheory \u001b[38;5;241m<\u001b[39m\u001b[38;5;241m=\u001b[39m \u001b[43mother\u001b[49m\u001b[38;5;241;43m.\u001b[39;49m\u001b[43mtheory\u001b[49m \u001b[38;5;129;01mand\u001b[39;00m\n\u001b[1;32m    274\u001b[0m             \u001b[38;5;28mself\u001b[39m\u001b[38;5;241m.\u001b[39mquantifier_free \u001b[38;5;241m>\u001b[39m\u001b[38;5;241m=\u001b[39m other\u001b[38;5;241m.\u001b[39mquantifier_free)\n",
      "\u001b[0;31mAttributeError\u001b[0m: 'list' object has no attribute 'theory'"
     ]
    }
   ],
   "source": [
    "with Solver(name=\"vampire\", logic=QF_LIA) as s:\n",
    "    s.add_assertion(formula)\n",
    "    if s.solve():\n",
    "        print(s.get_model())\n",
    "    else:\n",
    "        print(\"No solution found\")"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "\n",
    "### A Manual Proof Discipline\n",
    "Here were some old experiments using raw smtlib idioms\n",
    "\n",
    "Separation between module signature and proofs. Seperate files.\n",
    "Proof datatype to separate from bool definitions\n",
    "\n",
    "```z3\n",
    ";re\n",
    "(declare-datatypes () (\n",
    "  (Proof (Proof (getformula Bool)))\n",
    "))\n",
    "(define-fun proved ((p Proof)) Bool \n",
    "  (match p (\n",
    "    ((Proof p) (not p))\n",
    "  ))\n",
    ")\n",
    "(define-fun instan-lemma1 ((x Int)) Bool  (or (> x 0) (< x 0) (= x 0)))\n",
    "(define-const lemma1 Proof (Proof (forall ((x Int)) (instan-lemma1 x))))\n",
    "\n",
    "(push)\n",
    ";------------\n",
    "(assert (proved lemma1))\n",
    "(check-sat)\n",
    "(pop)\n",
    "\n",
    "```\n",
    "\n",
    "Other possibility for Proof type\n",
    "Forall (Expr)\n",
    "Exists (Expr)\n",
    "Eh. I don't really want to reflect everything.\n",
    "\n",
    "A shallower HOAS embedding\n",
    "\n",
    "```z3\n",
    ";re\n",
    "(declare-datatype Formula (\n",
    "  (Bool (bool1 Bool))\n",
    "  (Forall (forall1 (Array Int Formula)))\n",
    "  (Exists (exists1 (Array Int Formula)))\n",
    "))\n",
    "\n",
    "\n",
    "```\n",
    "\n",
    "```z3\n",
    ";re\n",
    "(declare-datatype Exists \n",
    "  (\n",
    "    (Pair (x Int) (pred (Array Int Bool)))\n",
    "  ))\n",
    "\n",
    "```\n",
    "\n",
    "Universals can be proven by having\n",
    "(define-const x)\n",
    "\n",
    "```\n",
    "; other hypothesis do not mention x\n",
    ";-----------------\n",
    "(instan-lemma1 x)\n",
    "```\n",
    "\n",
    "or by\n",
    "; other hypothesis can't possibly mention x\n",
    ";-----------\n",
    "(assert (proved lemma1))\n",
    "\n",
    "Universal can be used by using lemma1 or by instan-lemma1 on anything\n",
    "\n",
    "Existentials can be proved\n",
    "\n",
    "```\n",
    "----------\n",
    "(exists x (relx x))\n",
    "```\n",
    "\n",
    "or by\n",
    "\n",
    "```\n",
    "whatever\n",
    "---------\n",
    "(relx 3)\n",
    "\n",
    "```\n",
    "\n",
    "They can be used via\n",
    "\n",
    "```\n",
    "define-const x\n",
    "(relx x)\n",
    "; no other use of x\n",
    ";----------------\n",
    "conclusion may refer to x\n",
    "\n",
    "```\n",
    "\n",
    "ExistsE (ExistsE Bool)\n",
    "ForallE (ForallE (forallE Bool)) ; (apply?)\n",
    "Forall (ForallI (forallE Bool))\n",
    "\n",
    "(Exists Int Bool) ; can I pack the witness?\n",
    "\n",
    "define-const lemma Forall (forall x (forallE (instan-lemma x)))\n",
    "define-fun instan-lemma ((x Int)) ForallE\n",
    "\n",
    "The method is related to lambda lifting.\n",
    "We need to give names to different free parameter things.\n",
    "\n",
    "```z3\n",
    ";re\n",
    "(declare-fun even-fun (Int) Int)\n",
    "(define-fun even ((x Int)) Bool (= (* 2 (even-fun x)) x))\n",
    "\n",
    "(declare-fun odd-fun (Int) Int)\n",
    "(define-fun odd ((x Int)) Bool (= (+ 1 (* 2 (odd-fun x))) x))\n",
    "\n",
    ";(define-fun even ((x Int)) Bool (exists ((y Int)) (= (* 2 y) x)))\n",
    "(define-const even-axiom Bool (forall ((x Int)) (=> (exists ((y Int)) (= (* 2 y) x)) (= (* 2 (even-fun x)) x))))\n",
    ";(define-fun even ((x Int)) Bool (exists ((y Int)) (and (= y (even-fun x)) (= (* 2 y) x))))\n",
    "(define-fun not-even ((x Int)) Bool (forall ((y Int)) (distinct x (* 2 y))))\n",
    ";(define-fun instan-not-even ((x Int) (y Int)) (distinct x (2 * y)))\n",
    "\n",
    "(declare-const m Int)\n",
    "(declare-const n Int)\n",
    "(assert (odd m))\n",
    "(assert (odd n))\n",
    "(assert (not (even m)))\n",
    "(assert (not (even n)))\n",
    ";(assert (instan-not-even m (even-fun )))\n",
    "(assert-not (even (+ n m)))\n",
    ";(assert (not-even (+ n m)))\n",
    "(check-sat)\n",
    "(get-model)\n",
    "\n",
    "```\n",
    "\n",
    "```z3\n",
    ";re\n",
    "(define-fun even ((x Int) (y Int)) Bool (= (* 2 y) x))\n",
    "(define-fun odd ((x Int) (y Int)) Bool (= (+ 1 (* 2 y)) x))\n",
    "\n",
    "\n",
    "(declare-const m Int)\n",
    "(declare-const n Int)\n",
    "(declare-const mo Int)\n",
    "(declare-const no Int)\n",
    "(assert (odd m mo))\n",
    "(assert (odd n no))\n",
    "(assert-not (exists ((x Int)) (even (+ m n) x)))\n",
    "(check-sat)\n",
    "\n",
    "```\n",
    "\n",
    "```z3\n",
    ";re\n",
    "(define-fun even ((x Int)) Bool (exists ((y Int)) (= (* 2 y) x)))\n",
    "(define-fun odd ((x Int)) Bool (exists ((y Int)) (= (+ 1 (* 2 y)) x)))\n",
    "\n",
    "(declare-const m Int)\n",
    "(declare-const n Int)\n",
    "(assert (odd m))\n",
    "(assert (odd n))\n",
    ";--------------------------\n",
    "(assert-not (even (+ m n)))\n",
    "(check-sat)\n",
    "```\n",
    "\n",
    "```z3\n",
    "\n",
    "```z3\n",
    ";re\n",
    "(define-fun even ((x Int)) Bool (exists ((y Int)) (= (* 2 y) x)))\n",
    "(define-fun div4 ((x Int)) Bool (exists ((y Int)) (= (* 4 y) x)))\n",
    "\n",
    "(declare-const m Int)\n",
    "(assert (not (even m)))\n",
    ";---------------------------\n",
    "(assert-not (not (div4 m)))\n",
    "(check-sat)\n",
    "```\n",
    "\n",
    "Positively asserting existence isn't so bad. It is just asserting on a fresh variable.\n",
    "\n",
    "Only allow mention variable once on existential elim in hyps\n",
    "\n",
    "I'm a little worried about circulaity if I pack everything into one proof thing.\n",
    "\n",
    "existential lemma can be used in hyps.\n",
    "Instantiated version in a conc (give explicitly). Or just let it crank.\n",
    "\n",
    "```\n",
    "(declare-const e) appears only once in hyp for existential elim\n",
    "\n",
    "(assert-not\n",
    "  (=> (and (existse e) (instan-lemma1 ))\n",
    "      (and lemma1 )\n",
    "  )\n",
    "\n",
    ")\n",
    "\n",
    "```\n",
    "\n",
    "```z3\n",
    ";re\n",
    "(define-fun instan-theorem ((m Int) (n Int)) Bool (not (= (* m m) (* 2 (* n n)))))\n",
    "(define-const theorem Bool (forall ((m Int) (n Int)) (not (= (* m m) (* 2 (* n n))))))\n",
    "\n",
    "\n",
    "```"
   ]
  },
  {
   "cell_type": "markdown",
   "metadata": {},
   "source": [
    "### Python Hilbert Proof\n",
    "An older database style attempt\n",
    "\n",
    "Hilbert style proofs where the only inference rule is Z3_check.\n",
    "Register known theorems to central repository. Refer to them later.\n",
    "There are some basic theorems with quantifiers I couldn't get Z3 to do even the most basic step of, so this idea feels sunk.\n",
    "I guess I could also admit vampire.\n",
    "\n",
    "```python\n",
    "from z3 import *\n",
    "class Kernel():\n",
    "    def __init__(self):\n",
    "        self._formula = {}\n",
    "        self._schema = {}\n",
    "        self._explanations = {}\n",
    "    def axiom(self, name, formula):\n",
    "        assert name not in self._formula\n",
    "        self._formula[name] = (formula, None)\n",
    "        return name\n",
    "    def add_schema(self, name, schema):\n",
    "        assert name not in self._schema\n",
    "        self._schema[name] = schema\n",
    "        return name\n",
    "    def instan_schema(self, schema, name, *args):\n",
    "        #name = self.fresh(name)\n",
    "        self._formula[name] = (self._schema[schema](*args), (\"schema\", schema))\n",
    "        return name\n",
    "    def theorem(self, name, formula, *reasons, timeout=1):\n",
    "        assert name not in self._formula\n",
    "        s = Solver()\n",
    "        #s.set(\"timeout\", timeout)\n",
    "\n",
    "        s.add([ self._formula[reason][0] for reason in reasons])\n",
    "        # snity check here. s.check(). should return unknown. If returns unsat, we have an incosisitency\n",
    "        # Even if we have inconsistency but z3 isn't finding it, that is a small comfort. \n",
    "        s.add(Not(formula))\n",
    "        status = s.check()\n",
    "        if status == unsat:\n",
    "            self._formula[name] = (formula, reasons)\n",
    "            print(f\"Accepted {name}\")\n",
    "        elif status == sat:\n",
    "            print(f\"Counterexample : {s.model()}\")\n",
    "        elif status == unknown:\n",
    "            print(\"Failed: reason unknown\")\n",
    "        else:\n",
    "            print(f\"unexpected status {status}\")\n",
    "    #def calc(self):\n",
    "    #    pass\n",
    "    #def definition: # is definition distinct from axiom?\n",
    "\n",
    "t = Kernel()\n",
    "p = Real(\"p\")\n",
    "t.axiom(\"p_def\", p*p == 2)\n",
    "m,n = [ToReal(x) for x in Ints(\"m n\")]\n",
    "\n",
    "#t.theorem(\"p irrational\", p != m / n)\n",
    "#(m / n == p)\n",
    "def even(x):\n",
    "    y = FreshInt()\n",
    "    return Exists([y], 2 * y == x)\n",
    "def odd(x):\n",
    "    y = FreshInt()\n",
    "    return Exists([y], 2 * y + 1 == x)\n",
    "r,w = Ints(\"r w\")\n",
    "t.theorem(\"even odd\", ForAll([r], even(r) != odd(r)))\n",
    "def rational(p):\n",
    "    m = FreshInt(\"m\")\n",
    "    n = FreshInt(\"n\")\n",
    "    return Exists([m,n], And(n != 0, p == ToReal(m) / ToReal(n) ))\n",
    "rational(p)\n",
    "\n",
    "#t.theorem(\"mndef\", Implies(rational(p), p == m / n) )\n",
    "\n",
    "#t.theorem( Implies( p == m / n, Exists([r,w], And(p == r / w,  Or(And(even(r), odd(w)), And(odd(r), even(w)) ) ))))\n",
    "# how do you use this exists?\n",
    "\n",
    "#t.theorem(   )\n",
    "def evenodd(n,m):\n",
    "    return even(n) != even(m)\n",
    "\n",
    "t.theorem(0, Implies(  And(p == m / n, n != 0) , p * n == m  ))\n",
    "t.theorem(1, Implies(  And(p == m / n, n != 0, evenodd(m,n)) , p * n == m  ))\n",
    "#t.theorem(2, Implies(  And(p == m / n, n != 0, evenodd(m,n)) , 2 * n * n == m * m , \"p_def\"))\n",
    "t.theorem(3, even(r) != odd(r))\n",
    "#t.theorem(4, even(r * r) != odd(r * r))\n",
    "t.theorem(5, Implies(r == 2 * w, even(r*r) ))\n",
    "t.theorem(6, Implies(r == 2 * w, r*r == 4 * w * w ))\n",
    "\n",
    "#t.theorem(7, Implies(r == 2 * w, even(r) ))\n",
    "#t.theorem(8, Implies(even(r), even(r*r)), 5, 6, 7)\n",
    "t.theorem(\"even or odd\", Or(even(r), odd(r)))\n",
    "#t.theorem(\"even or odd r*r\", Or(even(r*r), odd(r*r)))\n",
    "t.theorem(8, ForAll([r], even(r) != odd(r)))\n",
    "t.theorem(9, even(r*r) != odd(r*r), 8)\n",
    "# I dunno. This idea is sunk. I can't even get this to work?\n",
    "# I could try not using the built in stuff I guess and wwork in FOL.\n",
    "\n",
    "\n",
    "#t.theorem(5, odd(r) == odd(r * r), 3, 4)\n",
    "#t.theorem(3, Implies(even(r * r), even(r) ))\n",
    "#t.theorem(3, Implies(  And(2 * n * n == m * m), even(m)))\n",
    "\n",
    "\n",
    "\n",
    "\n",
    "#t.theorem(0, Implies(  p == m / n , m * m == 2 * n * n   ), \"p_def\" )\n",
    "even(r * r) != odd(r * r)\n",
    "```\n",
    "\n",
    "```python\n",
    "t = Hilbert()\n",
    "\n",
    "Nat = Datatype('Nat')\n",
    "Nat.declare('Zero')\n",
    "Nat.declare('Succ', ('pred', Nat))\n",
    "# Create the datatype\n",
    "Nat = Nat.create()\n",
    "Zero = Nat.Zero\n",
    "Succ = Nat.Succ\n",
    "\n",
    "n,m = Consts(\"n m\", Nat)\n",
    "plus = Function(\"+\", Nat,Nat,Nat)\n",
    "DatatypeRef.__add__ = lambda self, rhs : plus(self, rhs)\n",
    "#type(n)\n",
    "plus_zero = t.axiom(\"plus_zero\", ForAll([m], Zero + m == m))\n",
    "plus_succ = t.axiom(\"plus_succ\", ForAll([n,m], Succ(n) + m == Succ(n + m)))\n",
    "plus_def = [plus_zero, plus_succ]\n",
    "def induction_schema(P):\n",
    "    return  Implies(\n",
    "                And(P(Zero), \n",
    "                    ForAll([n], Implies(P(n), P(Succ(n))))) ,\n",
    "                ForAll([n], P(n))\n",
    "    )\n",
    "induction = t.add_schema(\"induction\", induction_schema)\n",
    "\n",
    "t.theorem( \"plus_zero'\",  ForAll([m], Zero + m == m , patterns=[Zero + m]), plus_zero) # guardedc version of the axiom\n",
    "mylemma = t.instan(induction, \"mylemma\", lambda m : Zero + m == m + Zero )\n",
    "\n",
    "t.theorem( \"n_plus_zero\", ForAll([m], Zero + m == m + Zero) ,  plus_zero, plus_succ,  mylemma )\n",
    "t.theorem(\"test guarding\", Zero + Zero == Succ(Zero), \"plus_zero\")\n",
    "\n",
    "lte = Function(\"<=\", Nat, Nat, BoolSort())\n",
    "DatatypeRef.__le__ = lambda self, rhs : lte(self, rhs)\n",
    "\n",
    "t.axiom(\"zero lte\", ForAll([n], Zero <= n))\n",
    "t.axiom(\"succ lte\", ForAll([n,m], (Succ(n) <= Succ(m)) == (n <= m) ))\n",
    "lte_def = [\"zero lte\", \"succ lte\"]\n",
    "\n",
    "\n",
    "t.theorem(\"example\",  (Zero + Zero) <= Zero,  *plus_def, *lte_def  )\n",
    "\n",
    "\n",
    "reflect = Function(\"reflect\", Nat  , IntSort())\n",
    "x = Int(\"x\")\n",
    "t.axiom( \"reflect succ\", ForAll([x, n], (reflect(Succ(n)) == 1 + x) == (reflect(n) == x) ))\n",
    "t.axiom( \"reflect zero\", reflect(Zero) == 0) \n",
    "\n",
    "t.formula\n",
    "\n",
    "```"
   ]
  }
 ],
 "metadata": {
  "kernelspec": {
   "display_name": "Python 3",
   "language": "python",
   "name": "python3"
  },
  "language_info": {
   "codemirror_mode": {
    "name": "ipython",
    "version": 3
   },
   "file_extension": ".py",
   "mimetype": "text/x-python",
   "name": "python",
   "nbconvert_exporter": "python",
   "pygments_lexer": "ipython3",
   "version": "3.12.3"
  }
 },
 "nbformat": 4,
 "nbformat_minor": 2
}
