{
 "cells": [
  {
   "cell_type": "markdown",
   "id": "400213d9",
   "metadata": {},
   "source": [
    "Very confused what I am looking at with the ordered slotted egraph.\n",
    "\n",
    "Sound ematching, not complete.\n",
    "\n",
    "Maybe I shouldn't traverse redudnancy edges during ematching anyway, because those matches probably suck. They are going to be searching through dump expressions.\n",
    "\n",
    "Seed roots with non dump redundant pieces. This also means rebuilding may not want to compact away as much.\n",
    "\n",
    "equivalraince l(f(X,Y)) = f(l(X), l(Y)) is used to pull l up during canonization, but also can be used during ematching to pull l into the current target term to get f heads to match\n",
    "l(f(e1,e2)) =? f(X,Y)\n",
    "f(l(e1), l(e2)) =? f(X,Y)\n",
    "{X =? l(e1), Y =? l(e2)}\n",
    "\n",
    "\n",
    "traversing up  e -> ll(etop)\n",
    "becomes etop -> ddd(e)  which is not fun.\n",
    "\n",
    "((X*Y)*Z) associativity is no problemo.\n",
    "\n",
    "Ok what about lambda then. Lambda is not equivariant. lam(l1(B)) = l2(lam(B))\n",
    "\n",
    "\n",
    "l(e) = f(X,g(Y,Z))\n",
    "l(f(e1,e2)) = f(X,g(Y,Z))\n",
    "f(l(e1),l(e2)) = ..\n",
    "l(e1) = X, l(e2) = g(Y,Z)\n",
    "\n",
    "\n",
    "l being least in lpo does track. That would pull l to the outside.\n",
    "l(f(a,f(b,c)))\n",
    "dump is uneccessary?\n",
    "\n",
    "What if l was just add a variable to front. Always l1 l0 with the rest being understand to be true.\n",
    "\n",
    "l0(l1(l0()))\n",
    "\n",
    "coherent. uhh. No. But this won't let us do as much sharing.\n",
    "y = l0(l1(v))\n",
    " x =? l1(l0(v))? no.\n",
    "\n",
    " shift(l)\n",
    "\n",
    "l(n, m, X)\n",
    "\n",
    "Huh. No.\n",
    "\n",
    "pulling l outward\n",
    "but always e -> l(e)\n",
    "is tought to get to happen.\n",
    "\n",
    "\n",
    "\n",
    "dump nodes\n",
    "floordiv(x, 3) is kind of a skolem node. We could just not normalize these further. (even thought dump dump  probably merges and dumps should push down)\n",
    "\n",
    "\n",
    "l\n",
    "\n",
    "\n",
    "`e -> l(e)`  then we want `e > l`\n",
    "`f(l(X)) -> l(f(X))` then we want `f > l`\n",
    "So `l` has to be least in the ordered. Every eid should have a weight corresponding to it's size.\n",
    "\n",
    "In this simplified subsystem, I can't have incompatable l0(e) = l1(e).\n",
    "\n",
    "x != y % not sure it can do much with this.\n",
    "swap(x) = y\n",
    "swap(y) = x\n",
    "swap(swap(X)) = X\n",
    "\n",
    "f(e1(x,y), ) -> e2(x)\n",
    "\n",
    "swap(e2(x)) = e2(y)\n",
    "ordering the slots.\n",
    "Is this the thing rudi and I did? maybe.\n",
    "x > y > z, normalize all e. Is there a point to e(x,y,z) form? We just know it'll be 0,1,2?\n",
    "\n",
    "x + y = x + swap(x)\n",
    "\n",
    "e(x,y) = e2(y)\n",
    "e(x,y) = swap(y)\n",
    "\n",
    "e -> swap(e)\n",
    "implies \n",
    "e > swap\n",
    "swap weight = 0?\n",
    "swap(swap(X)) = X\n",
    "swap(f(X,Y)) = f(swap(X), swap(Y))\n",
    "swap(x) = y\n",
    "swap(y) = x\n",
    "\n",
    "x:1,swap:1,y:2 ? swap(x) = y is same weight, it will prefer swaps\n",
    "y -> swap(x)\n",
    "swap_xy == weight(x) - weight(y) ? for multivariate. Analog of weight(lift) = cod(lift) - dom(lift)\n",
    "x > y > z > ... to get sorting.\n",
    "\n",
    "\n",
    "e1(x,y) = e2(y)\n",
    "swap(e1(x,y)) = swap(e2(y))\n",
    "e1(y,x) = e2(x)\n",
    "sure.\n",
    "x,y,z\n",
    "\n",
    "e2(x) = swap(e1(y,x))\n",
    "\n",
    " hmm. My previous example did loop.\n",
    "\n",
    "e at different level should never be being compared directly without l in between.\n",
    "And at a level it's abritrary\n",
    "\n",
    "lpo = the worst subterm gets better (what is worst? The worst term is the top one)\n"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 53,
   "id": "78796752",
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "Overwriting /tmp/swap.p\n"
     ]
    }
   ],
   "source": [
    "%%file /tmp/swap.p\n",
    "\n",
    "cnf(swap, axiom, swap(x) = y).\n",
    "cnf(swapy, axiom, swap(y) = x).\n",
    "cnf(swap_involution, axiom, swap(swap(X)) = X).\n",
    "cnf(swap_mul, axiom, swap(mul(X,Y)) = mul(swap(X), swap(Y))).\n",
    "\n",
    "% mul(swap(X), Y) = swap(mul(X, swap(Y))). That's a good point. remove swap from first and add to all other args.\n",
    "\n",
    "cnf(swap_zero, axiom, swap(zero) = zero).\n",
    "cnf(mul_zero, axiom, mul(zero,zero) = zero).\n",
    "cnf(mulzero, axiom, mul(x,zero) = zero).\n",
    "cnf(mulzero, axiom, mul(y,zero) = zero).\n",
    "%cnf(mulxy, axiom, mul(x,y) = mul(y,x)).\n"
   ]
  },
  {
   "cell_type": "markdown",
   "id": "ac3ccb16",
   "metadata": {},
   "source": [
    "Hmm. So I can't get kbo to work. I think slotted would want to push swap down? Maybe this is why I need e(x,y,z)?\n",
    "\n",
    "x,y high to get rid of them in redundancies?\n",
    "swap low? to push swap(mul) -> mul(swap, swap) ?\n",
    "symmettries screw you?\n"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 56,
   "id": "95cf818a",
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "% Preprocessing class: FSSSSMSSSSSNFFN.\n",
      "% Configuration: G-E--_302_C18_F1_URBAN_RG_S04BN\n",
      "% (lift_lambdas = 1, lambda_to_forall = 1,unroll_only_formulas = 1, sine = Auto)\n",
      "% No SInE strategy applied\n",
      "% Search class: FUUPS-FFSM21-SFFFFFNN\n",
      "% partial match(1): FUUPS-FFSF21-SFFFFFNN\n",
      "% Configuration: G-E--_208_C12_00_F1_SE_CS_PI_SP_PS_S5PRR_RG_S04AN\n",
      "% Presaturation interreduction done\n",
      "\n",
      "% No proof found!\n",
      "% SZS status Satisfiable\n",
      "% Processed positive unit clauses:\n",
      "cnf(i_0_9, plain, (y=swap(x))).\n",
      "cnf(i_0_13, plain, (swap(zero)=zero)).\n",
      "cnf(i_0_11, plain, (swap(swap(X1))=X1)).\n",
      "cnf(i_0_15, plain, (mul(x,zero)=zero)).\n",
      "cnf(i_0_14, plain, (mul(zero,zero)=zero)).\n",
      "cnf(i_0_19, plain, (mul(swap(X1),X2)=swap(mul(X1,swap(X2))))).\n",
      "cnf(i_0_26, plain, (mul(zero,swap(X1))=swap(mul(zero,X1)))).\n",
      "\n",
      "% Processed negative unit clauses:\n",
      "\n",
      "% Processed non-unit clauses:\n",
      "\n",
      "% Unprocessed positive unit clauses:\n",
      "\n",
      "% Unprocessed negative unit clauses:\n",
      "\n",
      "% Unprocessed non-unit clauses:\n",
      "\n",
      "\n"
     ]
    }
   ],
   "source": [
    "! eprover-ho -t LPO4  -s  --auto --precedence=\"y > x > mul > zero > swap\"  --print-saturated /tmp/swap.p"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": null,
   "id": "f8c82be5",
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "% Preprocessing class: FSSSSMSSSSSNFFN.\n",
      "% Configuration: G-E--_302_C18_F1_URBAN_RG_S04BN\n",
      "% (lift_lambdas = 1, lambda_to_forall = 1,unroll_only_formulas = 1, sine = Auto)\n",
      "% No SInE strategy applied\n",
      "% Search class: FUUPS-FFSS21-SFFFFFNN\n",
      "% partial match(1): FUUPS-FFSF21-SFFFFFNN\n",
      "% Configuration: G-E--_208_C12_00_F1_SE_CS_PI_SP_PS_S5PRR_RG_S04AN\n",
      "setting user weights\n",
      "^C\n"
     ]
    }
   ],
   "source": [
    "! eprover-ho -t KBO6  -s --auto --precedence=\"mul > zero > x > y > swap\" --order-weights=\"mul:1,zero:1,swap:1,x:3,y:4\"  --print-saturated /tmp/swap.p"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": null,
   "id": "9c8f9acf",
   "metadata": {},
   "outputs": [],
   "source": []
  },
  {
   "cell_type": "code",
   "execution_count": 60,
   "id": "f454b0bf",
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "% Preprocessing class: FSMSSMSSSSSNFFN.\n",
      "% Configuration: G-E--_208_C18_F1_SE_CS_SOS_SP_PS_S5PRR_RG_S04AN\n",
      "% (lift_lambdas = 1, lambda_to_forall = 1,unroll_only_formulas = 1, sine = Auto)\n",
      "% No SInE strategy applied\n",
      "% Search class: FGHSF-FFSF21-SFFFFFNN\n",
      "% Configuration: SAT001_MinMin_p005000_rr_RG\n",
      "% Presaturation interreduction done\n",
      "\n",
      "% Proof found!\n",
      "% SZS status Unsatisfiable\n",
      "% SZS output start CNFRefutation\n",
      "cnf(clause_4, negated_conjecture, (f(X2)=X2|~big_f(X1,f(X2))|X1!=g(X2)), file('/tmp/test.p', clause_4)).\n",
      "cnf(clause_6, negated_conjecture, (big_f(X1,f(X2))|f(X2)=X2|X1!=g(X2)), file('/tmp/test.p', clause_6)).\n",
      "cnf(clause_9, negated_conjecture, (big_f(h(X1,X2),f(X1))|h(X1,X2)=X2|f(X1)!=X1), file('/tmp/test.p', clause_9)).\n",
      "cnf(clause_10, negated_conjecture, (f(X1)!=X1|h(X1,X2)!=X2|~big_f(h(X1,X2),f(X1))), file('/tmp/test.p', clause_10)).\n",
      "cnf(clause_1, axiom, (X1=a|~big_f(X1,X2)), file('/tmp/test.p', clause_1)).\n",
      "cnf(clause_3, axiom, (big_f(X1,X2)|X1!=a|X2!=b), file('/tmp/test.p', clause_3)).\n",
      "cnf(c_0_6, negated_conjecture, (f(X2)=X2|~big_f(X1,f(X2))|X1!=g(X2)), inference(fof_simplification,[status(thm)],[clause_4])).\n",
      "cnf(c_0_7, negated_conjecture, (big_f(X1,f(X2))|f(X2)=X2|X1!=g(X2)), inference(fof_simplification,[status(thm)],[clause_6])).\n",
      "cnf(c_0_8, negated_conjecture, (f(X2)=X2|~big_f(X1,f(X2))|X1!=g(X2)), c_0_6).\n",
      "cnf(c_0_9, negated_conjecture, (big_f(h(X1,X2),f(X1))|h(X1,X2)=X2|f(X1)!=X1), inference(fof_simplification,[status(thm)],[clause_9])).\n",
      "cnf(c_0_10, negated_conjecture, (big_f(X1,f(X2))|f(X2)=X2|X1!=g(X2)), c_0_7).\n",
      "cnf(c_0_11, negated_conjecture, (f(X1)=X1|~big_f(g(X1),f(X1))), inference(er,[status(thm)],[c_0_8])).\n",
      "cnf(c_0_12, negated_conjecture, (f(X1)!=X1|h(X1,X2)!=X2|~big_f(h(X1,X2),f(X1))), inference(fof_simplification,[status(thm)],[clause_10])).\n",
      "cnf(c_0_13, plain, (X1=a|~big_f(X1,X2)), inference(fof_simplification,[status(thm)],[clause_1])).\n",
      "cnf(c_0_14, negated_conjecture, (big_f(h(X1,X2),f(X1))|h(X1,X2)=X2|f(X1)!=X1), c_0_9).\n",
      "cnf(c_0_15, negated_conjecture, (f(X1)=X1), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_10]), c_0_11])).\n",
      "cnf(c_0_16, negated_conjecture, (f(X1)!=X1|h(X1,X2)!=X2|~big_f(h(X1,X2),f(X1))), c_0_12).\n",
      "cnf(c_0_17, plain, (X1=a|~big_f(X1,X2)), c_0_13).\n",
      "cnf(c_0_18, negated_conjecture, (h(X1,X2)=X2|big_f(h(X1,X2),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_15]), c_0_15])])).\n",
      "cnf(c_0_19, plain, (big_f(X1,X2)|X1!=a|X2!=b), inference(fof_simplification,[status(thm)],[clause_3])).\n",
      "cnf(c_0_20, negated_conjecture, (h(X1,X2)!=X2|~big_f(h(X1,X2),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_15]), c_0_15])])).\n",
      "cnf(c_0_21, negated_conjecture, (h(X1,X2)=a|h(X1,X2)=X2), inference(spm,[status(thm)],[c_0_17, c_0_18])).\n",
      "cnf(c_0_22, plain, (big_f(X1,X2)|X1!=a|X2!=b), c_0_19).\n",
      "cnf(c_0_23, negated_conjecture, (h(X1,X2)=a|~big_f(X2,X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).\n",
      "cnf(c_0_24, plain, (big_f(a,b)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_22])])).\n",
      "cnf(c_0_25, negated_conjecture, (~big_f(a,X1)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_23])])).\n",
      "cnf(c_0_26, plain, ($false), inference(sr,[status(thm)],[c_0_24, c_0_25]), ['proof']).\n",
      "% SZS output end CNFRefutation\n"
     ]
    }
   ],
   "source": [
    "! eprover-ho --auto --cpu-limit=1 --silent --proof-object /tmp/test.p"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 116,
   "id": "3982c279",
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "Overwriting /tmp/lift.p\n"
     ]
    }
   ],
   "source": [
    "%%file /tmp/lift.p\n",
    "\n",
    "cnf(mulname, axiom, mul(id, l(zero)) = e0).\n",
    "cnf(union, axiom, e0 = l(zero)).\n",
    "cnf(redundnant, axiom, mul(l(id), l(l(zero))) = l(l(zero))). % if you don't have \"union\" injectivity + homo will derive it.\n",
    "\n",
    "\n",
    "\n",
    "cnf(mul_id_id, axiom, mul(id,id) = e1). % x**2. no particular identities apply\n",
    "\n",
    "cnf(mul_zero_id, axiom, mul(zero,zero) = e2).\n",
    "cnf(mulzerozero, axiom, e2 = zero).\n",
    "\n",
    "cnf(mul_comm, axiom, mul(id, l(zero)) = mul(l(zero),id)).\n",
    "cnf(mul_assoc, axiom, mul(id,mul(id,l(zero))) = mul(mul(id,id),l(zero))).\n",
    "\n",
    "\n",
    "\n",
    "cnf(linj, axiom  , l(X) != l(Y) | X = Y).\n",
    "%cnf(lequivf, axiom, l(f(X)) = f(l(X))).\n",
    "cnf(mul_homo, axiom, l(mul(X,Y)) = mul(l(X),l(Y)))."
   ]
  },
  {
   "cell_type": "code",
   "execution_count": null,
   "id": "86e0379b",
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "% (lift_lambdas = 1, lambda_to_forall = 1,unroll_only_formulas = 1, sine = (null))\n",
      "^C\n"
     ]
    }
   ],
   "source": [
    "! eprover-ho --precedence=\"mul > zero > e1 > e0 > e2 > id > l\" -t LPO4  /tmp/lift.p --silent --print-oriented-eqlits-as-rules --print-saturated"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 124,
   "id": "8bb1704c",
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "% (lift_lambdas = 1, lambda_to_forall = 1,unroll_only_formulas = 1, sine = (null))\n",
      "setting user weights\n",
      "\n",
      "% No proof found!\n",
      "% SZS status Satisfiable\n",
      "% Processed positive unit clauses:\n",
      "cnf(i_0_16, plain, (zero->e2)).\n",
      "cnf(i_0_12, plain, (e0->l(e2))).\n",
      "cnf(i_0_14, plain, (mul(id,id)->e1)).\n",
      "cnf(i_0_15, plain, (mul(e2,e2)->e2)).\n",
      "cnf(i_0_11, plain, (mul(id,l(e2))->l(e2))).\n",
      "cnf(i_0_20, plain, (mul(l(X1),l(X2))->l(mul(X1,X2)))).\n",
      "cnf(i_0_17, plain, (mul(l(e2),id)->l(e2))).\n",
      "cnf(i_0_18, plain, (mul(e1,l(e2))->l(e2))).\n",
      "\n",
      "% Processed negative unit clauses:\n",
      "\n",
      "% Processed non-unit clauses:\n",
      "cnf(i_0_19, plain, (X1->X2|l(X1)!->l(X2))).\n",
      "\n",
      "% Unprocessed positive unit clauses:\n",
      "\n",
      "% Unprocessed negative unit clauses:\n",
      "\n",
      "% Unprocessed non-unit clauses:\n",
      "\n",
      "\n"
     ]
    }
   ],
   "source": [
    "! eprover-ho --precedence=\"mul > zero > id > e2 > e0  > e1 > l\" -t KBO6 --order-weights=\"id:2,zero:1,l:1,e0:2,e1:2,e2:1,mul:1\" /tmp/lift.p --silent --print-oriented-eqlits-as-rules --print-saturated"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 111,
   "id": "a5974b85",
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "Overwriting /tmp/lift.p\n"
     ]
    }
   ],
   "source": [
    "%%file /tmp/lift.p\n",
    "\n",
    "cnf(a1, axiom, l(e1) = e2).\n",
    "%cnf(a2, axiom, l(e2) = e3).\n",
    "cnf(a3, axiom, f(l(e1)) = e4).\n",
    "% l injective also\n",
    "cnf(linj, axiom  , l(X) != l(Y) | X = Y).\n",
    "cnf(lequivf, axiom, l(f(X)) = f(l(X))).\n"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": null,
   "id": "5cd1d330",
   "metadata": {},
   "outputs": [],
   "source": [
    "! eprover-ho   --print-oriented-eqlits-as-rules  --term-ordering=LPO4 --precedence=\"f  > e2 > e4 > e1 >  l\" --print-saturated /tmp/lift.p"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": null,
   "id": "2a349c08",
   "metadata": {},
   "outputs": [],
   "source": [
    "from dataclasses import dataclass, field\n",
    "\n",
    "type Thin = list[bool]\n",
    "\n",
    "@dataclass\n",
    "class Node:\n",
    "    f : str\n",
    "    args : tuple\n",
    "\n",
    "@dataclass\n",
    "class EGraph:\n",
    "    memo : dict\n",
    "    uf : list[int]\n",
    "\n"
   ]
  },
  {
   "cell_type": "markdown",
   "id": "8071c266",
   "metadata": {},
   "source": [
    "tree-like proof relevant partial orders. I expected to have to sometimes do LCA or traverse down the uf?\n",
    "\n",
    "unary rewriting theory combination. If I have a normalization procedure for a string system encoded as unary symbols,\n",
    "I don't think I can actually hurt it by adding in new ground eqautions?\n",
    "Including ones.\n",
    "\n",
    "\n",
    "Slotted - information goes down, thinning information goes up\n",
    "thin is alpha euqal on the nose. slotted is disequal but g related\n",
    "\n",
    "\n",
    "\n",
    "\n",
    "\n",
    "\n",
    "String knuth bendix with \"fat\" characters\n",
    "coproduct knuth bendix.\n",
    "\n",
    "Frex\n",
    "\n",
    "\n",
    "Tower of equivalances\n",
    "\n",
    "sin(pi * (?x + 1)) = cos(pi * ?x)\n",
    "\n",
    "Ematching needs an unapply interface that is group aware?\n",
    "\n",
    "offset, linear enodes, \n",
    "proofs need to record how ematching worked.\n",
    "\n",
    "versioned egraphs on edges?\n",
    "\n",
    "\n",
    "Representations of Relations. But we also have things on the nodes.\n",
    "Convert constraint graphs into constraint trees\n",
    "\n",
    "Skolems\n",
    "fn reify(x : FatId) -> RawId\n",
    "\n",
    "\n",
    "\n",
    "What the hell just happened in ordered slotted\n",
    "v10 = lambda *args: args[10]\n",
    "sin is still lifted to sin'\n",
    "\n",
    "AppliedId\n",
    "\n",
    "e12(2,4,6)\n",
    "f(v1,v0) is fine. Arguments are not linked to slots.\n",
    "\n",
    "Using the opposite form of thinning. Sure.\n",
    "type Wide = list[int]\n",
    "\n",
    "\n",
    "Named thinnings. Maybe this is a useful concept to figure out what is the relation of thinning and slotted/redundancy. A thinning is a bitvector recipe saying how to find one list as a subsequence of anothet list. A named thinning is a boolean valued struct as a recipe on how to extract a subrecord from another record (take a subset of the keys/names). The *args combinators I wrote in my blog post can be made to work on **kwargs by using named thinnings. This can all be discussed without the concept of renaming, I dunno if it is useful without renaming though.\n",
    "\n",
    "Named thinning + permutative renamings does seem like an independent factoring of partial injective renamings.Philip Zucker: Using records for the sin cos model feels weird but is fine. Instead of R^3, use {x : R, y: R, z : R}. Its still mathematical.\n",
    "\n",
    "\n",
    "My symbols were all (T -> R) -> (T -> R) lifted, but Var wasn't\n",
    "\n",
    "Id\n",
    "PComp(Const())\n",
    "\n",
    "Maybe PComp is fused with App. Or\n",
    "\n",
    "App == PComp\n",
    "def pcomp(f):\n",
    "    return lambda g0: lmabda g1: ... lambda *args:\n",
    "def app(f, g):\n",
    "    return lambda *gs: lambda *args: f(g(*args), *[g(*args) for g in gs])\n",
    "\n",
    "\n",
    "enum Node {\n",
    "    App(Node, Node),\n",
    "    Symbol( name : String, arity : usize)\n",
    "    PComp\n",
    "    //Id\n",
    "}\n",
    "\n",
    "Var = Symbol(\"Id\", arity=1)\n",
    "\n"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": null,
   "id": "64af74bf",
   "metadata": {},
   "outputs": [
    {
     "data": {
      "text/plain": [
       "-10"
      ]
     },
     "execution_count": 4,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "type NThin = dict[str, bool]\n",
    "\n",
    "def nthin(t : NThin, d : dict) -> dict:\n",
    "    return {k : v for k,v in d.items() if t[k]}\n",
    "\n",
    "def nlift(t : NThin, f):\n",
    "    return lambda **kwargs: f(**nthin(t, kwargs))\n",
    "\n",
    "# oh interesting. They have their own defaults... But then lift needs to carry the default? And dump is a no-op\n",
    "def foo(x=0, y=1):\n",
    "    return x - y\n",
    "\n",
    "nlift({\"x\": True, \"y\": True, \"z\" : False}, foo)(x=10, y=20, z=200)\n",
    "\n",
    "\n",
    "\n",
    "\n"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 13,
   "id": "6e3937ca",
   "metadata": {},
   "outputs": [
    {
     "data": {
      "text/plain": [
       "-97"
      ]
     },
     "execution_count": 13,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "def replace_thin(t : dict, d : dict) -> dict:\n",
    "    return {**d, **t}\n",
    "def rlift(t, f):\n",
    "    return lambda **kwargs: f(**replace_thin(t, kwargs))\n",
    "\n",
    "rlift({\"y\": 100}, foo)(x=3)"
   ]
  }
 ],
 "metadata": {
  "kernelspec": {
   "display_name": "philzook58.github.io (3.12.3)",
   "language": "python",
   "name": "python3"
  },
  "language_info": {
   "codemirror_mode": {
    "name": "ipython",
    "version": 3
   },
   "file_extension": ".py",
   "mimetype": "text/x-python",
   "name": "python",
   "nbconvert_exporter": "python",
   "pygments_lexer": "ipython3",
   "version": "3.12.3"
  }
 },
 "nbformat": 4,
 "nbformat_minor": 5
}
